// Numbas version: finer_feedback_settings
{"name": "Friction: Determine direction of $F_f$", "extensions": ["quantities"], "custom_part_types": [{"source": {"pk": 19, "author": {"name": "William Haynes", "pk": 2530}, "edit_page": "/part_type/19/edit"}, "name": "Engineering Accuracy with units", "short_name": "engineering-answer", "description": "<p>A value with units marked right if within an adjustable % error of the correct value. &nbsp;Marked close if within a wider margin of error.</p>", "help_url": "", "input_widget": "string", "input_options": {"correctAnswer": "siground(settings['correctAnswer'],4)", "hint": {"static": true, "value": ""}, "allowEmpty": {"static": true, "value": true}}, "can_be_gap": true, "can_be_step": true, "marking_script": "mark:\nswitch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)\n\ninterpreted_answer:\nqty(student_scalar, student_units)\n\n\n\ncorrect_quantity:\nsettings[\"correctAnswer\"]\n\n\n\ncorrect_units:\nunits(correct_quantity)\n\n\nallowed_notation_styles:\n[\"plain\",\"en\"]\n\nmatch_student_number:\nmatchnumber(studentAnswer,allowed_notation_styles)\n\nstudent_scalar:\nmatch_student_number[1]\n\nstudent_units:\nreplace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")\n\ngood_units:\ntry(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n\n\nstudent_quantity:\nswitch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n\n\npercent_error:\ntry(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   \n\nright:\npercent_error <= settings['right']\n\n\nclose:\nright_sign and percent_error <= settings['close']\n\nright_sign:\nsign(student_scalar) = sign(correct_quantity)", "marking_notes": [{"name": "mark", "description": "This is the main marking note. It should award credit and provide feedback based on the student's answer.", "definition": "switch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)"}, {"name": "interpreted_answer", "description": "A value representing the student's answer to this part.", "definition": "qty(student_scalar, student_units)\n\n"}, {"name": "correct_quantity", "description": "", "definition": "settings[\"correctAnswer\"]\n\n"}, {"name": "correct_units", "description": "", "definition": "units(correct_quantity)\n"}, {"name": "allowed_notation_styles", "description": "", "definition": "[\"plain\",\"en\"]"}, {"name": "match_student_number", "description": "", "definition": "matchnumber(studentAnswer,allowed_notation_styles)"}, {"name": "student_scalar", "description": "", "definition": "match_student_number[1]"}, {"name": "student_units", "description": "<p>Modify&nbsp;the&nbsp;unit portion of&nbsp;the student's answer by</p>\n<p>1. replacing \"ohms\" with \"ohm\" &nbsp;case insensitive</p>\n<p>2.&nbsp;replacing '-'&nbsp;with ' '&nbsp;</p>\n<p>3. replacing '&deg;' with ' deg'&nbsp;</p>\n<p>to allow answers like 10 ft-lb and 30&deg;</p>", "definition": "replace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")"}, {"name": "good_units", "description": "", "definition": "try(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n"}, {"name": "student_quantity", "description": "<p>This fixes the student answer for two common errors. &nbsp;</p>\n<p>If student_units are wrong&nbsp;&nbsp;- replace with correct units</p>\n<p>If student_scalar has the wrong sign - replace with right sign</p>\n<p>If student makes both errors, only one gets fixed.</p>", "definition": "switch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n"}, {"name": "percent_error", "description": "", "definition": "try(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   "}, {"name": "right", "description": "", "definition": "percent_error <= settings['right']\n"}, {"name": "close", "description": "<p>Only marked close if the student actually has the right sign.</p>", "definition": "right_sign and percent_error <= settings['close']"}, {"name": "right_sign", "description": "", "definition": "sign(student_scalar) = sign(correct_quantity) "}], "settings": [{"name": "correctAnswer", "label": "Correct Quantity.", "help_url": "", "hint": "The correct answer given as a JME quantity.", "input_type": "code", "default_value": "", "evaluate": true}, {"name": "right", "label": "% Accuracy for right.", "help_url": "", "hint": "Question will be considered correct if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "0.2", "evaluate": true}, {"name": "close", "label": "% Accuracy for close.", "help_url": "", "hint": "Question will be considered close if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "1.0", "evaluate": true}, {"name": "C1", "label": "Close with units.", "help_url": "", "hint": "Partial Credit for close value with appropriate units. &nbsp;if correct answer is 100 N and close is &plusmn;1%,<br />99 &nbsp;N is accepted.", "input_type": "percent", "default_value": "75"}, {"name": "C2", "label": "No units or wrong sign", "help_url": "", "hint": "Partial credit for forgetting units or using wrong sign.<br />If the correct answer is 100 N, both 100 and -100 N are accepted.", "input_type": "percent", "default_value": "50"}, {"name": "C3", "label": "Close, no units.", "help_url": "", "hint": "Partial Credit for close value but forgotten units.<br />This value would be close&nbsp;if the expected units were provided. &nbsp;If the correct answer is 100 N, and close is &plusmn;1%,<br />99 is accepted.", "input_type": "percent", "default_value": "25"}], "public_availability": "always", "published": true, "extensions": ["quantities"]}], "resources": [["question-resources/hw5sol1.png", "hw5sol1.png"], ["question-resources/block_TxlOZUk.svg", "block_TxlOZUk.svg"]], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Friction: Determine direction of $F_f$", "tags": [], "metadata": {"description": "<p>Friction</p>", "licence": "Creative Commons Attribution-NonCommercial 4.0 International"}, "statement": "<p>Determine whether the block shown is in equilibrium and find the magnitude and direction of the friction force when $\\mu_s = 0.20$, $\\mu_k = 0.15$, $\\theta=\\var{th}^\\circ$ and $P=\\var{P}\\,\\text{N}$</p>\n<p style=\"text-align: center;\"><img src=\"resources/question-resources/block_TxlOZUk.svg\" width=\"240\" height=\"240\"/></p>", "advice": "<p><strong><img src=\"resources/question-resources/hw5sol1.png\" width=\"202\" height=\"144\"/></strong></p>\n<p><strong>(a)</strong></p>\n<p>Assume equilibrium</p>\n<p>$+\\!\\!\\nearrow\\sum F_y=0,\\quad\\implies\\quad N-(800\\,\\text{N})\\cos 25^\\circ +(\\var{P}\\,\\text{N})\\sin(\\var{th}^\\circ-25^\\circ)=0$</p>\n<p style=\"padding-left: 30px;\">$N=\\var{normal_force}\\,\\text{N}$</p>\n<p>$+\\!\\!\\searrow\\sum F_x=0,\\quad\\implies\\quad -F_\\mu+(800\\,\\text{N})\\sin 25^\\circ-(\\var{P}\\,\\text{N})\\cos&nbsp;(\\var{th}^\\circ-25^\\circ)=0$</p>\n<p style=\"padding-left: 30px;\">$F_\\mu=\\var{friction_force_trial}\\,\\text{N}$ is required for equilibrium.</p>\n<p>Maximum friction force: $F_m=\\mu_s N=0.2\\times\\var{normal_force}=\\var{max_friction_force}\\,\\text{N}$</p>\n<div data-jme-visible=\"abs({friction_force_trial})&gt;{mu_s}*{normal_force}\">\n<p>Since $|F_\\mu| &gt; F_m$, the block is sliding,&nbsp;</p>\n<p><strong>(b)</strong></p>\n<p>and the magmitude of the friction force is</p>\n<p>\\[F_\\mu=\\mu_k N=\\var{abs(friction_force)}\\,\\text{N},\\quad\\var{direction_string}\\]&nbsp;</p>\n</div>\n<p></p>\n<div data-jme-visible=\"abs({friction_force_trial})&lt;={mu_s}*{normal_force}\">\n<p>Since $|F_\\mu| \\le F_m$, the block is in equilibrium, and</p>\n<p><strong>(b)</strong></p>\n<p>The magmitude of the friction force is</p>\n<p>\\[F_\\mu=\\var{abs(friction_force)}\\,\\text{N},\\quad\\var{direction_string}\\]&nbsp;</p>\n</div>", "rulesets": {}, "extensions": ["quantities"], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"mu_s": {"name": "mu_s", "group": "Ungrouped variables", "definition": "0.2", "description": "", "templateType": "number", "can_override": false}, "th": {"name": "th", "group": "Ungrouped variables", "definition": "random(25 .. 70#5)", "description": "", "templateType": "randrange", "can_override": false}, "mu_k": {"name": "mu_k", "group": "Ungrouped variables", "definition": "0.15", "description": "", "templateType": "anything", "can_override": false}, "friction_force_trial": {"name": "friction_force_trial", "group": "Ungrouped variables", "definition": "precround(800*sin(radians(25))-P*cos(radians(th-25)),1)", "description": "<p>Positive if up the slope</p>", "templateType": "anything", "can_override": false}, "equil": {"name": "equil", "group": "Ungrouped variables", "definition": "if(abs(friction_force_trial)<=(mu_s*normal_force),[10,0],[0,10])", "description": "", "templateType": "anything", "can_override": false}, "direction": {"name": "direction", "group": "Ungrouped variables", "definition": "if(friction_force<0,[0,10],[10,0])", "description": "", "templateType": "anything", "can_override": false}, "max_friction_force": {"name": "max_friction_force", "group": "Ungrouped variables", "definition": "precround(mu_s*normal_force,1)", "description": "", "templateType": "anything", "can_override": false}, "normal_force": {"name": "normal_force", "group": "Ungrouped variables", "definition": "precround(800*cos(radians(25))-P*sin(radians(th-25)),1)", "description": "", "templateType": "anything", "can_override": false}, "friction_force": {"name": "friction_force", "group": "Ungrouped variables", "definition": "if(abs(friction_force_trial)<=(mu_s*normal_force),friction_force_trial,sign(friction_force_trial)*precround(mu_k*normal_force,1))", "description": "", "templateType": "anything", "can_override": false}, "direction_string": {"name": "direction_string", "group": "Ungrouped variables", "definition": "if(friction_force<0,\"down the slope\",\"up the slope\")", "description": "", "templateType": "anything", "can_override": false}, "P": {"name": "P", "group": "Ungrouped variables", "definition": "random(100 .. 500#20)", "description": "", "templateType": "randrange", "can_override": false}, "W": {"name": "W", "group": "Ungrouped variables", "definition": "\"800 N\"", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["th", "P", "normal_force", "mu_s", "mu_k", "friction_force_trial", "max_friction_force", "equil", "friction_force", "direction", "direction_string", "W"], "variable_groups": [], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Is the block in static equilibrium?</p>\n<p style=\"padding-left: 40px;\">[[0]]</p>\n<p>The magnitude of the friction force is [[1]]</p>\n<p>The friction force acts &nbsp;[[2]].</p>", "gaps": [{"type": "1_n_2", "useCustomName": true, "customName": "In equilibrium", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "radiogroup", "displayColumns": "0", "showBlankOption": true, "showCellAnswerState": true, "choices": ["<p>Yes</p>", "<p>No</p>"], "matrix": "{equil}"}, {"type": "engineering-answer", "useCustomName": true, "customName": "Friction", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(abs(friction_force),'N')", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": true, "customName": "Direction", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "radiogroup", "displayColumns": "2", "showBlankOption": true, "showCellAnswerState": true, "choices": ["<p>UP the slope $\\nwarrow$</p>", "<p>DOWN the slope $\\searrow$</p>"], "matrix": "{direction}"}], "sortAnswers": false}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "Vladimir Vingoradov", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/1862/"}, {"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}, {"name": "Haoyu Huang", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/17657/"}], "resources": ["question-resources/hw5sol1.png", "question-resources/block_TxlOZUk.svg"]}]}], "contributors": [{"name": "Vladimir Vingoradov", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/1862/"}, {"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}, {"name": "Haoyu Huang", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/17657/"}]}