// Numbas version: finer_feedback_settings
{"name": "Frame: A-frame Difficulty 4", "extensions": ["geogebra", "quantities"], "custom_part_types": [{"source": {"pk": 19, "author": {"name": "William Haynes", "pk": 2530}, "edit_page": "/part_type/19/edit"}, "name": "Engineering Accuracy with units", "short_name": "engineering-answer", "description": "<p>A value with units marked right if within an adjustable % error of the correct value. &nbsp;Marked close if within a wider margin of error.</p>", "help_url": "", "input_widget": "string", "input_options": {"correctAnswer": "siground(settings['correctAnswer'],4)", "hint": {"static": true, "value": ""}, "allowEmpty": {"static": true, "value": true}}, "can_be_gap": true, "can_be_step": true, "marking_script": "mark:\nswitch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)\n\ninterpreted_answer:\nqty(student_scalar, student_units)\n\n\n\ncorrect_quantity:\nsettings[\"correctAnswer\"]\n\n\n\ncorrect_units:\nunits(correct_quantity)\n\n\nallowed_notation_styles:\n[\"plain\",\"en\"]\n\nmatch_student_number:\nmatchnumber(studentAnswer,allowed_notation_styles)\n\nstudent_scalar:\nmatch_student_number[1]\n\nstudent_units:\nreplace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")\n\ngood_units:\ntry(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n\n\nstudent_quantity:\nswitch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n\n\npercent_error:\ntry(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   \n\nright:\npercent_error <= settings['right']\n\n\nclose:\nright_sign and percent_error <= settings['close']\n\nright_sign:\nsign(student_scalar) = sign(correct_quantity)", "marking_notes": [{"name": "mark", "description": "This is the main marking note. It should award credit and provide feedback based on the student's answer.", "definition": "switch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)"}, {"name": "interpreted_answer", "description": "A value representing the student's answer to this part.", "definition": "qty(student_scalar, student_units)\n\n"}, {"name": "correct_quantity", "description": "", "definition": "settings[\"correctAnswer\"]\n\n"}, {"name": "correct_units", "description": "", "definition": "units(correct_quantity)\n"}, {"name": "allowed_notation_styles", "description": "", "definition": "[\"plain\",\"en\"]"}, {"name": "match_student_number", "description": "", "definition": "matchnumber(studentAnswer,allowed_notation_styles)"}, {"name": "student_scalar", "description": "", "definition": "match_student_number[1]"}, {"name": "student_units", "description": "<p>Modify&nbsp;the&nbsp;unit portion of&nbsp;the student's answer by</p>\n<p>1. replacing \"ohms\" with \"ohm\" &nbsp;case insensitive</p>\n<p>2.&nbsp;replacing '-'&nbsp;with ' '&nbsp;</p>\n<p>3. replacing '&deg;' with ' deg'&nbsp;</p>\n<p>to allow answers like 10 ft-lb and 30&deg;</p>", "definition": "replace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")"}, {"name": "good_units", "description": "", "definition": "try(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n"}, {"name": "student_quantity", "description": "<p>This fixes the student answer for two common errors. &nbsp;</p>\n<p>If student_units are wrong&nbsp;&nbsp;- replace with correct units</p>\n<p>If student_scalar has the wrong sign - replace with right sign</p>\n<p>If student makes both errors, only one gets fixed.</p>", "definition": "switch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n"}, {"name": "percent_error", "description": "", "definition": "try(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   "}, {"name": "right", "description": "", "definition": "percent_error <= settings['right']\n"}, {"name": "close", "description": "<p>Only marked close if the student actually has the right sign.</p>", "definition": "right_sign and percent_error <= settings['close']"}, {"name": "right_sign", "description": "", "definition": "sign(student_scalar) = sign(correct_quantity) "}], "settings": [{"name": "correctAnswer", "label": "Correct Quantity.", "help_url": "", "hint": "The correct answer given as a JME quantity.", "input_type": "code", "default_value": "", "evaluate": true}, {"name": "right", "label": "% Accuracy for right.", "help_url": "", "hint": "Question will be considered correct if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "0.2", "evaluate": true}, {"name": "close", "label": "% Accuracy for close.", "help_url": "", "hint": "Question will be considered close if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "1.0", "evaluate": true}, {"name": "C1", "label": "Close with units.", "help_url": "", "hint": "Partial Credit for close value with appropriate units. &nbsp;if correct answer is 100 N and close is &plusmn;1%,<br />99 &nbsp;N is accepted.", "input_type": "percent", "default_value": "75"}, {"name": "C2", "label": "No units or wrong sign", "help_url": "", "hint": "Partial credit for forgetting units or using wrong sign.<br />If the correct answer is 100 N, both 100 and -100 N are accepted.", "input_type": "percent", "default_value": "50"}, {"name": "C3", "label": "Close, no units.", "help_url": "", "hint": "Partial Credit for close value but forgotten units.<br />This value would be close&nbsp;if the expected units were provided. &nbsp;If the correct answer is 100 N, and close is &plusmn;1%,<br />99 is accepted.", "input_type": "percent", "default_value": "25"}], "public_availability": "always", "published": true, "extensions": ["quantities"]}], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Frame: A-frame Difficulty 4", "tags": [], "metadata": {"description": "<p>An A-frame structure supporting a&nbsp;force and&nbsp;a moment. The feet are at the different&nbsp;vertical positions so the solution will require simultaneous equations, unless you rotate the coordinate system.</p>", "licence": "Creative Commons Attribution-NonCommercial 4.0 International"}, "statement": "<p>The A-frame shown supports a {display(abs(F_d))} force $D$ at the midpoint of member $AB$ and a {display(abs(M_E)) } {if(me&gt;0,\"counter-\",\"\")}clockwise moment $M_E$ at $E$ as shown. Determine the <em>scalar</em> components of the reactions at pins $A$ and $C$. &nbsp;</p>\n<table style=\"margin-left: auto; margin-right: auto;\">\n<tbody>\n<tr>\n<td>\n<p>{geogebra_applet('qe7xhnse',[[\"B\",B],[\"C\",C],['units','\"'+units[1]+'\"']]+ggb_load&nbsp;)}</p>\n</td>\n</tr>\n</tbody>\n</table>", "advice": "<p><strong>Strategy:</strong><br/><strong></strong></p>\n<p><strong>&nbsp; &nbsp; &nbsp;</strong>&nbsp;&nbsp;Draw a free body diagram of the entire frame and each of its legs. &nbsp;Note that members<em>&nbsp;AB</em> or <em>BC</em>&nbsp;are not&nbsp;two-force bodies;&nbsp;therefore&nbsp;the&nbsp;direction of the forces at <em>A</em> and <em>C</em>&nbsp;don't align with the members&nbsp;and should be represented with <em>x</em> and <em>y</em> components. &nbsp;Since all of the free&nbsp;body diagrams have&nbsp;four unknowns, there is no simple&nbsp;way to find one unknown using just one equation. &nbsp;The general approach in this case is to write two equations which contain the same two unknowns, then solve them simultaneously. &nbsp;</p>\n<p>{geogebra_applet('p6rz9bqx',[[\"B\",B],[\"C\",C]] + ggb_load )}</p>\n<p><strong>Known:</strong></p>\n<p style=\"padding-left: 30px;\">$M_E = \\var{display(abs(M_E))}$ and $D = \\var{display(abs(F_D))}$</p>\n<p><strong>Determine necessary angles:</strong></p>\n<p style=\"padding-left: 30px;\">$\\begin{align}\\alpha &amp;= \\tan^{-1}\\left(\\frac{\\var{y1}}{\\var{x1}}\\right) = \\var{siground(alpha,4)}&deg; &nbsp;&amp;&nbsp;\\beta &amp;= \\tan^{-1}\\left(\\frac{\\var{y2}}{\\var{x2}}\\right) = \\var{siground(beta,4)}&deg;\\end{align}$</p>\n<p><strong>Use&nbsp;FBDs II and III to generate&nbsp;two independent equations&nbsp;containing $B_x$ and&nbsp;$B_y$:</strong></p>\n<div>\n<p style=\"padding-left: 30px;\"><strong></strong>$\\begin{align} \\textrm{II: }&nbsp;\\Sigma M_A &amp;= 0 &nbsp; &nbsp;&amp; \\textrm{III: } \\Sigma M_C &amp;= 0\\\\B_x (\\var{y1}) + B_y (\\var{x1})&nbsp;&amp;=&nbsp;&nbsp;\\simplify{{sign(fd)} D &nbsp; ({y1/2})}&amp;&nbsp;&nbsp; - B_x (\\var{y2}) + B_y (\\var{x2})&nbsp;&amp;= &nbsp;\\simplify{ -&nbsp;{sign(me)} M_E &nbsp;} \\\\ &nbsp;\\var{scalar(y1)} B_x + \\var{scalar(x1)}B_y&amp; = \\var{fd scalar( y1/2) }\\qquad (1) &amp; &nbsp;- \\var{scalar(y2)} B_x + \\var{Scalar(x2)} B_y &amp;= &nbsp;\\var{scalar(-M_E ) }\\qquad (2)\\end{align}$</p>\n<p><strong></strong><strong>Solve equations (1) and (2) simultaneously to find $B_x$ and $B_y$:</strong></p>\n<p style=\"padding-left: 30px;\">$\\qquad B_x = \\var{display(Bx)} &nbsp;\\qquad\\qquad\\qquad\\qquad\\qquad&nbsp; B_y&nbsp;= \\var{display(By)}$</p>\n<p><strong>From FBD III:</strong></p>\n<p style=\"padding-left: 30px;\">$\\begin{align} \\textrm{III: }\\Sigma F_x &amp;= 0 &nbsp;&amp; &nbsp;&nbsp;\\textrm{III: }\\Sigma F_y &amp;= 0&nbsp;\\\\&nbsp;&nbsp;C_x &amp;= B_x &amp; C_y &amp; = B_y \\\\&nbsp;&amp;= \\var{display(abs(Cx))} &nbsp;\\var{if(vecC[0]&lt;0,latex('\\\\leftarrow'),latex('\\\\rightarrow'))}&amp; &amp;=&nbsp;\\var{display(abs(Cy))}&nbsp;\\var{if(vecC[1]&lt;0,latex('\\\\downarrow'),latex('\\\\uparrow'))}&nbsp;\\end{align}$</p>\n<p><strong>From FBD II:</strong></p>\n</div>\n<div>\n<p style=\"padding-left: 30px;\">$\\begin{align} \\textrm{II: }\\Sigma F_x &amp;= 0 &amp; \\textrm{II: } \\Sigma F_y &amp;= 0\\\\ A_x &amp;= \\simplify{B_x - {sign(fd)} D } &nbsp;&amp; A_y &amp;= -B_y &nbsp;\\\\&amp;= \\simplify[!collectNumbers]{({display(-qty(vecC[0],units[0]))}) +&nbsp;({-F_D})}&amp; &amp;= &nbsp;- ( \\var{display(qty(vecC[1],units[0]))}) \\\\ &amp;=\\var{display(abs(qty(vecA[0],units[0])))} \\var{if(vecA[0]&lt;0,latex('\\\\leftarrow'),latex('\\\\rightarrow'))} &amp; &amp;=&nbsp;\\var{display(abs(qty(vecA[1],units[0])))} \\var{if(vecA[1]&lt;0,latex('\\\\downarrow'),latex('\\\\uparrow'))}\\end{align}$ &nbsp;</p>\n</div>", "rulesets": {}, "extensions": ["geogebra", "quantities"], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"vecW": {"name": "vecW", "group": "vectors", "definition": "vector(0,-w)", "description": "<p>always 0 in this problem</p>", "templateType": "anything", "can_override": false}, "Cy": {"name": "Cy", "group": "quantities", "definition": "By", "description": "<p>( sign(fd)} D &nbsp; y1/2 &nbsp;+ -sign(me) M_E ) /(x1+x2)</p>", "templateType": "anything", "can_override": false}, "fd": {"name": "fd", "group": "Ungrouped variables", "definition": "random(-500..500#25 except 0)", "description": "<p>the magnitude of the force</p>", "templateType": "anything", "can_override": false}, "me": {"name": "me", "group": "Ungrouped variables", "definition": "random(100..2000#100) random(1,-1)", "description": "<p>the magnitude of the moment</p>", "templateType": "anything", "can_override": false}, "x2": {"name": "x2", "group": "quantities", "definition": "qty(abs(b[0]-c[0]),units[1])", "description": "", "templateType": "anything", "can_override": false}, "F_w": {"name": "F_w", "group": "quantities", "definition": "qty(w,units[0])", "description": "", "templateType": "anything", "can_override": false}, "alpha": {"name": "alpha", "group": "quantities", "definition": "degrees(angle(vector(1,0),B))", "description": "", "templateType": "anything", "can_override": false}, "x1": {"name": "x1", "group": "quantities", "definition": "qty(B[0],units[1])", "description": "", "templateType": "anything", "can_override": false}, "F_d": {"name": "F_d", "group": "quantities", "definition": "qty(fd,units[0])", "description": "", "templateType": "anything", "can_override": false}, "Cx": {"name": "Cx", "group": "quantities", "definition": "Bx", "description": "<p>( C_y x2 &nbsp;+ -sign(me) M_E )/ y2</p>", "templateType": "anything", "can_override": false}, "det": {"name": "det", "group": "quantities", "definition": "det(matrix(\n    [scalar(y1),scalar(x1)],\n    [-scalar(y2),scalar(x2)]\n))", "description": "", "templateType": "anything", "can_override": false}, "dety": {"name": "dety", "group": "quantities", "definition": "det(matrix(\n    [scalar(y1)       ,fd scalar(y1) /2],\n    [-scalar(y2)      ,- me]\n))", "description": "", "templateType": "anything", "can_override": false}, "vecC": {"name": "vecC", "group": "vectors", "definition": "vector(scalar(-cx),scalar(cy))", "description": "<p>reactions at C&nbsp;for both versions as a vector</p>", "templateType": "anything", "can_override": false}, "y2": {"name": "y2", "group": "quantities", "definition": "qty(abs(b[1]-c[1]),units[1])", "description": "", "templateType": "anything", "can_override": false}, "M_E": {"name": "M_E", "group": "quantities", "definition": "qty(me,units[0]+\"*\"+units[1])", "description": "", "templateType": "anything", "can_override": false}, "sigma_f": {"name": "sigma_f", "group": "vectors", "definition": "vecW+vecA+vecC + vecD", "description": "", "templateType": "anything", "can_override": false}, "B": {"name": "B", "group": "Ungrouped variables", "definition": "vector(random(3..5),random(3..5))", "description": "<p>location of pin B</p>", "templateType": "anything", "can_override": false}, "vecD": {"name": "vecD", "group": "vectors", "definition": "vector(scalar(f_d),0)", "description": "", "templateType": "anything", "can_override": false}, "w": {"name": "w", "group": "Ungrouped variables", "definition": "0//random(100..500#25)", "description": "", "templateType": "anything", "can_override": false}, "y1": {"name": "y1", "group": "quantities", "definition": "qty(B[1],units[1])", "description": "", "templateType": "anything", "can_override": false}, "beta": {"name": "beta", "group": "quantities", "definition": "degrees(angle(vector(-1,0),B-C))", "description": "", "templateType": "anything", "can_override": false}, "units": {"name": "units", "group": "Ungrouped variables", "definition": "random(['N','m'],['lb','ft'])", "description": "", "templateType": "anything", "can_override": false}, "Bx": {"name": "Bx", "group": "quantities", "definition": "qty(detx/det,units[0])", "description": "", "templateType": "anything", "can_override": false}, "description": {"name": "description", "group": "quantities", "definition": "display(abs(F_d)) +  \" force $D$ at the midpoint of member <i>AB</i> and\"+\n  display(abs(M_E)) + if(me>0,\" counter-\",\" \") + \"clockwise moment $M_E$ \"", "description": "", "templateType": "anything", "can_override": false}, "detx": {"name": "detx", "group": "quantities", "definition": "det(matrix(\n    [fd scalar(y1)/2 ,scalar(x1)],\n    [- me            ,scalar(x2)]\n))", "description": "", "templateType": "anything", "can_override": false}, "By": {"name": "By", "group": "quantities", "definition": "qty(dety/det,units[0])", "description": "", "templateType": "anything", "can_override": false}, "vecA": {"name": "vecA", "group": "vectors", "definition": "-(vecD+VecC)", "description": "<p>reactions at A for both versions as a vector</p>", "templateType": "anything", "can_override": false}, "C": {"name": "C", "group": "Ungrouped variables", "definition": "vector(B[0]+ random(3..5),random(-4..4 except 0))", "description": "<p>Location of pin C. &nbsp;Pin A is at (0,0)</p>", "templateType": "anything", "can_override": false}, "ggb_load": {"name": "ggb_load", "group": "Ungrouped variables", "definition": "[['fd',fd],['me',me],['w',0]]", "description": "<p style=\"text-align: left;\">This will select either the force or the moment, but not both. &nbsp;Weight is always zero.</p>\n<p style=\"text-align: left;\"></p>\n<p style=\"text-align: left;\"></p>", "templateType": "anything", "can_override": false}, "debug": {"name": "debug", "group": "Ungrouped variables", "definition": "true", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "B[1] > C[1] ", "maxRuns": 100}, "ungrouped_variables": ["B", "C", "fd", "me", "w", "units", "ggb_load", "debug"], "variable_groups": [{"name": "quantities", "variables": ["F_w", "F_d", "M_E", "alpha", "beta", "x1", "y1", "x2", "y2", "description", "Cy", "Cx", "Bx", "By", "det", "detx", "dety"]}, {"name": "vectors", "variables": ["vecW", "vecA", "vecC", "vecD", "sigma_f"]}], "functions": {"display": {"parameters": [["Q", "quantity"]], "type": "string", "language": "jme", "definition": "string(siground(q,4))"}}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p style=\"padding-left: 40px;\">$A_x = $&nbsp;[[0]]&nbsp;</p>\n<p style=\"padding-left: 40px;\">$A_y = $&nbsp;[[1]]</p>\n<p style=\"padding-left: 40px;\">$C_x = $&nbsp;[[2]]&nbsp;</p>\n<p style=\"padding-left: 40px;\">$C_y = $&nbsp;[[3]]&nbsp;</p>", "gaps": [{"type": "engineering-answer", "useCustomName": true, "customName": "$A_x$", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(vecA[0],units[0])", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "$A_y$", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(vecA[1],units[0])", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "$C_x$", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(vecC[0],units[0])", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "$C_y$", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(vecC[1],units[0])", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}], "sortAnswers": false}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}], "resources": []}]}], "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}]}