// Numbas version: finer_feedback_settings
{"name": "Machine: Excavator", "extensions": ["geogebra", "quantities"], "custom_part_types": [{"source": {"pk": 19, "author": {"name": "William Haynes", "pk": 2530}, "edit_page": "/part_type/19/edit"}, "name": "Engineering Accuracy with units", "short_name": "engineering-answer", "description": "<p>A value with units marked right if within an adjustable % error of the correct value. &nbsp;Marked close if within a wider margin of error.</p>", "help_url": "", "input_widget": "string", "input_options": {"correctAnswer": "siground(settings['correctAnswer'],4)", "hint": {"static": true, "value": ""}, "allowEmpty": {"static": true, "value": true}}, "can_be_gap": true, "can_be_step": true, "marking_script": "mark:\nswitch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)\n\ninterpreted_answer:\nqty(student_scalar, student_units)\n\n\n\ncorrect_quantity:\nsettings[\"correctAnswer\"]\n\n\n\ncorrect_units:\nunits(correct_quantity)\n\n\nallowed_notation_styles:\n[\"plain\",\"en\"]\n\nmatch_student_number:\nmatchnumber(studentAnswer,allowed_notation_styles)\n\nstudent_scalar:\nmatch_student_number[1]\n\nstudent_units:\nreplace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")\n\ngood_units:\ntry(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n\n\nstudent_quantity:\nswitch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n\n\npercent_error:\ntry(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   \n\nright:\npercent_error <= settings['right']\n\n\nclose:\nright_sign and percent_error <= settings['close']\n\nright_sign:\nsign(student_scalar) = sign(correct_quantity)", "marking_notes": [{"name": "mark", "description": "This is the main marking note. It should award credit and provide feedback based on the student's answer.", "definition": "switch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)"}, {"name": "interpreted_answer", "description": "A value representing the student's answer to this part.", "definition": "qty(student_scalar, student_units)\n\n"}, {"name": "correct_quantity", "description": "", "definition": "settings[\"correctAnswer\"]\n\n"}, {"name": "correct_units", "description": "", "definition": "units(correct_quantity)\n"}, {"name": "allowed_notation_styles", "description": "", "definition": "[\"plain\",\"en\"]"}, {"name": "match_student_number", "description": "", "definition": "matchnumber(studentAnswer,allowed_notation_styles)"}, {"name": "student_scalar", "description": "", "definition": "match_student_number[1]"}, {"name": "student_units", "description": "<p>Modify&nbsp;the&nbsp;unit portion of&nbsp;the student's answer by</p>\n<p>1. replacing \"ohms\" with \"ohm\" &nbsp;case insensitive</p>\n<p>2.&nbsp;replacing '-'&nbsp;with ' '&nbsp;</p>\n<p>3. replacing '&deg;' with ' deg'&nbsp;</p>\n<p>to allow answers like 10 ft-lb and 30&deg;</p>", "definition": "replace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")"}, {"name": "good_units", "description": "", "definition": "try(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n"}, {"name": "student_quantity", "description": "<p>This fixes the student answer for two common errors. &nbsp;</p>\n<p>If student_units are wrong&nbsp;&nbsp;- replace with correct units</p>\n<p>If student_scalar has the wrong sign - replace with right sign</p>\n<p>If student makes both errors, only one gets fixed.</p>", "definition": "switch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n"}, {"name": "percent_error", "description": "", "definition": "try(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   "}, {"name": "right", "description": "", "definition": "percent_error <= settings['right']\n"}, {"name": "close", "description": "<p>Only marked close if the student actually has the right sign.</p>", "definition": "right_sign and percent_error <= settings['close']"}, {"name": "right_sign", "description": "", "definition": "sign(student_scalar) = sign(correct_quantity) "}], "settings": [{"name": "correctAnswer", "label": "Correct Quantity.", "help_url": "", "hint": "The correct answer given as a JME quantity.", "input_type": "code", "default_value": "", "evaluate": true}, {"name": "right", "label": "% Accuracy for right.", "help_url": "", "hint": "Question will be considered correct if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "0.2", "evaluate": true}, {"name": "close", "label": "% Accuracy for close.", "help_url": "", "hint": "Question will be considered close if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "1.0", "evaluate": true}, {"name": "C1", "label": "Close with units.", "help_url": "", "hint": "Partial Credit for close value with appropriate units. &nbsp;if correct answer is 100 N and close is &plusmn;1%,<br />99 &nbsp;N is accepted.", "input_type": "percent", "default_value": "75"}, {"name": "C2", "label": "No units or wrong sign", "help_url": "", "hint": "Partial credit for forgetting units or using wrong sign.<br />If the correct answer is 100 N, both 100 and -100 N are accepted.", "input_type": "percent", "default_value": "50"}, {"name": "C3", "label": "Close, no units.", "help_url": "", "hint": "Partial Credit for close value but forgotten units.<br />This value would be close&nbsp;if the expected units were provided. &nbsp;If the correct answer is 100 N, and close is &plusmn;1%,<br />99 is accepted.", "input_type": "percent", "default_value": "25"}], "public_availability": "always", "published": true, "extensions": ["quantities"]}], "resources": [["question-resources/shearFBD.png", "shearFBD.png"], ["question-resources/shearFBD.gif", "shearFBD.gif"], ["question-resources/shear.gif", "shear.gif"], ["question-resources/shearFBD_XWaIQSM.png", "shearFBD_XWaIQSM.png"], ["question-resources/FBD.png", "FBD.png"], ["question-resources/FBD_UvxNguf.png", "FBD_UvxNguf.png"], ["question-resources/FBD_ZA7RKPO.png", "FBD_ZA7RKPO.png"]], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Machine: Excavator", "tags": ["Equilibrium", "equilibrium", "machine", "Mechanics", "mechanics", "rigid body", "Rigid Body", "Statics", "statics"], "metadata": {"description": "<p>Determine the&nbsp;forces in&nbsp;the pistons of a hydraulicly operated excavator.</p>", "licence": "Creative Commons Attribution-NonCommercial 4.0 International"}, "statement": "<p>The motion of the excavator is controlled by three hydraulic cylinders. &nbsp;Determine the force exerted by cylinders $BC$ and $DE$ when supporting a {load} downward vertical load at $G$. &nbsp;Neglect the weights of the parts.</p>\n<table style=\"margin-left: auto; margin-right: auto;\">\n<tbody>\n<tr>\n<td>{geogebra_applet('cbh4g7zm', {map(dx(n),n,1..6)} + {map(dy(n),n,1..6)} )}</td>\n</tr>\n</tbody>\n</table>\n<p>Diagram units are in:&nbsp;&nbsp;[{units[1]}]</p>\n<p></p>", "advice": "<p><strong>Draw and number&nbsp;neat, labeled&nbsp;free body diagrams.</strong></p>\n<p><img src=\"resources/question-resources/FBD_ZA7RKPO.png\" width=\"655\" height=\"317\"/></p>\n<p style=\"padding-left: 30px;\">Given: $W$ = {load}</p>\n<p><strong>Find the necessarys angles.</strong></p>\n<p style=\"padding-left: 30px;\">Let $\\theta_B$ &nbsp;and $\\theta_D$ be the angles that <em>BC</em>&nbsp; and <em>DE</em> make with the horizontal.</p>\n<p style=\"padding-left: 30px;\">$\\theta_B = \\tan^{-1} \\left(\\dfrac{\\var{y1+y2}}{\\var{x2}}\\right) = \\var{siground(theta_B,4)}&deg; \\qquad&nbsp;\\theta_D&nbsp;= \\tan^{-1} \\left(\\dfrac{\\var{y4}}{\\var{x4}} \\right)= \\var{siground(theta_D,4)}&deg;$</p>\n<p><strong>Use Free Body Diagram I&nbsp;to solve for the force in cylinder $BC$.</strong></p>\n<p style=\"padding-left: 30px;\">$\\begin{align} \\textrm{I: } \\Sigma M_A &amp;= 0 \\\\ &nbsp;BC_x ( \\var{y2}) -&nbsp;BC_y(\\var{x1 + x2}) + W ( \\var{x_t}) &amp;= 0\\\\ BC ( -\\var{scalar(y2)}&nbsp;\\cos \\theta_B +&nbsp;\\var{scalar(x1+x2)}&nbsp;\\sin \\theta_B) &amp;= &nbsp;\\var{scalar(x_t)}\\ W \\\\ BC &amp;= W \\left(\\dfrac{\\var{scalar(x_t)}}{-\\var{scalar(y2)}&nbsp;\\cos \\var{siground(theta_B,4)}&deg; +&nbsp;\\var{scalar(x1+x2)}&nbsp;\\sin \\var{siground(theta_B,4)}&deg;}\\right) \\\\&nbsp; &amp;= \\var{load} \\left(\\dfrac{\\var{x_t}}{\\var{siground(dperp_B,4)}}\\right) &nbsp;\\\\ &amp;= \\var{siground(BC,4)}&nbsp;\\textrm{&nbsp;(compression)}&nbsp;\\end{align}$&nbsp;</p>\n<p></p>\n<p><strong>Use Free Body Diagram II to solve for the force in cylinder $DE$.</strong></p>\n<p style=\"padding-left: 30px;\">$\\begin{align} \\textrm{II: } \\Sigma M_F &amp;= 0 \\\\ &nbsp;-DE_x ( \\var{y5}) - DE_y(\\var{x5}) + W ( \\var{x6}) &amp;= 0\\\\ DE&nbsp;( \\var{scalar(y5)} \\cos \\theta_D&nbsp;+&nbsp;\\var{scalar(x5)} \\sin \\theta_d) &amp;= &nbsp;\\var{scalar(x6)} W \\\\ DE&nbsp;&amp;= &nbsp;W \\left(\\dfrac{\\var{x6}}{\\var{siground(dperp_D,4)}}\\right) &nbsp;\\\\ &amp;= \\var{siground(DE,4)} \\textrm{ (tension)} \\end{align}$&nbsp;</p>", "rulesets": {}, "extensions": ["geogebra", "quantities"], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"y3": {"name": "y3", "group": "inputs", "definition": "round((1.2  L random(0.9..1.1#0.05)),units[2])", "description": "", "templateType": "anything", "can_override": false}, "x3": {"name": "x3", "group": "inputs", "definition": "round((2.6 L random(0.9..1.1#0.05)),units[2])", "description": "", "templateType": "anything", "can_override": false}, "BC": {"name": "BC", "group": "inputs", "definition": "load x_t /dperp_B", "description": "", "templateType": "anything", "can_override": false}, "x2": {"name": "x2", "group": "inputs", "definition": "round((3.9 L random(0.9..1.1#0.05)),units[2])", "description": "", "templateType": "anything", "can_override": false}, "units": {"name": "units", "group": "inputs", "definition": "random(['kN','m', \"5 cm\"],['lbf','ft', \"0.05 ft\"])", "description": "<p>units[2] large rounding base</p>\n<p>units[3] small rounding base</p>", "templateType": "anything", "can_override": false}, "DE": {"name": "DE", "group": "inputs", "definition": "load x6 /dperp_D", "description": "", "templateType": "anything", "can_override": false}, "dperp_D": {"name": "dperp_D", "group": "inputs", "definition": "y5 cos(radians(theta_D)) + x5 sin(radians(theta_d))", "description": "", "templateType": "anything", "can_override": false}, "y6": {"name": "y6", "group": "inputs", "definition": "round((9.4  L random(0.9..1.1#0.05)),units[2])", "description": "", "templateType": "anything", "can_override": false}, "L": {"name": "L", "group": "inputs", "definition": "round(qty(random(12..24),'in') in units[1],units[2])", "description": "<p>scaling value for all lengths, which are also randomly perterbed actual value as scaled.</p>", "templateType": "anything", "can_override": false}, "dperp_B": {"name": "dperp_B", "group": "inputs", "definition": "abs(y2 cos(radians(theta_B)) - (x1+x2) sin(radians(theta_B)))", "description": "<p>perpendicular distance from A to force BC</p>", "templateType": "anything", "can_override": false}, "theta_d": {"name": "theta_d", "group": "outputs", "definition": "degrees(arctan(scalar(y4/x4)))", "description": "", "templateType": "anything", "can_override": false}, "y4": {"name": "y4", "group": "inputs", "definition": "round((1.0  L random(0.9..1.1#0.05)),units[2])", "description": "", "templateType": "anything", "can_override": false}, "y1": {"name": "y1", "group": "inputs", "definition": "round((1.2  L random(0.9..1.1#0.05)),units[2])", "description": "", "templateType": "anything", "can_override": false}, "y2": {"name": "y2", "group": "inputs", "definition": "round((4.2  L random(0.9..1.1#0.05)),units[2])", "description": "", "templateType": "anything", "can_override": false}, "x1": {"name": "x1", "group": "inputs", "definition": "round((1.7  L random(0.9..1.1#0.05)),units[2])", "description": "", "templateType": "anything", "can_override": false}, "x_t": {"name": "x_t", "group": "inputs", "definition": "x1+x2+x3+x4+x5+x6", "description": "", "templateType": "anything", "can_override": false}, "y5": {"name": "y5", "group": "inputs", "definition": "round((2.3 L random(0.9..1.1#0.05)),units[2])", "description": "", "templateType": "anything", "can_override": false}, "x5": {"name": "x5", "group": "inputs", "definition": "round((0.8 L random(0.9..1.1#0.05)),units[2])", "description": "", "templateType": "anything", "can_override": false}, "load": {"name": "load", "group": "inputs", "definition": "siground(qty(random(3000..5000),'lbf') in units[0],2)", "description": "", "templateType": "anything", "can_override": false}, "theta_b": {"name": "theta_b", "group": "outputs", "definition": "degrees(arctan(scalar((y1+y2)/x2)))", "description": "", "templateType": "anything", "can_override": false}, "x6": {"name": "x6", "group": "inputs", "definition": "round((5.2  L random(0.9..1.1#0.05)),units[2])", "description": "", "templateType": "anything", "can_override": false}, "x4": {"name": "x4", "group": "inputs", "definition": "round((7.6 L random(0.9..1.1#0.05)),units[2])", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": [], "variable_groups": [{"name": "Unnamed group", "variables": []}, {"name": "inputs", "variables": ["L", "x1", "x2", "x3", "x4", "x5", "x6", "y1", "y2", "y3", "y4", "y5", "y6", "load", "units", "x_t", "dperp_B", "BC", "dperp_D", "DE"]}, {"name": "outputs", "variables": ["theta_b", "theta_d"]}], "functions": {"dy": {"parameters": [["ix", "number"]], "type": "number", "language": "jme", "definition": "let(n,ix,['y' + n,'\\\"'+ scalar(precround(eval(expression('y' + n)),2))\n    + '\\\"'])"}, "dx": {"parameters": [["ix", "number"]], "type": "number", "language": "jme", "definition": "let(n,ix,['x' + n,'\\\"'+ scalar(precround(eval(expression('x' + n)),2))+ '\\\"'])"}}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p style=\"padding-left: 40px;\">$BC$ =&nbsp;[[0]]&nbsp;[[1]]&nbsp;&nbsp;</p>\n<p style=\"padding-left: 40px;\">$DE$ =&nbsp;[[2]]&nbsp;[[3]]&nbsp;</p>", "gaps": [{"type": "engineering-answer", "useCustomName": true, "customName": "BC", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "{BC}", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": true, "customName": "T or C", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "choices": ["Tension", "Compression"], "matrix": [0, "1"], "distractors": ["", ""]}, {"type": "engineering-answer", "useCustomName": true, "customName": "DE", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "{DE}", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": true, "customName": "T or C", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "choices": ["Tension", "Compression"], "matrix": ["1", "0"], "distractors": ["", ""]}], "sortAnswers": false}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}], "resources": ["question-resources/shearFBD.png", "question-resources/shearFBD.gif", "question-resources/shear.gif", "question-resources/shearFBD_XWaIQSM.png", "question-resources/FBD.png", "question-resources/FBD_UvxNguf.png", "question-resources/FBD_ZA7RKPO.png"]}]}], "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}]}