// Numbas version: finer_feedback_settings
{"name": "Internal Force: distributed load", "extensions": ["geogebra", "quantities"], "custom_part_types": [{"source": {"pk": 19, "author": {"name": "William Haynes", "pk": 2530}, "edit_page": "/part_type/19/edit"}, "name": "Engineering Accuracy with units", "short_name": "engineering-answer", "description": "<p>A value with units marked right if within an adjustable % error of the correct value. &nbsp;Marked close if within a wider margin of error.</p>", "help_url": "", "input_widget": "string", "input_options": {"correctAnswer": "siground(settings['correctAnswer'],4)", "hint": {"static": true, "value": ""}, "allowEmpty": {"static": true, "value": true}}, "can_be_gap": true, "can_be_step": true, "marking_script": "mark:\nswitch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)\n\ninterpreted_answer:\nqty(student_scalar, student_units)\n\n\n\ncorrect_quantity:\nsettings[\"correctAnswer\"]\n\n\n\ncorrect_units:\nunits(correct_quantity)\n\n\nallowed_notation_styles:\n[\"plain\",\"en\"]\n\nmatch_student_number:\nmatchnumber(studentAnswer,allowed_notation_styles)\n\nstudent_scalar:\nmatch_student_number[1]\n\nstudent_units:\nreplace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")\n\ngood_units:\ntry(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n\n\nstudent_quantity:\nswitch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n\n\npercent_error:\ntry(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   \n\nright:\npercent_error <= settings['right']\n\n\nclose:\nright_sign and percent_error <= settings['close']\n\nright_sign:\nsign(student_scalar) = sign(correct_quantity)", "marking_notes": [{"name": "mark", "description": "This is the main marking note. It should award credit and provide feedback based on the student's answer.", "definition": "switch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)"}, {"name": "interpreted_answer", "description": "A value representing the student's answer to this part.", "definition": "qty(student_scalar, student_units)\n\n"}, {"name": "correct_quantity", "description": "", "definition": "settings[\"correctAnswer\"]\n\n"}, {"name": "correct_units", "description": "", "definition": "units(correct_quantity)\n"}, {"name": "allowed_notation_styles", "description": "", "definition": "[\"plain\",\"en\"]"}, {"name": "match_student_number", "description": "", "definition": "matchnumber(studentAnswer,allowed_notation_styles)"}, {"name": "student_scalar", "description": "", "definition": "match_student_number[1]"}, {"name": "student_units", "description": "<p>Modify&nbsp;the&nbsp;unit portion of&nbsp;the student's answer by</p>\n<p>1. replacing \"ohms\" with \"ohm\" &nbsp;case insensitive</p>\n<p>2.&nbsp;replacing '-'&nbsp;with ' '&nbsp;</p>\n<p>3. replacing '&deg;' with ' deg'&nbsp;</p>\n<p>to allow answers like 10 ft-lb and 30&deg;</p>", "definition": "replace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")"}, {"name": "good_units", "description": "", "definition": "try(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n"}, {"name": "student_quantity", "description": "<p>This fixes the student answer for two common errors. &nbsp;</p>\n<p>If student_units are wrong&nbsp;&nbsp;- replace with correct units</p>\n<p>If student_scalar has the wrong sign - replace with right sign</p>\n<p>If student makes both errors, only one gets fixed.</p>", "definition": "switch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n"}, {"name": "percent_error", "description": "", "definition": "try(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   "}, {"name": "right", "description": "", "definition": "percent_error <= settings['right']\n"}, {"name": "close", "description": "<p>Only marked close if the student actually has the right sign.</p>", "definition": "right_sign and percent_error <= settings['close']"}, {"name": "right_sign", "description": "", "definition": "sign(student_scalar) = sign(correct_quantity) "}], "settings": [{"name": "correctAnswer", "label": "Correct Quantity.", "help_url": "", "hint": "The correct answer given as a JME quantity.", "input_type": "code", "default_value": "", "evaluate": true}, {"name": "right", "label": "% Accuracy for right.", "help_url": "", "hint": "Question will be considered correct if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "0.2", "evaluate": true}, {"name": "close", "label": "% Accuracy for close.", "help_url": "", "hint": "Question will be considered close if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "1.0", "evaluate": true}, {"name": "C1", "label": "Close with units.", "help_url": "", "hint": "Partial Credit for close value with appropriate units. &nbsp;if correct answer is 100 N and close is &plusmn;1%,<br />99 &nbsp;N is accepted.", "input_type": "percent", "default_value": "75"}, {"name": "C2", "label": "No units or wrong sign", "help_url": "", "hint": "Partial credit for forgetting units or using wrong sign.<br />If the correct answer is 100 N, both 100 and -100 N are accepted.", "input_type": "percent", "default_value": "50"}, {"name": "C3", "label": "Close, no units.", "help_url": "", "hint": "Partial Credit for close value but forgotten units.<br />This value would be close&nbsp;if the expected units were provided. &nbsp;If the correct answer is 100 N, and close is &plusmn;1%,<br />99 is accepted.", "input_type": "percent", "default_value": "25"}], "public_availability": "always", "published": true, "extensions": ["quantities"]}], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Internal Force: distributed load", "tags": [], "metadata": {"description": "<p>Find the shear, normal force and bending moment at a point in a beam supporting a uniformly distributed load.</p>", "licence": "Creative Commons Attribution-NonCommercial 4.0 International"}, "statement": "<p>The frame shown supports a uniformly distributed load $w = \\var{load}$. &nbsp;Member $BC$ is {BC} long. Determine the internal forces (shear, normal force, and bending moment) at point $D$, located {BD} to the right of point $B$.</p>\n<p>{geogebra_applet( 'pt6wrf9n', params)}</p>", "advice": "<p><strong>Find the reactions <em>B</em> and&nbsp;<em>C</em>.&nbsp;</strong></p>\n<p style=\"padding-left: 40px;\">Draw a Free body diagram of beam $BC$. &nbsp;</p>\n<p>{fbd1}</p>\n<p style=\"padding-left: 30px;\">Note that member&nbsp;$AB$&nbsp;is a two-force member, so force $AB$ acts along its axis.</p>\n<p style=\"padding-left: 30px;\">Replace the distributed load with an equivalent concentrated. The equivalent load acts at the center of member $BC$ with magnitude:</p>\n<p style=\"padding-left: 30px;\">$W = w \\cdot \\overline{BC} = (\\var{load}) (\\var{BC}) = \\var{Q_W}$.</p>\n<p></p>\n<p style=\"padding-left: 30px;\">Let the length of member &nbsp;$BC = \\ell = \\var{BC}$, then solve the equilibrium equations to find the reactions at $B$ and $C$.</p>\n<p style=\"padding-left: 30px;\">$\\begin{align}\\Sigma M_C &amp; = 0 &amp;\\Sigma F_x&nbsp;&amp;= 0 &amp;\\Sigma F_y &amp;= 0\\\\AB_y &nbsp;\\cdot \\ell&nbsp;&amp;= W \\cdot \\dfrac {\\ell}{2} &nbsp;&amp; &nbsp; C_x &amp;=AB_x &nbsp;&amp; C_y &amp;= W -AB_y&nbsp;\\\\AB \\sin \\var{abs(alpha)}&deg;&amp; = \\dfrac{ W }{2}&amp; C_x &amp;= AB \\cos \\var{abs(alpha)}&deg; &amp;&nbsp;C_y &amp;= W - AB&nbsp;\\sin \\var{abs(alpha)}&deg; \\\\ AB &amp;= \\var{display(Q_AB)}&nbsp;&amp; C_x &amp;= \\var{display(abs(Q_cx))}\\var{if(F_AB[0]&gt;0,latex('\\\\leftarrow'),latex('\\\\rightarrow'))} &amp; C_y &amp;= \\var{display(Q_Cy)} \\var{if(F_AB[1]&gt;0,latex('\\\\uparrow'),latex('\\\\downarrow'))}\\end{align}$</p>\n<p style=\"padding-left: 30px;\"></p>\n<p><strong>Take an imaginary cut through member <em>BC</em> at point&nbsp;<em>D</em> and draw a free body diagram of the right (or left) portion. </strong></p>\n<p>{fbd2}&nbsp;</p>\n<p style=\"padding-left: 40px;\">The length of segment $DC$ is $\\var{CD}$. There will be unknown internal shear ($V$), bending moment ($M_D$) and normal force&nbsp;($P$)&nbsp;at the cut. Assume that they act&nbsp;in the 'positive' directions established by the sign convention for shear and bending moments. &nbsp;</p>\n<p style=\"padding-left: 40px;\">Note that the downward force on this section of the beam&nbsp;($W'$) is&nbsp;due to the distributed force acting on this portion of the beam only and acts at the centroid of the loading.</p>\n<p style=\"padding-left: 40px;\">$W' = w \\cdot \\overline{DC} = (\\var{load}) (\\var{BC-BD}) = \\var{display(Q_W')}$</p>\n<p><strong>Solve for the internal forces at <em>D</em>.</strong></p>\n<p style=\"padding-left: 30px;\">$\\begin{align}\\Sigma M_D &amp; = 0 &amp;\\Sigma F_x &amp;= 0 &amp;\\Sigma F_y &amp;= 0\\\\ M_D&nbsp; &amp;= C_y (\\var{CD}) - W'(\\var{CD/2})&amp; &nbsp; P &amp;= \\var{if(F_AB[0]&gt;0,latex('-'),'')} C_x &nbsp;&amp; V&nbsp;&amp;= W'-C_y\\\\ &amp;= \\var{display(Q_MD)} &amp; &amp;= \\var{display(-Q_Cx)} &amp; &amp;= \\var{display(Q_V)}&nbsp;\\end{align}$</p>\n<p>&nbsp;&nbsp;</p>", "rulesets": {}, "extensions": ["geogebra", "quantities"], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"BC": {"name": "BC", "group": "Inputs", "definition": "qty(random(4..12#2),units[1])", "description": "", "templateType": "anything", "can_override": false}, "units": {"name": "units", "group": "Ungrouped variables", "definition": "['lb','ft']", "description": "", "templateType": "anything", "can_override": false}, "load": {"name": "load", "group": "Inputs", "definition": "qty(random(10..50#5), units[0\n  ]+\"/\"+units[1])", "description": "", "templateType": "anything", "can_override": false}, "Q_W'": {"name": "Q_W'", "group": "soon", "definition": "CD Load\n", "description": "", "templateType": "anything", "can_override": false}, "Q_V": {"name": "Q_V", "group": "soon", "definition": "Q_W'-Q_Cy", "description": "", "templateType": "anything", "can_override": false}, "Q_Cy": {"name": "Q_Cy", "group": "soon", "definition": "Q_W - qty(F_AB[1], units[0])", "description": "", "templateType": "anything", "can_override": false}, "Q_AB": {"name": "Q_AB", "group": "soon", "definition": "qty(abs(F_AB), units[0])", "description": "", "templateType": "anything", "can_override": false}, "params": {"name": "params", "group": "Ungrouped variables", "definition": "[alpha: [definition: radians(alpha), visible: false],\nn: [definition: n, visible: false],\nshow: [definition: 0, visible: false],\n]", "description": "<p>n = 1..10 &nbsp;</p>", "templateType": "anything", "can_override": false}, "F_AB": {"name": "F_AB", "group": "soon", "definition": "precround(scalar(Q_W)/2 * vector(1/tan(radians(alpha+180)), 1),4)", "description": "<p>alpha is defined from negative x axis,&nbsp;</p>", "templateType": "anything", "can_override": false}, "BD": {"name": "BD", "group": "Inputs", "definition": "n/10 BC\n", "description": "<p>n/10 is fraction of length BC</p>", "templateType": "anything", "can_override": false}, "Q_W": {"name": "Q_W", "group": "soon", "definition": "load * bc", "description": "<p>W</p>", "templateType": "anything", "can_override": false}, "alpha": {"name": "alpha", "group": "Inputs", "definition": "random(20..60#5) random(-1,1)", "description": "", "templateType": "anything", "can_override": false}, "sum": {"name": "sum", "group": "soon", "definition": "precround(F_AB+F_C+F_W,3)", "description": "<p>check</p>", "templateType": "anything", "can_override": false}, "Q_Cx": {"name": "Q_Cx", "group": "soon", "definition": "qty((F_AB[0]), units[0])", "description": "", "templateType": "anything", "can_override": false}, "F_W": {"name": "F_W", "group": "soon", "definition": "vector(-scalar(Q_W),0)", "description": "", "templateType": "anything", "can_override": false}, "CD": {"name": "CD", "group": "Inputs", "definition": "BC-BD", "description": "", "templateType": "anything", "can_override": false}, "Q_MD": {"name": "Q_MD", "group": "soon", "definition": "(CD Q_Cy - CD Q_W'/2)", "description": "<p>Q</p>", "templateType": "anything", "can_override": false}, "F_C": {"name": "F_C", "group": "soon", "definition": "precround(-(F_AB+F_W),4)", "description": "", "templateType": "anything", "can_override": false}, "applet": {"name": "applet", "group": "Ungrouped variables", "definition": "geogebra_applet( 'pt6wrf9n', params)", "description": "", "templateType": "anything", "can_override": false}, "n": {"name": "n", "group": "Inputs", "definition": "random(2..8# 0.5)", "description": "<p>Position of cut as ratio of BC</p>", "templateType": "anything", "can_override": false}, "fbd1": {"name": "fbd1", "group": "Ungrouped variables", "definition": "geogebra_applet('pt6wrf9n', params + [show: [definition: 1, visible: false]])", "description": "", "templateType": "anything", "can_override": false}, "fbd2": {"name": "fbd2", "group": "Ungrouped variables", "definition": "geogebra_applet('pt6wrf9n', params + \n[show: [definition: 2, visible: false],\n n: [definition: min(5,n), visible: false]\n])", "description": "<p>just setting n=5 to make the diagram larger</p>", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["units", "applet", "params", "fbd1", "fbd2"], "variable_groups": [{"name": "soon", "variables": ["Q_W", "F_AB", "F_W", "F_C", "sum", "Q_AB", "Q_Cx", "Q_Cy", "Q_MD", "Q_W'", "Q_V"]}, {"name": "Inputs", "variables": ["BC", "alpha", "load", "n", "BD", "CD"]}], "functions": {"display": {"parameters": [["q", "quantity"]], "type": "string", "language": "jme", "definition": "string(siground(q,4))"}}, "preamble": {"js": "question.signals.on('adviceDisplayed',function() {\n try{\n    var app = question.scope.variables.applet.app;  \n    app.setVisible(\"fbd\", true,false);\n    app.setValue(\"fbd\", 1);\n  }\n  catch(err){}  \n})\n", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p style=\"padding-left: 40px;\">$V_D &nbsp;= $ [[1]]</p>\n<p style=\"padding-left: 40px;\">$P_D = $ [[0]]</p>\n<p style=\"padding-left: 40px;\">$M_D =$ [[2]]</p>", "gaps": [{"type": "engineering-answer", "useCustomName": true, "customName": "$P_D$", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "-Q_Cx", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "$V_D$", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "Q_V", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "$M_D$", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "Q_MD", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}], "sortAnswers": false}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}], "resources": []}]}], "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}]}