// Numbas version: exam_question_groups
{"name": "Machine: lopping shears", "extensions": ["/srv/numbas/media/user-extensions/extracted/41/geogebra", "/srv/numbas/media/user-extensions/extracted/55/quantities", "/srv/numbas/media/user-extensions/extracted/40/visjs", "/srv/numbas/media/user-extensions/extracted/98/weh"], "custom_part_types": [{"help_url": "", "name": "Engineering Accuracy with units", "input_options": {"correctAnswer": "siground(settings['correctAnswer'],4)", "hint": {"value": "", "static": true}, "allowEmpty": {"value": true, "static": true}}, "marking_script": "mark:\nswitch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)\n\n\ninterpreted_answer:\nqty(student_scalar, student_units)\n\n\n\ncorrect_quantity:\nsettings[\"correctAnswer\"]\n\n\n\ncorrect_units:\nunits(correct_quantity)\n\n\nallowed_notation_styles:\n[\"plain\",\"en\"]\n\nmatch_student_number:\nmatchnumber(studentAnswer,allowed_notation_styles)\n\nstudent_scalar:\nmatch_student_number[1]\n\nstudent_units:\njoin(\nsplit(\njoin(\nsplit(studentAnswer[len(match_student_number[0])..len(studentAnswer)]\n,'-'),' ')\n,'\u00b0'),' deg')\n\n\ngood_units:\ntry(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n\n\nstudent_quantity:\nswitch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n\n\npercent_error:\ntry(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   \n\nright:\npercent_error <= settings['right']\n\n\nclose:\nright_sign and percent_error <= settings['close']\n\nright_sign:\nsign(student_scalar) = sign(correct_quantity)", "description": "<p>A value with units marked right if within an adjustable % error of the correct value. &nbsp;Marked close if within a wider margin of error.</p>", "published": false, "can_be_gap": true, "extensions": ["quantities"], "marking_notes": [{"description": "This is the main marking note. It should award credit and provide feedback based on the student's answer.", "definition": "switch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)\n", "name": "mark"}, {"description": "A value representing the student's answer to this part.", "definition": "qty(student_scalar, student_units)\n\n", "name": "interpreted_answer"}, {"description": "", "definition": "settings[\"correctAnswer\"]\n\n", "name": "correct_quantity"}, {"description": "", "definition": "units(correct_quantity)\n", "name": "correct_units"}, {"description": "", "definition": "[\"plain\",\"en\"]", "name": "allowed_notation_styles"}, {"description": "", "definition": "matchnumber(studentAnswer,allowed_notation_styles)", "name": "match_student_number"}, {"description": "", "definition": "match_student_number[1]", "name": "student_scalar"}, {"description": "<p>Does clumsy substitution to</p>\n<p>1.&nbsp;replace '-'&nbsp;with ' '&nbsp;</p>\n<p>2. replace '&deg;' with ' deg'&nbsp;</p>\n<p>to allow answers like 10 ft-lb and 30&deg;</p>", "definition": "join(\nsplit(\njoin(\nsplit(studentAnswer[len(match_student_number[0])..len(studentAnswer)]\n,'-'),' ')\n,'\u00b0'),' deg')\n", "name": "student_units"}, {"description": "", "definition": "try(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n", "name": "good_units"}, {"description": "<p>This fixes the student answer for two common errors. &nbsp;</p>\n<p>If student_units are wrong&nbsp;&nbsp;- replace with correct units</p>\n<p>If student_scalar has the wrong sign - replace with right sign</p>\n<p>If student makes both errors, only one gets fixed.</p>", "definition": "switch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n", "name": "student_quantity"}, {"description": "", "definition": "try(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   ", "name": "percent_error"}, {"description": "", "definition": "percent_error <= settings['right']\n", "name": "right"}, {"description": "<p>Only marked close if the student actually has the right sign.</p>", "definition": "right_sign and percent_error <= settings['close']", "name": "close"}, {"description": "", "definition": "sign(student_scalar) = sign(correct_quantity) ", "name": "right_sign"}], "short_name": "engineering-answer", "input_widget": "string", "settings": [{"label": "Correct Quantity.", "hint": "The correct answer given as a JME quantity.", "default_value": "", "name": "correctAnswer", "help_url": "", "input_type": "code", "evaluate": true}, {"label": "% Accuracy for right.", "hint": "Question will be considered correct if the scalar part of the student's answer is within this % of correct value.", "default_value": "0.2", "name": "right", "help_url": "", "input_type": "code", "evaluate": true}, {"label": "% Accuracy for close.", "hint": "Question will be considered close if the scalar part of the student's answer is within this % of correct value.", "default_value": "1.0", "name": "close", "help_url": "", "input_type": "code", "evaluate": true}, {"label": "Close with units.", "hint": "Partial Credit for close value with appropriate units. &nbsp;if correct answer is 100 N and close is &plusmn;1%,<br />99 &nbsp;N is accepted.", "default_value": "75", "name": "C1", "help_url": "", "input_type": "percent"}, {"label": "No units or wrong sign", "hint": "Partial credit for forgetting units or using wrong sign.<br />If the correct answer is 100 N, both 100 and -100 N are accepted.", "default_value": "50", "name": "C2", "help_url": "", "input_type": "percent"}, {"label": "Close, no units.", "hint": "Partial Credit for close value but forgotten units.<br />This value would be close&nbsp;if the expected units were provided. &nbsp;If the correct answer is 100 N, and close is &plusmn;1%,<br />99 is accepted.", "default_value": "25", "name": "C3", "help_url": "", "input_type": "percent"}], "source": {"pk": 19, "author": {"pk": 2530, "name": "William Haynes"}, "edit_page": "/part_type/19/edit"}, "can_be_step": true, "public_availability": "restricted"}], "resources": [["question-resources/shearFBD.png", "/srv/numbas/media/question-resources/shearFBD.png"], ["question-resources/shearFBD.gif", "/srv/numbas/media/question-resources/shearFBD.gif"], ["question-resources/shear.gif", "/srv/numbas/media/question-resources/shear.gif"], ["question-resources/shearFBD_XWaIQSM.png", "/srv/numbas/media/question-resources/shearFBD_XWaIQSM.png"]], "navigation": {"preventleave": false, "showfrontpage": false, "allowregen": true}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"variable_groups": [{"variables": ["theta", "dp_1", "ma_1", "dp_2", "ma_2", "ma_t", "version", "given", "Statement", "answer"], "name": "Unnamed group"}], "variablesTest": {"maxRuns": 100, "condition": ""}, "parts": [{"prompt": "<p>Mechanical advantage =&nbsp;[[0]]&nbsp; &nbsp;&nbsp;$\\var{if(version=0,latex('A'),latex('P'))}$ =&nbsp;[[1]]&nbsp;</p>", "unitTests": [], "scripts": {}, "showCorrectAnswer": true, "showFeedbackIcon": true, "sortAnswers": false, "variableReplacementStrategy": "originalfirst", "type": "gapfill", "gaps": [{"marks": "4", "variableReplacementStrategy": "originalfirst", "extendBaseMarkingAlgorithm": true, "type": "engineering-answer", "showCorrectAnswer": true, "settings": {"close": "1.0", "C2": "50", "correctAnswer": "qty(MA_t,\"\")", "C1": "75", "right": "0.2", "C3": "25"}, "customMarkingAlgorithm": "", "unitTests": [], "showFeedbackIcon": true, "variableReplacements": [], "scripts": {}}, {"marks": "4", "variableReplacementStrategy": "originalfirst", "extendBaseMarkingAlgorithm": true, "type": "engineering-answer", "showCorrectAnswer": true, "settings": {"close": "1.0", "C2": "50", "correctAnswer": "answer[version]", "C1": "75", "right": "0.2", "C3": "25"}, "customMarkingAlgorithm": "", "unitTests": [], "showFeedbackIcon": true, "variableReplacements": [], "scripts": {}}], "customMarkingAlgorithm": "", "marks": 0, "extendBaseMarkingAlgorithm": true, "variableReplacements": []}], "contributors": [{"profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/", "name": "William Haynes"}], "extensions": ["geogebra", "quantities", "visjs", "weh"], "functions": {"dist": {"language": "jme", "definition": "let(n,ix,['d_' + n,'\\\"'+ scalar(precround(eval(expression('d_' + n)),2))+ ' ' + units[1] + '\\\"'])", "type": "list", "parameters": [["ix", "number"]]}}, "advice": "<p><strong>1. Draw and number a neat, consistent, labeled set of free body diagrams for&nbsp;the parts of the shears.</strong></p>\n<p style=\"padding-left: 30px;\">Note link DE is a two force member with a slope of $m = \\dfrac{\\var{d_7}} {\\var{d_9 - d_4}}$, in tension. &nbsp;The free body diagrams must be consistent with each other and internal forces must occur in action-reaction pairs.</p>\n<p style=\"padding-left: 30px;\"><img src=\"resources/question-resources/shearFBD_XWaIQSM.png\" width=\"805\" height=\"303\"/></p>\n<p style=\"padding-left: 30px;\"></p>\n<p><strong>2. Find the necessary angle.</strong></p>\n<p style=\"padding-left: 30px;\">Let $\\theta$ be the angle that <em>DE</em> makes with the horizontal.</p>\n<p style=\"padding-left: 30px;\">$\\theta = \\tan^{-1} \\dfrac{\\var{d_7}}{\\var{d_9-d_4}} = \\var{siground(theta,4)}&deg;$</p>\n<p><strong>3. Use FBD II&nbsp;&nbsp;and take moments about <em>C</em> to find the tension in <em>DE</em>.</strong></p>\n<p style=\"padding-left: 30px;\">$\\begin{align}\\textrm{II: }\\Sigma M_C &amp;= 0 \\\\&nbsp;DE_x (\\var{d_6}) + DE_y(\\var{d_3}) &amp;= P \\,(\\var{d_3 + d_4&nbsp;})\\\\&nbsp;DE &amp;= \\left ( \\dfrac{\\var{scalar(d_3 + d_4&nbsp;)}}{\\var{scalar(d_6)} \\cos \\theta + \\var{scalar(d_3)}\\sin \\theta}&nbsp;\\right)&nbsp; P &nbsp;\\\\&amp;= \\left(\\dfrac{\\var{d_3 + d_4&nbsp;}}{\\var{siground(dp_1,4)}} \\right )P \\\\&amp;= \\var{siground(ma_1,4)} \\,P \\textrm{ (tension) } \\end{align}$</p>\n<p><strong>4.&nbsp;FBD III&nbsp;and take moments about B&nbsp;to find the cutting force at A.</strong></p>\n<p style=\"padding-left: 30px;\">$\\begin{align}\\textrm{III: }\\Sigma M_A &amp;= 0 \\\\ A (\\var{d_1}) &amp;=DE_x (\\var{d_5 +d_6 +d_7 }) + DE_y(\\var{d_8-d_1}) \\\\ A&nbsp;&amp;= \\left ( \\dfrac{\\var{scalar(d_5 +d_6 +d_7 )}\\cos \\theta + \\var{scalar(d_8-d_1)}\\sin \\theta}{\\var{scalar(d_1 )}} \\right) DE \\\\&amp;= \\left(\\dfrac{\\var{siground(dp_2,4)}}{\\var{d_1 }} \\right )DE&nbsp;\\\\&amp;= \\var{siground(ma_2,4)} \\,DE&nbsp;\\\\ &amp;= \\var{siground(ma_2,4)}\\, ( \\var{siground(ma_1,4)} \\,P )\\\\ &amp;= \\var{siground(ma_t,4)} P \\end{align}$</p>\n<p><strong>5. &nbsp;Find&nbsp;the mechanical advantage</strong></p>\n<p style=\"padding-left: 30px;\">Mechanical Advantage = &nbsp;$\\dfrac{A}{P} = \\var{siground(ma_t,4)}$</p>\n<p><strong>6. Find the {statement[version]}.</strong></p>\n<p style=\"padding-left: 30px;\">$A = \\var{siground(ma_t,4)} P$, so for <span data-jme-visible=\"version=1\">$A$</span> <span data-jme-visible=\"version=0\">$P$</span> = {given[version]}:</p>\n<p style=\"padding-left: 30px;\"></p>\n<p style=\"padding-left: 30px;\"><span data-jme-visible=\"version=0\">$A = \\var{siground(ma_t,4)} (\\var{given[0]}) = \\var{siground(answer[0],4)}$</span></p>\n<p style=\"padding-left: 30px;\"><span data-jme-visible=\"version=1\">$P = \\dfrac{\\var{given[1]}}{\\var{siground(ma_t,4)}} = \\var{siground(answer[1],4)}$</span></p>", "rulesets": {}, "metadata": {"description": "<p>Determine the mechanical advantage of a pair of compound lopping shears.</p>", "licence": "Creative Commons Attribution-NonCommercial 4.0 International"}, "statement": "<p>{geogebra_applet('csq3j7yg', {map(dist(n),n,1..9)} )}</p>\n<p>Determine</p>\n<p style=\"padding-left: 30px;\">a) the mechanical advantage of the lopping shears shown, and&nbsp;</p>\n<p style=\"padding-left: 30px;\">b) the {statement[version]}.</p>\n<p>Note: The mechanical advantage of the shears is the ratio of the cutting force developed at $A$ to the force $P$ applied to the handles.</p>\n<p style=\"padding-left: 30px;\"></p>", "variables": {"ma_2": {"description": "<p>mechanical advantage 2</p>", "templateType": "anything", "definition": "scalar(dp_2)/scalar(d_1)", "name": "ma_2", "group": "Unnamed group"}, "version": {"description": "", "templateType": "anything", "definition": "random(0,1)", "name": "version", "group": "Unnamed group"}, "L": {"description": "", "templateType": "anything", "definition": "round(qty(random(24..36#2),'in') as units[1],units[2])", "name": "L", "group": "Ungrouped variables"}, "d_6": {"description": "", "templateType": "anything", "definition": "round((2.59 / 24.4 L random(0.9..1.1#0.05)),units[3])", "name": "d_6", "group": "Ungrouped variables"}, "d_5": {"description": "", "templateType": "anything", "definition": "round((0.47 / 24.4 L random(0.9..1.1#0.05)),units[3])", "name": "d_5", "group": "Ungrouped variables"}, "d_2": {"description": "", "templateType": "anything", "definition": "round((3 / 24.4 L random(0.9..1.1#0.05)),units[2])", "name": "d_2", "group": "Ungrouped variables"}, "ma_t": {"description": "<p>total mechanical advantage</p>", "templateType": "anything", "definition": "ma_1 * ma_2", "name": "ma_t", "group": "Unnamed group"}, "Statement": {"description": "", "templateType": "anything", "definition": "[\"shearing force created at <i>A</i> when a \" + \n  given[0] + \" force is applied to the handles at <i>P</i>\", \n  'force <i>P</i> that must be applied to the handles in order to create a '\n  + given[1] + ' shearing force at <i>A</i>']", "name": "Statement", "group": "Unnamed group"}, "d_8": {"description": "", "templateType": "anything", "definition": "L-d_9", "name": "d_8", "group": "Ungrouped variables"}, "d_1": {"description": "", "templateType": "anything", "definition": "round((4.1 / 24.4 L random(0.9..1.1#0.05)),units[2])", "name": "d_1", "group": "Ungrouped variables"}, "d_9": {"description": "", "templateType": "anything", "definition": "d_4+ round((0.4 / 24.4 L random(0.9..1.1#0.05)),units[3])", "name": "d_9", "group": "Ungrouped variables"}, "theta": {"description": "", "templateType": "anything", "definition": "degrees(atan2(scalar(d_7),scalar(d_9-d_4)))", "name": "theta", "group": "Unnamed group"}, "answer": {"description": "", "templateType": "anything", "definition": "[given[0] * ma_t, given[1] / ma_t]", "name": "answer", "group": "Unnamed group"}, "dp_1": {"description": "<p>perpendicular distance from DE to C</p>", "templateType": "anything", "definition": "(d_6 cos(radians(theta)) + d_3 sin(radians(theta)))", "name": "dp_1", "group": "Unnamed group"}, "given": {"description": "<p><strong>]]</strong></p>", "templateType": "anything", "definition": "[qty(random(10..40#5),units[0]),qty(random(50..200#10),units[0])]", "name": "given", "group": "Unnamed group"}, "d_7": {"description": "", "templateType": "anything", "definition": "round((1.3 / 24.4 L random(0.9..1.1#0.05)),units[3])", "name": "d_7", "group": "Ungrouped variables"}, "d_3": {"description": "", "templateType": "anything", "definition": "round((5.2 / 24.4 L random(0.9..1.1#0.05)),units[2])", "name": "d_3", "group": "Ungrouped variables"}, "units": {"description": "<p>units[2] large rounding base</p>\n<p>units[3] small rounding base</p>", "templateType": "anything", "definition": "random(['N','mm', \"10 mm\", \"1 mm\" ],['lb','in', \"0.1 in\", \"0.1 in\"])", "name": "units", "group": "Ungrouped variables"}, "dp_2": {"description": "<p>perpendicular distance from DE to B</p>", "templateType": "anything", "definition": "(d_5 + d_6 +d_7 ) cos(radians(theta)) + (d_8-d_1) sin(radians(theta))", "name": "dp_2", "group": "Unnamed group"}, "d_4": {"description": "", "templateType": "anything", "definition": "L - (d_1 + d_2 + d_3)", "name": "d_4", "group": "Ungrouped variables"}, "ma_1": {"description": "<p>mechanical advantage 1</p>", "templateType": "anything", "definition": "scalar(d_3 + d_4 )/scalar(dp_1)", "name": "ma_1", "group": "Unnamed group"}}, "preamble": {"js": "", "css": ""}, "ungrouped_variables": ["units", "L", "d_1", "d_2", "d_3", "d_4", "d_5", "d_6", "d_7", "d_8", "d_9"], "tags": ["machine", "mechanics", "Mechanics", "Rigid body equilibrium", "rigid body equilibrium", "statics", "Statics"], "name": "Machine: lopping shears"}]}], "contributors": [{"profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/", "name": "William Haynes"}]}