// Numbas version: finer_feedback_settings
{"name": "Centroidal moment of inertia: box beam", "extensions": ["geogebra", "quantities"], "custom_part_types": [{"source": {"pk": 19, "author": {"name": "William Haynes", "pk": 2530}, "edit_page": "/part_type/19/edit"}, "name": "Engineering Accuracy with units", "short_name": "engineering-answer", "description": "<p>A value with units marked right if within an adjustable % error of the correct value. &nbsp;Marked close if within a wider margin of error.</p>", "help_url": "", "input_widget": "string", "input_options": {"correctAnswer": "siground(settings['correctAnswer'],4)", "hint": {"static": true, "value": ""}, "allowEmpty": {"static": true, "value": true}}, "can_be_gap": true, "can_be_step": true, "marking_script": "mark:\nswitch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)\n\ninterpreted_answer:\nqty(student_scalar, student_units)\n\n\n\ncorrect_quantity:\nsettings[\"correctAnswer\"]\n\n\n\ncorrect_units:\nunits(correct_quantity)\n\n\nallowed_notation_styles:\n[\"plain\",\"en\"]\n\nmatch_student_number:\nmatchnumber(studentAnswer,allowed_notation_styles)\n\nstudent_scalar:\nmatch_student_number[1]\n\nstudent_units:\nreplace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")\n\ngood_units:\ntry(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n\n\nstudent_quantity:\nswitch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n\n\npercent_error:\ntry(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   \n\nright:\npercent_error <= settings['right']\n\n\nclose:\nright_sign and percent_error <= settings['close']\n\nright_sign:\nsign(student_scalar) = sign(correct_quantity)", "marking_notes": [{"name": "mark", "description": "This is the main marking note. It should award credit and provide feedback based on the student's answer.", "definition": "switch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)"}, {"name": "interpreted_answer", "description": "A value representing the student's answer to this part.", "definition": "qty(student_scalar, student_units)\n\n"}, {"name": "correct_quantity", "description": "", "definition": "settings[\"correctAnswer\"]\n\n"}, {"name": "correct_units", "description": "", "definition": "units(correct_quantity)\n"}, {"name": "allowed_notation_styles", "description": "", "definition": "[\"plain\",\"en\"]"}, {"name": "match_student_number", "description": "", "definition": "matchnumber(studentAnswer,allowed_notation_styles)"}, {"name": "student_scalar", "description": "", "definition": "match_student_number[1]"}, {"name": "student_units", "description": "<p>Modify&nbsp;the&nbsp;unit portion of&nbsp;the student's answer by</p>\n<p>1. replacing \"ohms\" with \"ohm\" &nbsp;case insensitive</p>\n<p>2.&nbsp;replacing '-'&nbsp;with ' '&nbsp;</p>\n<p>3. replacing '&deg;' with ' deg'&nbsp;</p>\n<p>to allow answers like 10 ft-lb and 30&deg;</p>", "definition": "replace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")"}, {"name": "good_units", "description": "", "definition": "try(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n"}, {"name": "student_quantity", "description": "<p>This fixes the student answer for two common errors. &nbsp;</p>\n<p>If student_units are wrong&nbsp;&nbsp;- replace with correct units</p>\n<p>If student_scalar has the wrong sign - replace with right sign</p>\n<p>If student makes both errors, only one gets fixed.</p>", "definition": "switch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n"}, {"name": "percent_error", "description": "", "definition": "try(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   "}, {"name": "right", "description": "", "definition": "percent_error <= settings['right']\n"}, {"name": "close", "description": "<p>Only marked close if the student actually has the right sign.</p>", "definition": "right_sign and percent_error <= settings['close']"}, {"name": "right_sign", "description": "", "definition": "sign(student_scalar) = sign(correct_quantity) "}], "settings": [{"name": "correctAnswer", "label": "Correct Quantity.", "help_url": "", "hint": "The correct answer given as a JME quantity.", "input_type": "code", "default_value": "", "evaluate": true}, {"name": "right", "label": "% Accuracy for right.", "help_url": "", "hint": "Question will be considered correct if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "0.2", "evaluate": true}, {"name": "close", "label": "% Accuracy for close.", "help_url": "", "hint": "Question will be considered close if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "1.0", "evaluate": true}, {"name": "C1", "label": "Close with units.", "help_url": "", "hint": "Partial Credit for close value with appropriate units. &nbsp;if correct answer is 100 N and close is &plusmn;1%,<br />99 &nbsp;N is accepted.", "input_type": "percent", "default_value": "75"}, {"name": "C2", "label": "No units or wrong sign", "help_url": "", "hint": "Partial credit for forgetting units or using wrong sign.<br />If the correct answer is 100 N, both 100 and -100 N are accepted.", "input_type": "percent", "default_value": "50"}, {"name": "C3", "label": "Close, no units.", "help_url": "", "hint": "Partial Credit for close value but forgotten units.<br />This value would be close&nbsp;if the expected units were provided. &nbsp;If the correct answer is 100 N, and close is &plusmn;1%,<br />99 is accepted.", "input_type": "percent", "default_value": "25"}], "public_availability": "always", "published": true, "extensions": ["quantities"]}], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Centroidal moment of inertia: box beam", "tags": ["Mechanics", "mechanics", "moment of inertia", "parallel axis theorem", "rectangular beam", "Statics", "statics"], "metadata": {"description": "<p>Calculate the centroidal moment of intertia of a rectangle with a excentric rectangular hole.</p>", "licence": "Creative Commons Attribution-NonCommercial 4.0 International"}, "statement": "<p>The cross-section of a hollow box beam is shown. &nbsp;Grid units are [{units}].&nbsp;</p>\n<table style=\"margin-left: auto; margin-right: auto;\">\n<tbody>\n<tr>\n<td>\n<p style=\"text-align: center;\">{geogebra_applet('pv5edgwm',[[\"A\",A],[\"B\",B]])}</p>\n</td>\n</tr>\n</tbody>\n</table>", "advice": "<p><strong>Part 1: Big Rectangle:</strong></p>\n<p style=\"padding-left: 30px;\">$b_1$ = {b1}, $h_1$ = {h1}, &nbsp;$A_1$ = {A1}, $\\bar{y}_1$ = {ybar1}</p>\n<p><strong>Part 2: Small rectangle (removed):</strong></p>\n<p style=\"padding-left: 30px;\">$b_2$ = {b2}, $h_2$ = {h2}, &nbsp;$A_2$ = {A2}, $\\bar{y}_2$ = {ybar2}</p>\n<p><strong>Total Area:</strong></p>\n<p style=\"padding-left: 30px;\">$A = \\Sigma A_i = A_1 - A_2 =$ {display(Area)}</p>\n<p><strong>Centroid:</strong></p>\n<p style=\"padding-left: 30px;\">$\\bar{y} = \\dfrac{\\Sigma A_i \\bar{y}_i}{\\Sigma A_i} = \\dfrac{A_1 \\bar{y}_1 - A_2 \\bar{y}_2}{ A_1 - A_2}$ = {display(ybar)}</p>\n<p><strong>Moments of inerta about the <em>x</em>&nbsp;axis:</strong></p>\n<p style=\"padding-left: 30px;\">$I_{x_1} = \\dfrac{b_1 (h_1)^3}{3} = $ {display(Ix1)}</p>\n<p style=\"padding-left: 30px;\">$I_{x_2}= \\bar{I} + A d^2 = \\dfrac{b_2&nbsp;(h_2)^3}{12} + A_2 (\\bar{y}_2)^2 = $ {display(Ix2)}</p>\n<p style=\"padding-left: 30px;\">$I_x = I_{x_1} - I_{x_2}= $ {display(Ix)}</p>\n<p><strong>Centroidal Moment of inertia:</strong></p>\n<p style=\"padding-left: 30px;\">Using the parallel axis theorem,</p>\n<p style=\"padding-left: 30px;\">$I = \\bar{I} + A d^2 \\implies&nbsp;\\bar{I}_x= I_x - A (\\bar{y})^2$ = {display(Ibar)}</p>", "rulesets": {}, "extensions": ["geogebra", "quantities"], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"Ix1": {"name": "Ix1", "group": "Unnamed group", "definition": "b1 h1 h1 h1/3", "description": "", "templateType": "anything", "can_override": false}, "b1": {"name": "b1", "group": "Unnamed group", "definition": "qty(A[0],units)", "description": "<p>base for large rectangle</p>", "templateType": "anything", "can_override": false}, "A": {"name": "A", "group": "Ungrouped variables", "definition": "vector(random(4..20#2),random(8..18))", "description": "<p>top right corner of large rectangle</p>", "templateType": "anything", "can_override": false}, "Area": {"name": "Area", "group": "Unnamed group", "definition": "A1-A2", "description": "", "templateType": "anything", "can_override": false}, "ybar": {"name": "ybar", "group": "Unnamed group", "definition": "(A1 ybar1 - A2 ybar2)/(A1-A2)", "description": "", "templateType": "anything", "can_override": false}, "A1": {"name": "A1", "group": "Unnamed group", "definition": "b1 h1", "description": "", "templateType": "anything", "can_override": false}, "debug": {"name": "debug", "group": "Ungrouped variables", "definition": "false", "description": "", "templateType": "anything", "can_override": false}, "Ibar": {"name": "Ibar", "group": "Unnamed group", "definition": "Ix - area ybar ybar", "description": "", "templateType": "anything", "can_override": false}, "t": {"name": "t", "group": "Ungrouped variables", "definition": "random(1..3)", "description": "<p>thickness of sides and top</p>", "templateType": "anything", "can_override": false}, "Ix": {"name": "Ix", "group": "Unnamed group", "definition": "Ix1 - Ix2", "description": "", "templateType": "anything", "can_override": false}, "b2": {"name": "b2", "group": "Unnamed group", "definition": "qty(A[0] - 2 t,units)", "description": "", "templateType": "anything", "can_override": false}, "h2": {"name": "h2", "group": "Unnamed group", "definition": "qty(A[1]-t - B[1],units)", "description": "", "templateType": "anything", "can_override": false}, "ybar1": {"name": "ybar1", "group": "Unnamed group", "definition": "h1/2", "description": "", "templateType": "anything", "can_override": false}, "B": {"name": "B", "group": "Ungrouped variables", "definition": "vector(t, random(2..(A[1]-t-1)))", "description": "", "templateType": "anything", "can_override": false}, "h1": {"name": "h1", "group": "Unnamed group", "definition": "qty(A[1],units)", "description": "<p>height of large rectangle</p>", "templateType": "anything", "can_override": false}, "Ix2": {"name": "Ix2", "group": "Unnamed group", "definition": "b2 h2 h2 h2 /12 + A2 ybar2 ybar2", "description": "", "templateType": "anything", "can_override": false}, "ybar2": {"name": "ybar2", "group": "Unnamed group", "definition": "qty(B[1] + scalar(h2/2),units)", "description": "", "templateType": "anything", "can_override": false}, "A2": {"name": "A2", "group": "Unnamed group", "definition": "b2 h2", "description": "", "templateType": "anything", "can_override": false}, "units": {"name": "units", "group": "Ungrouped variables", "definition": "random('in','mm','cm','ft')", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "B[0]<A[0]/2 and\nybar1<>ybar2", "maxRuns": 100}, "ungrouped_variables": ["t", "A", "B", "units", "debug"], "variable_groups": [{"name": "Unnamed group", "variables": ["b1", "h1", "b2", "h2", "A1", "A2", "ybar", "ybar1", "ybar2", "Ix1", "Ix2", "Ix", "Area", "Ibar"]}], "functions": {"display": {"parameters": [["q", "quantity"]], "type": "string", "language": "jme", "definition": "string(siground(q,5))"}}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Determine the vertical location of the centroid of the cross section. &nbsp;</p>\n<p>$\\bar{y} = $&nbsp;[[0]] <span data-jme-visible=\"debug\">{display(ybar)}</span></p>\n<p>Determine the moment of inertia of the shape about a horizontal ${x'}$ axis (the <em>neutral axis</em>) passing through the centroid. &nbsp;</p>\n<p>${I_{x'} = }$&nbsp;[[1]] <span data-jme-visible=\"debug\">{display(Ibar)}</span></p>", "gaps": [{"type": "engineering-answer", "useCustomName": false, "customName": "", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "{ybar}", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": false, "customName": "", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "{Ibar}", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}], "sortAnswers": false}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}], "resources": []}]}], "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}]}