// Numbas version: exam_results_page_options
{"name": "Equilibrium of a particle: two suspended loads", "extensions": ["/srv/numbas/media/user-extensions/extracted/41/geogebra", "/srv/numbas/media/user-extensions/extracted/55/quantities", "/srv/numbas/media/user-extensions/extracted/98/weh"], "custom_part_types": [{"marking_script": "mark:\nswitch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)\n\n\ninterpreted_answer:\nqty(student_scalar, student_units)\n\n\n\ncorrect_quantity:\nsettings[\"correctAnswer\"]\n\n\n\ncorrect_units:\nunits(correct_quantity)\n\n\nallowed_notation_styles:\n[\"plain\",\"en\"]\n\nmatch_student_number:\nmatchnumber(studentAnswer,allowed_notation_styles)\n\nstudent_scalar:\nmatch_student_number[1]\n\nstudent_units:\njoin(\nsplit(\njoin(\nsplit(studentAnswer[len(match_student_number[0])..len(studentAnswer)]\n,'-'),' ')\n,'\u00b0'),' deg')\n\n\ngood_units:\ntry(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n\n\nstudent_quantity:\nswitch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n\n\npercent_error:\ntry(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   \n\nright:\npercent_error <= settings['right']\n\n\nclose:\nright_sign and percent_error <= settings['close']\n\nright_sign:\nsign(student_scalar) = sign(correct_quantity)", "name": "Engineering Accuracy with units", "extensions": ["quantities"], "public_availability": "restricted", "input_options": {"correctAnswer": "siground(settings['correctAnswer'],4)", "allowEmpty": {"value": true, "static": true}, "hint": {"value": "", "static": true}}, "published": false, "input_widget": "string", "can_be_step": true, "marking_notes": [{"definition": "switch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)\n", "name": "mark", "description": "This is the main marking note. It should award credit and provide feedback based on the student's answer."}, {"definition": "qty(student_scalar, student_units)\n\n", "name": "interpreted_answer", "description": "A value representing the student's answer to this part."}, {"definition": "settings[\"correctAnswer\"]\n\n", "name": "correct_quantity", "description": ""}, {"definition": "units(correct_quantity)\n", "name": "correct_units", "description": ""}, {"definition": "[\"plain\",\"en\"]", "name": "allowed_notation_styles", "description": ""}, {"definition": "matchnumber(studentAnswer,allowed_notation_styles)", "name": "match_student_number", "description": ""}, {"definition": "match_student_number[1]", "name": "student_scalar", "description": ""}, {"definition": "join(\nsplit(\njoin(\nsplit(studentAnswer[len(match_student_number[0])..len(studentAnswer)]\n,'-'),' ')\n,'\u00b0'),' deg')\n", "name": "student_units", "description": "<p>Does clumsy substitution to</p>\n<p>1.&nbsp;replace '-'&nbsp;with ' '&nbsp;</p>\n<p>2. replace '&deg;' with ' deg'&nbsp;</p>\n<p>to allow answers like 10 ft-lb and 30&deg;</p>"}, {"definition": "try(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n", "name": "good_units", "description": ""}, {"definition": "switch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n", "name": "student_quantity", "description": "<p>This fixes the student answer for two common errors. &nbsp;</p>\n<p>If student_units are wrong&nbsp;&nbsp;- replace with correct units</p>\n<p>If student_scalar has the wrong sign - replace with right sign</p>\n<p>If student makes both errors, only one gets fixed.</p>"}, {"definition": "try(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   ", "name": "percent_error", "description": ""}, {"definition": "percent_error <= settings['right']\n", "name": "right", "description": ""}, {"definition": "right_sign and percent_error <= settings['close']", "name": "close", "description": "<p>Only marked close if the student actually has the right sign.</p>"}, {"definition": "sign(student_scalar) = sign(correct_quantity) ", "name": "right_sign", "description": ""}], "description": "<p>A value with units marked right if within an adjustable % error of the correct value. &nbsp;Marked close if within a wider margin of error.</p>", "short_name": "engineering-answer", "settings": [{"label": "Correct Quantity.", "name": "correctAnswer", "evaluate": true, "default_value": "", "help_url": "", "input_type": "code", "hint": "The correct answer given as a JME quantity."}, {"label": "% Accuracy for right.", "name": "right", "evaluate": true, "default_value": "0.2", "help_url": "", "input_type": "code", "hint": "Question will be considered correct if the scalar part of the student's answer is within this % of correct value."}, {"label": "% Accuracy for close.", "name": "close", "evaluate": true, "default_value": "1.0", "help_url": "", "input_type": "code", "hint": "Question will be considered close if the scalar part of the student's answer is within this % of correct value."}, {"label": "Close with units.", "name": "C1", "default_value": "75", "help_url": "", "input_type": "percent", "hint": "Partial Credit for close value with appropriate units. &nbsp;if correct answer is 100 N and close is &plusmn;1%,<br />99 &nbsp;N is accepted."}, {"label": "No units or wrong sign", "name": "C2", "default_value": "50", "help_url": "", "input_type": "percent", "hint": "Partial credit for forgetting units or using wrong sign.<br />If the correct answer is 100 N, both 100 and -100 N are accepted."}, {"label": "Close, no units.", "name": "C3", "default_value": "25", "help_url": "", "input_type": "percent", "hint": "Partial Credit for close value but forgotten units.<br />This value would be close&nbsp;if the expected units were provided. &nbsp;If the correct answer is 100 N, and close is &plusmn;1%,<br />99 is accepted."}], "help_url": "", "source": {"pk": 19, "edit_page": "/part_type/19/edit", "author": {"pk": 2530, "name": "William Haynes"}}, "can_be_gap": true}], "resources": [], "navigation": {"showfrontpage": false, "preventleave": false, "allowregen": true}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"ungrouped_variables": [], "parts": [{"prompt": "<p>Draw a free body diagram of point <em>A</em> and apply the equations of equilibrium to determine the tensions in cables <em>AB</em> and <em>AC</em>.</p>\n<p>$T_{AB}$ =&nbsp;[[0]] <span data-jme-visible=\"debug\">{siground(TB,4)}</span></p>\n<p>$T_{AC}$ =&nbsp;[[1]] <span data-jme-visible=\"debug\">{siground(TC,4)}</span></p>", "customMarkingAlgorithm": "", "gaps": [{"customMarkingAlgorithm": "", "scripts": {}, "variableReplacementStrategy": "originalfirst", "extendBaseMarkingAlgorithm": true, "showCorrectAnswer": true, "unitTests": [], "marks": "4", "showFeedbackIcon": true, "settings": {"right": "0.2", "C1": "75", "C2": "50", "close": "1.0", "correctAnswer": "TB", "C3": "25"}, "type": "engineering-answer", "variableReplacements": []}, {"customMarkingAlgorithm": "", "scripts": {}, "variableReplacementStrategy": "originalfirst", "extendBaseMarkingAlgorithm": true, "showCorrectAnswer": true, "unitTests": [], "marks": "4", "showFeedbackIcon": true, "settings": {"right": "0.2", "C1": "75", "C2": "50", "close": "1.0", "correctAnswer": "TC", "C3": "25"}, "type": "engineering-answer", "variableReplacements": []}], "extendBaseMarkingAlgorithm": true, "scripts": {}, "variableReplacements": [], "marks": 0, "variableReplacementStrategy": "originalfirst", "sortAnswers": false, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "type": "gapfill"}, {"prompt": "<p>Draw a free body diagram of point <em>D</em>&nbsp;and apply the equations of equilibrium to determine the tensions in cables <em>CE&nbsp;</em>and <em>CD</em>.</p>\n<p>$T_{CE}$ =&nbsp;[[0]] <span data-jme-visible=\"debug\">{siground(TE,4)}</span></p>\n<p>$T_{CD}$ =&nbsp;[[1]] <span data-jme-visible=\"debug\">{siground(TD,4)}</span></p>", "customMarkingAlgorithm": "", "gaps": [{"customMarkingAlgorithm": "", "scripts": {}, "variableReplacementStrategy": "originalfirst", "extendBaseMarkingAlgorithm": true, "showCorrectAnswer": true, "unitTests": [], "marks": "4", "showFeedbackIcon": true, "settings": {"right": "0.2", "C1": "75", "C2": "50", "close": "1.0", "correctAnswer": "TE", "C3": "25"}, "type": "engineering-answer", "variableReplacements": []}, {"customMarkingAlgorithm": "", "scripts": {}, "variableReplacementStrategy": "originalfirst", "extendBaseMarkingAlgorithm": true, "showCorrectAnswer": true, "unitTests": [], "marks": "4", "showFeedbackIcon": true, "settings": {"right": "0.2", "C1": "75", "C2": "50", "close": "1.0", "correctAnswer": "TD", "C3": "25"}, "type": "engineering-answer", "variableReplacements": []}], "extendBaseMarkingAlgorithm": true, "scripts": {}, "variableReplacements": [], "marks": 0, "variableReplacementStrategy": "originalfirst", "sortAnswers": false, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "type": "gapfill"}], "variable_groups": [{"variables": ["a1", "a2", "a3", "w1", "w2", "units", "debug"], "name": "input"}, {"variables": ["TC", "TE", "TD", "TB"], "name": "magnitudes"}, {"variables": ["forceB", "forceC", "forceE", "forceD"], "name": "check"}], "name": "Equilibrium of a particle: two suspended loads", "rulesets": {}, "metadata": {"licence": "Creative Commons Attribution-NonCommercial 4.0 International", "description": "<p>Find the tensions is a system of cables supporting two loads.</p>"}, "preamble": {"css": "", "js": ""}, "extensions": ["geogebra", "quantities", "weh"], "variablesTest": {"condition": "a1 <> a2 and a3>a2+10 and scalar(TD) > 0", "maxRuns": 100}, "variables": {"debug": {"templateType": "anything", "definition": "false", "name": "debug", "description": "", "group": "input"}, "w1": {"templateType": "anything", "definition": "qty(random(10..100#5),units)", "name": "w1", "description": "", "group": "input"}, "a1": {"templateType": "anything", "definition": "random(10..80#5)", "name": "a1", "description": "<p>angle of rope AB from horizontal, as shown on diagram.</p>", "group": "input"}, "TC": {"templateType": "anything", "definition": "w1 cos(radians(a1))/sin(radians(a1+a2))\n\n", "name": "TC", "description": "", "group": "magnitudes"}, "TD": {"templateType": "anything", "definition": "TC cos(radians(a2))-TE cos(radians(a3))", "name": "TD", "description": "", "group": "magnitudes"}, "TB": {"templateType": "anything", "definition": "w1 cos(radians(a2))/sin(radians(a1+a2))", "name": "TB", "description": "", "group": "magnitudes"}, "TE": {"templateType": "anything", "definition": "(W2 + TC sin(radians(A2)))/sin(radians(A3))", "name": "TE", "description": "", "group": "magnitudes"}, "w2": {"templateType": "anything", "definition": "qty(random(10..100#5),units)", "name": "w2", "description": "", "group": "input"}, "a3": {"templateType": "anything", "definition": "random(20..70#5)", "name": "a3", "description": "<p>angle of CE as shown on diagram.</p>", "group": "input"}, "forceC": {"templateType": "anything", "definition": "vector(cos(radians(a2)),sin(radians(a2))) scalar(TC)\n", "name": "forceC", "description": "", "group": "check"}, "a2": {"templateType": "anything", "definition": "random(10..70#5)", "name": "a2", "description": "<p>a2 is the angle from the horizontal to AC. &nbsp;Note angle shown on diagram is the complement of a2</p>", "group": "input"}, "forceB": {"templateType": "anything", "definition": "vector(-cos(radians(A1)),sin(radians(A1))) scalar(TB)", "name": "forceB", "description": "", "group": "check"}, "forceE": {"templateType": "anything", "definition": "vector(cos(radians(a3)),sin(radians(a3))) scalar(TE)", "name": "forceE", "description": "", "group": "check"}, "units": {"templateType": "anything", "definition": "'lb'", "name": "units", "description": "", "group": "input"}, "forceD": {"templateType": "anything", "definition": "vector(scalar(TD),0)", "name": "forceD", "description": "", "group": "check"}}, "tags": ["2-d", "equilibrium", "mechanics", "particle", "statics"], "statement": "<p>Two loads $W_1 = \\var{w1}$, and&nbsp;$W_2 = \\var{w2}$ are suspended from a system of cables as shown. &nbsp;Determine the tensions in the cables.</p>\n<p>{geogebra_applet('vzqtfrbq',[['a1',a1+\"&deg;\"],['a2',a2+\"&deg;\"],['a3',a3+\"&deg;\"]])}</p>\n<p></p>", "functions": {}, "contributors": [{"profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/", "name": "William Haynes"}], "type": "question", "advice": "<p>Begin by drawing a free body diagram of particle <em>A</em>, then apply the equations of equilibrium.</p>\n<p>$\\begin{align}\\\\ \\Sigma F_x = 0 &nbsp;\\\\-{T_B}_x + {T_C}_x = 0 \\end{align} \\\\ T_B \\cos(\\var{a1}&deg;) = T_C \\sin(\\var{90-a2}&deg;) \\qquad(1)$</p>\n<p>$\\begin{align}\\\\ \\Sigma F_y = 0 &nbsp;\\\\{T_B}_y + {T_C}_y&nbsp;- W_1 = 0 \\end{align} \\\\ T_B&nbsp;\\sin(\\var{a1}&deg;) +&nbsp;T_C \\cos(\\var{90-a2})&deg; = W_1 \\qquad(2)$</p>\n<p>Solve (1) and (2) simultaneously for $T_B$ and $T_C$.</p>\n<p>$T_B = \\var{siground(TB,4)}, \\quad T_C = \\var{siground(TC,4)}$</p>\n<hr/>\n<p>With $T_C$ known, use a similar approach on particle $C$ to find $T_D$ and $T_E$. &nbsp;</p>\n<p>Since cable<em> CD</em> is horizontal $T_D$&nbsp;has no <em>y</em>-component, so&nbsp;the $\\Sigma F_y = 0$ equation only has one unknown&nbsp;and&nbsp;simultaneous equations are unnecessary in this case.</p>\n<p>$T_D&nbsp;= \\var{siground(TD,4)}, \\quad T_E = \\var{siground(TE,4)}$</p>"}]}], "contributors": [{"profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/", "name": "William Haynes"}]}