// Numbas version: finer_feedback_settings
{"name": "Parallel Axis Theorem: common shapes", "extensions": ["geogebra"], "custom_part_types": [], "resources": [["question-resources/MOI_table.pdf", "MOI_table.pdf"], ["question-resources/MOI_table.png", "MOI_table.png"]], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Parallel Axis Theorem: common shapes", "tags": ["Mechanics", "mechanics", "moment of inertia", "parallel axis theorem", "second moment of area", "Statics", "statics"], "metadata": {"description": "<p>Write expressions for the moment of inertia of simple shapes about various axes.</p>", "licence": "Creative Commons Attribution-NonCommercial 4.0 International"}, "statement": "<p>Write an expression for the&nbsp;moment of inertia&nbsp;with respect to the {if(axis=1,\"centroidal \",\"\")} $x$&nbsp;axis for&nbsp;the {shape} with {description}. &nbsp;</p>\n<table style=\"margin-left: auto; margin-right: auto;\">\n<tbody>\n<tr>\n<td>{geogebra_applet('j25dzmbq', [[\"A\",A],[\"shape\",index +1],[\"axis\",axis],[\"reference\",ref]])}</td>\n</tr>\n</tbody>\n</table>", "advice": "<p>Some cases are found in the table below. &nbsp;For others use the parallel axis theorem.</p>\n<p>\\[I = \\bar{I} + A d^2\\]</p>\n<p>&nbsp;where:</p>\n<p style=\"padding-left: 30px;\">$\\bar{I}$ is the centroidal moment of inertia, i.e.,&nbsp;the moment of inertia of the shape with respect to an axis which is parallel to the desired axis and passes throught the shape's centroid.</p>\n<p style=\"padding-left: 30px;\">$A$ is the area of the shape.</p>\n<p style=\"padding-left: 30px;\">$d$ is the distance between the desired axis and the parallel centroidal axis.</p>\n<p style=\"padding-left: 30px;\"></p>\n<p style=\"padding-left: 30px;\"></p>\n<p><img src=\"resources/question-resources/MOI_table.png\" width=\"300\" height=\"732\" style=\"display: block; margin-left: auto; margin-right: auto;\"/></p>\n<p></p>\n<p></p>\n<p></p>", "rulesets": {}, "extensions": ["geogebra"], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"distances": {"name": "distances", "group": "Ungrouped variables", "definition": "let(d,string(down),u,string(up),[[\"0\",\"0\",\"0\",\"0\"],\n  [d,d,d,d],\n  [u,u,u,u],\n  [\"0\",\"c\",\"c+(\"+d+\")\",\"c-(\"+u+\")\"],\n  [\"0\",\"c\",\"c-\"+d,\"c+\"+u]])", "description": "", "templateType": "anything", "can_override": false}, "names": {"name": "names", "group": "Ungrouped variables", "definition": "[\"rectangle\",\"triangle\",\"circle\",\"semicircle\",\"quarter circle\"]", "description": "", "templateType": "anything", "can_override": false}, "ref": {"name": "ref", "group": "Ungrouped variables", "definition": "if(axis>3,random(1..3),0)", "description": "<p>The reference axis,&nbsp;needed for axis 4 and 5. &nbsp;</p>\n<p>1 centroid</p>\n<p>2 bottom of shape</p>\n<p>3 top of shape</p>", "templateType": "anything", "can_override": false}, "shape": {"name": "shape", "group": "parts", "definition": "names[index]", "description": "", "templateType": "anything", "can_override": false}, "index": {"name": "index", "group": "Ungrouped variables", "definition": "random(0..4)", "description": "<p>rectangle, triangle, circle, semicircle, quarter circle</p>", "templateType": "anything", "can_override": false}, "area": {"name": "area", "group": "parts", "definition": "expression(areas[index])", "description": "", "templateType": "anything", "can_override": false}, "areas": {"name": "areas", "group": "Ungrouped variables", "definition": "[\"b*h\", \"b h / 2\", \"pi r^2\", \"pi r^2 /2\" , \"pi r^2 /4\"]", "description": "", "templateType": "anything", "can_override": false}, "maximums": {"name": "maximums", "group": "Ungrouped variables", "definition": "[\"h/2\",\n  if(quadrant<2,\"2 h/3\",\"h/3\")\n  ,\"r\",\n  if(quadrant=0,\"(r-(4 r)/(3 pi))\",if(quadrant=2,\"(4 r)/(3 pi)\",\"r\")),\n  if(quadrant<=1,\"(r-(4 r)/(3 pi))\", \"(4 r)/(3 pi)\")]", "description": "<p>Distance from centroid to the top of the shape.</p>\n<p></p>\n<p>maximums = {abs(y(A) / 2), Max(2 / 3 y(A), (-y(A)) / 3), radius, If(quadrant \u225f 0, radius - ybar, quadrant \u225f 1, radius, quadrant \u225f 2, ybar, quadrant \u225f 3, radius), If(quadrant &le; 1, radius - ybar, ybar)}</p>", "templateType": "anything", "can_override": false}, "quadrant": {"name": "quadrant", "group": "parts", "definition": "if(A[0]>0 and A[1]>0,0,\n  if(A[0]<0 and A[1] > 0, 1,\n    if(A[0]<0 and A[1]<0,2,3)))", "description": "<p>Calculates which quadrant point A is in. &nbsp;Needed to determine orientation of the triangle, semi and quarter circles.</p>", "templateType": "anything", "can_override": false}, "even_quadrant": {"name": "even_quadrant", "group": "Ungrouped variables", "definition": "mod(quadrant,2)=0 ", "description": "<p>usefull utility</p>", "templateType": "anything", "can_override": false}, "Ibars": {"name": "Ibars", "group": "Ungrouped variables", "definition": "[\"b*h^3/12\",\n  \"b*h^3/36\", \n  \"pi r^4/4\", \n  if(even_quadrant,\"0.1098 r^4\",\"pi r^4/8\"),\n  \"0.0549 r^4\"]", "description": "", "templateType": "anything", "can_override": false}, "minimums": {"name": "minimums", "group": "Ungrouped variables", "definition": "[\"h/2\",\n  if(quadrant<2,\"h/3\",\"2 h/3\")\n  ,\"r\",\n  if(quadrant=0,\"(4 r)/(3 pi)\",if(quadrant=2,\"r-(4 r)/(3 pi)\",\"r\")),\n  if(quadrant<=1,\"(4 r)/(3 pi)\", \"r-(4 r)/(3 pi)\")]", "description": "<p>distance from centroid to bottom of the shape</p>", "templateType": "anything", "can_override": false}, "Ibar": {"name": "Ibar", "group": "parts", "definition": "expression(Ibars[index])", "description": "", "templateType": "anything", "can_override": false}, "axis": {"name": "axis", "group": "Ungrouped variables", "definition": "random(1..5)", "description": "<p>1 through centroid</p>\n<p>2 bottom of shape</p>\n<p>3 top of shape</p>\n<p>4 below shape</p>\n<p>5 above shape</p>", "templateType": "anything", "can_override": false}, "down": {"name": "down", "group": "parts", "definition": "expression(minimums[index])", "description": "<p>distance from centroid to bottom of shape</p>", "templateType": "anything", "can_override": false}, "up": {"name": "up", "group": "parts", "definition": "expression(maximums[index])", "description": "", "templateType": "anything", "can_override": false}, "d": {"name": "d", "group": "parts", "definition": "expression(distances[axis-1][ref])", "description": "", "templateType": "anything", "can_override": false}, "description": {"name": "description", "group": "parts", "definition": "if(index<2,\"base  $b$ and height $h$\", \"radius $r$\")", "description": "", "templateType": "anything", "can_override": false}, "Ix": {"name": "Ix", "group": "parts", "definition": "(string(Ibar) + \"+ \" + string(area) + \"*(\"+ string(d) + \")^2\")", "description": "", "templateType": "anything", "can_override": false}, "A": {"name": "A", "group": "Ungrouped variables", "definition": "vector(random([2,3],[3,2],[-3,2],[-2,3])) random(-1,1)\n", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["A", "index", "axis", "ref", "names", "areas", "Ibars", "maximums", "even_quadrant", "minimums", "distances"], "variable_groups": [{"name": "parts", "variables": ["shape", "area", "description", "Ibar", "up", "quadrant", "down", "d", "Ix"]}], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>$\\var{if(axis=1,latex(\"\\\\bar\\{I\\}_x\"), latex(\"I_x\"))}= $&nbsp;[[0]]</p>\n<p>Note: If your answer produces an error message,&nbsp;try&nbsp;putting an asterisk&nbsp;before&nbsp;parentheses, like so:&nbsp;<code>b h * (h/2)^2&nbsp;</code></p>\n<p></p>\n<p><span data-jme-visible=\"false\">\\[I_x=\\simplify{{expression(Ix)}}\\]</span></p>\n<p></p>", "gaps": [{"type": "jme", "useCustomName": true, "customName": "MOI", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "answer": "{expression(Ix)}", "showPreview": true, "checkingType": "absdiff", "checkingAccuracy": 0.001, "failureRate": 1, "vsetRangePoints": 5, "vsetRange": [0, 1], "checkVariableNames": true, "singleLetterVariables": true, "allowUnknownFunctions": false, "implicitFunctionComposition": false, "caseSensitive": false, "valuegenerators": []}], "sortAnswers": false}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}], "resources": ["question-resources/MOI_table.pdf", "question-resources/MOI_table.png"]}]}], "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}]}