// Numbas version: finer_feedback_settings
{"name": "Moment of Inertia: rectangles", "extensions": ["geogebra", "quantities"], "custom_part_types": [{"source": {"pk": 19, "author": {"name": "William Haynes", "pk": 2530}, "edit_page": "/part_type/19/edit"}, "name": "Engineering Accuracy with units", "short_name": "engineering-answer", "description": "<p>A value with units marked right if within an adjustable % error of the correct value. &nbsp;Marked close if within a wider margin of error.</p>", "help_url": "", "input_widget": "string", "input_options": {"correctAnswer": "siground(settings['correctAnswer'],4)", "hint": {"static": true, "value": ""}, "allowEmpty": {"static": true, "value": true}}, "can_be_gap": true, "can_be_step": true, "marking_script": "mark:\nswitch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)\n\ninterpreted_answer:\nqty(student_scalar, student_units)\n\n\n\ncorrect_quantity:\nsettings[\"correctAnswer\"]\n\n\n\ncorrect_units:\nunits(correct_quantity)\n\n\nallowed_notation_styles:\n[\"plain\",\"en\"]\n\nmatch_student_number:\nmatchnumber(studentAnswer,allowed_notation_styles)\n\nstudent_scalar:\nmatch_student_number[1]\n\nstudent_units:\nreplace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")\n\ngood_units:\ntry(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n\n\nstudent_quantity:\nswitch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n\n\npercent_error:\ntry(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   \n\nright:\npercent_error <= settings['right']\n\n\nclose:\nright_sign and percent_error <= settings['close']\n\nright_sign:\nsign(student_scalar) = sign(correct_quantity)", "marking_notes": [{"name": "mark", "description": "This is the main marking note. It should award credit and provide feedback based on the student's answer.", "definition": "switch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)"}, {"name": "interpreted_answer", "description": "A value representing the student's answer to this part.", "definition": "qty(student_scalar, student_units)\n\n"}, {"name": "correct_quantity", "description": "", "definition": "settings[\"correctAnswer\"]\n\n"}, {"name": "correct_units", "description": "", "definition": "units(correct_quantity)\n"}, {"name": "allowed_notation_styles", "description": "", "definition": "[\"plain\",\"en\"]"}, {"name": "match_student_number", "description": "", "definition": "matchnumber(studentAnswer,allowed_notation_styles)"}, {"name": "student_scalar", "description": "", "definition": "match_student_number[1]"}, {"name": "student_units", "description": "<p>Modify&nbsp;the&nbsp;unit portion of&nbsp;the student's answer by</p>\n<p>1. replacing \"ohms\" with \"ohm\" &nbsp;case insensitive</p>\n<p>2.&nbsp;replacing '-'&nbsp;with ' '&nbsp;</p>\n<p>3. replacing '&deg;' with ' deg'&nbsp;</p>\n<p>to allow answers like 10 ft-lb and 30&deg;</p>", "definition": "replace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")"}, {"name": "good_units", "description": "", "definition": "try(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n"}, {"name": "student_quantity", "description": "<p>This fixes the student answer for two common errors. &nbsp;</p>\n<p>If student_units are wrong&nbsp;&nbsp;- replace with correct units</p>\n<p>If student_scalar has the wrong sign - replace with right sign</p>\n<p>If student makes both errors, only one gets fixed.</p>", "definition": "switch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n"}, {"name": "percent_error", "description": "", "definition": "try(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   "}, {"name": "right", "description": "", "definition": "percent_error <= settings['right']\n"}, {"name": "close", "description": "<p>Only marked close if the student actually has the right sign.</p>", "definition": "right_sign and percent_error <= settings['close']"}, {"name": "right_sign", "description": "", "definition": "sign(student_scalar) = sign(correct_quantity) "}], "settings": [{"name": "correctAnswer", "label": "Correct Quantity.", "help_url": "", "hint": "The correct answer given as a JME quantity.", "input_type": "code", "default_value": "", "evaluate": true}, {"name": "right", "label": "% Accuracy for right.", "help_url": "", "hint": "Question will be considered correct if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "0.2", "evaluate": true}, {"name": "close", "label": "% Accuracy for close.", "help_url": "", "hint": "Question will be considered close if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "1.0", "evaluate": true}, {"name": "C1", "label": "Close with units.", "help_url": "", "hint": "Partial Credit for close value with appropriate units. &nbsp;if correct answer is 100 N and close is &plusmn;1%,<br />99 &nbsp;N is accepted.", "input_type": "percent", "default_value": "75"}, {"name": "C2", "label": "No units or wrong sign", "help_url": "", "hint": "Partial credit for forgetting units or using wrong sign.<br />If the correct answer is 100 N, both 100 and -100 N are accepted.", "input_type": "percent", "default_value": "50"}, {"name": "C3", "label": "Close, no units.", "help_url": "", "hint": "Partial Credit for close value but forgotten units.<br />This value would be close&nbsp;if the expected units were provided. &nbsp;If the correct answer is 100 N, and close is &plusmn;1%,<br />99 is accepted.", "input_type": "percent", "default_value": "25"}], "public_availability": "always", "published": true, "extensions": ["quantities"]}], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Moment of Inertia: rectangles", "tags": ["Area moment of inertia", "area moment of inertia", "Mechanics", "mechanics", "moment of inertia", "second moment of area", "Statics", "statics"], "metadata": {"description": "<p>Find moment of inertia wrt the <em>x</em>- and <em>y</em>- axes for a shape made up of two rectangles.</p>", "licence": "Creative Commons Attribution-NonCommercial 4.0 International"}, "statement": "<p>The formula $I = \\dfrac{b h^3}{3}$ gives the moment of inertia of a rectangle about an axis passing through its base. &nbsp;</p>\n<p>Use this formula to find the moment of inertia of the shape below about both the $x$- and $y$- axes. &nbsp;Grid units are [{unit}].&nbsp;</p>\n<table style=\"margin-left: auto; margin-right: auto; width: 364px;\" height=\"41\">\n<tbody>\n<tr>\n<td style=\"text-align: center; width: 347px;\">{geogebra_applet('b3fxtwx6',[['A',A],['B',B]])}</td>\n</tr>\n</tbody>\n</table>\n<p style=\"text-align: center;\"></p>", "advice": "<p>To find $I_x$ consider the shape to be made of&nbsp;a&nbsp;&nbsp;$ (\\var{b_1}&nbsp;\\times&nbsp;\\var{h_1})$ rectangle and a&nbsp;$(\\var{b_2}&nbsp;\\times&nbsp;\\var{h_2})$&nbsp;{if(removed_x,'removed','')} rectangle.</p>\n<p>$&nbsp;\\quad I_x = \\dfrac{&nbsp;(\\var{b_1}) (\\var{h_1})^3}{3} \\, &nbsp;\\var{if(removed_x, '&ndash;' , '+')} \\,\\dfrac{&nbsp;(\\var{b_2}) (\\var{h_2})^3}{3} = \\var{siground(ix units, 4)} \\\\$</p>\n<p>To find $I_y$ consider the shape to be made of&nbsp;a&nbsp;&nbsp;$ (\\var{b_3}&nbsp;\\times&nbsp;\\var{h_3})$ rectangle and a&nbsp;$(\\var{b_4}&nbsp;\\times&nbsp;\\var{h_4})$ {if(removed_y,'removed','')} rectangle.</p>\n<p>$\\quad I_y = \\dfrac{&nbsp;(\\var{b_3}) (\\var{h_3})^3}{3} \\, &nbsp;\\var{if(removed_y, '&ndash;' , '+')} \\,\\dfrac{&nbsp;(\\var{b_4}) (\\var{h_4})^3}{3}= \\var{siground(Iy units,4)}$</p>", "rulesets": {}, "extensions": ["geogebra", "quantities"], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"Iy": {"name": "Iy", "group": "Unnamed group", "definition": "MOI(b_3,h_3) + if(removed_y,-1,1) MOI(b_4,h_4)", "description": "", "templateType": "anything", "can_override": false}, "removed_x": {"name": "removed_x", "group": "Unnamed group", "definition": "not(add)", "description": "<p>remove</p>", "templateType": "anything", "can_override": false}, "add": {"name": "add", "group": "Ungrouped variables", "definition": "random(true,false)", "description": "", "templateType": "anything", "can_override": false}, "h_2": {"name": "h_2", "group": "Unnamed group", "definition": "B[1]", "description": "", "templateType": "anything", "can_override": false}, "b_4": {"name": "b_4", "group": "Unnamed group", "definition": "abs(b[1]-A[1])", "description": "", "templateType": "anything", "can_override": false}, "removed_y": {"name": "removed_y", "group": "Unnamed group", "definition": "add and (B[1]>A[1])", "description": "", "templateType": "anything", "can_override": false}, "h_4": {"name": "h_4", "group": "Unnamed group", "definition": "A[0]", "description": "", "templateType": "anything", "can_override": false}, "A": {"name": "A", "group": "Ungrouped variables", "definition": "vector(random(1..10),random(1..10))", "description": "", "templateType": "anything", "can_override": false}, "h_1": {"name": "h_1", "group": "Unnamed group", "definition": "A[1]", "description": "", "templateType": "anything", "can_override": false}, "unit": {"name": "unit", "group": "Ungrouped variables", "definition": "random('in', 'cm', 'ft', 'm')", "description": "", "templateType": "anything", "can_override": false}, "Ix": {"name": "Ix", "group": "Unnamed group", "definition": "MOI(b_1,h_1) + if(removed_x,-1,1) MOI(b_2,h_2)", "description": "", "templateType": "anything", "can_override": false}, "b_3": {"name": "b_3", "group": "Unnamed group", "definition": "B[1]", "description": "", "templateType": "anything", "can_override": false}, "units": {"name": "units", "group": "Ungrouped variables", "definition": "qty(1,unit+'^4')", "description": "", "templateType": "anything", "can_override": false}, "b_1": {"name": "b_1", "group": "Unnamed group", "definition": "A[0]", "description": "", "templateType": "anything", "can_override": false}, "h_3": {"name": "h_3", "group": "Unnamed group", "definition": "B[0]", "description": "", "templateType": "anything", "can_override": false}, "B": {"name": "B", "group": "Ungrouped variables", "definition": "vector(random(1..10),random(1..10))", "description": "", "templateType": "anything", "can_override": false}, "b_2": {"name": "b_2", "group": "Unnamed group", "definition": "abs(B[0]-A[0])", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "if(add, B[0]>A[0], B[0]<A[0] and B[1] < A[1]) and A[1] <> B[1]", "maxRuns": 100}, "ungrouped_variables": ["A", "B", "add", "unit", "units"], "variable_groups": [{"name": "Unnamed group", "variables": ["b_1", "h_1", "b_2", "h_2", "b_3", "h_3", "b_4", "h_4", "Ix", "removed_x", "removed_y", "Iy"]}], "functions": {"MOI": {"parameters": [["b", "number"], ["h", "number"]], "type": "number", "language": "jme", "definition": "b * h * h * h /3"}}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p style=\"padding-left: 40px;\">$I_x = $ [[0]]</p>\n<p style=\"padding-left: 40px;\">$I_y = $ [[1]]</p>", "gaps": [{"type": "engineering-answer", "useCustomName": true, "customName": "$I_x$", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "Ix units", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "$I_y$", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "Iy units", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}], "sortAnswers": false}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}], "resources": []}]}], "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}]}