// Numbas version: finer_feedback_settings
{"name": "Machine: Vice grips", "extensions": ["geogebra", "quantities"], "custom_part_types": [], "resources": [["question-resources/vice-grips-fbd.png", "vice-grips-fbd.png"], ["question-resources/vicegrips2.vwx", "vicegrips2_MKKeCV0.vwx"]], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Machine: Vice grips", "tags": ["Equilibrium", "equilibrium", "machine", "Mechanics", "mechanics", "Statics", "statics"], "metadata": {"description": "Find the mechanical advantage of a pair of 'vice-grip' pliers.", "licence": "Creative Commons Attribution-NonCommercial 4.0 International"}, "statement": "<p>Determine the compressive force&nbsp;on object&nbsp;$A$ in terms of squeezing forces&nbsp;$P$ when the vice grip pliers are in the position shown. (Assume that the forces acting on $A$ are vertical.)</p>\n<table style=\"margin-left: auto; margin-right: auto;\">\n<tbody>\n<tr>\n<td>{applet}</td>\n</tr>\n</tbody>\n</table>\n<p style=\"text-align: center;\"></p>", "advice": "<p>Start by drawing neat, correct, labeled free body diagrams for the lower handle and the lower jaw. &nbsp;Recognize&nbsp;that piece&nbsp;$CE$ is a two force member in compression so the force on&nbsp;pin $E$&nbsp;acts down and left&nbsp;along a line passing through $C$ and $E$.</p>\n<p><img src=\"resources/question-resources/vice-grips-fbd.png\" width=\"565\" height=\"208\" style=\"display: block; margin-left: auto; margin-right: auto;\"/></p>\n<p><strong></strong><strong>Geometry</strong></p>\n<p>Determine angle $\\theta$ from the geometry of the pliers. &nbsp;</p>\n<p>\\[\\theta = \\tan^{-1}\\left( \\dfrac{\\var{q(x3)}}{\\var{q(y1)}}\\right) = \\var{d(theta)}&deg;\\]</p>\n<p><strong>FBD II</strong></p>\n<p>\\[\\begin{align}\\text{II: } \\Sigma M_B &amp;= 0\\\\ CE_x \\,(\\var{q(y2)}) -CE_y\\,(\\var{q(x2)}) + P \\,( \\var{q(x2+x3+x4)})&amp; = 0 \\\\ CE\\left( \\var{y2} \\sin \\theta -&nbsp;\\var{x2} \\cos \\theta \\right) &amp;= &nbsp;-P \\,( \\var{q(x2)}+\\var{q(x3)}+\\var{q(x4)})\\\\CE &amp;=-P \\,\\left( \\dfrac{\\var{q(x2+x3+x4)}}{\\var{-q(dperp)}} \\right)\\\\ &amp;= \\var{d(ce)} P \\\\ \\\\ \\text{II: }\\Sigma F_x &amp;= 0\\\\ B_x &amp;= CE_x\\\\ &amp;= CE \\sin{\\theta}\\\\ &amp;= (&nbsp;\\var{d(ce)} P&nbsp;)\\,(\\sin \\var{d(theta)}&deg;)\\\\&amp;= \\var{d(bx)} P\\end{align}\\]</p>\n<p><strong>FBD I</strong></p>\n<p>\\[\\begin{align} \\text{I: }\\Sigma M_D &amp;=0\\\\ A\\,(\\var{q(x1)}) &amp;= B_x\\,(\\var{q(y1)} + \\var{q(y2)})\\\\ A &amp;= (\\var{d(bx)} P) \\,&nbsp;\\left(\\dfrac{\\var{q(y1+y2)}}{\\var{q(x1)}}\\right)\\\\ &nbsp;&amp;= \\var{d(A)} P\\end{align}\\]</p>\n<p></p>", "rulesets": {}, "extensions": ["geogebra", "quantities"], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"dperp": {"name": "dperp", "group": "solution", "definition": "x2 cos(radians(theta)) - y2 sin(radians(theta))", "description": "<p>\\var{x2} \\cos \\theta - \\var{y2} \\sin \\theta</p>", "templateType": "anything", "can_override": false}, "Y2": {"name": "Y2", "group": "Ungrouped variables", "definition": "9 + random(-2..2)", "description": "", "templateType": "anything", "can_override": false}, "Y1": {"name": "Y1", "group": "Ungrouped variables", "definition": "24+random(-3..3)", "description": "", "templateType": "anything", "can_override": false}, "X1": {"name": "X1", "group": "Ungrouped variables", "definition": "43 + random(-5..5)\n", "description": "<p>x</p>", "templateType": "anything", "can_override": false}, "X3": {"name": "X3", "group": "Ungrouped variables", "definition": "61 + random(-6..6)", "description": "", "templateType": "anything", "can_override": false}, "X2": {"name": "X2", "group": "Ungrouped variables", "definition": "35 + random(-5..5)", "description": "", "templateType": "anything", "can_override": false}, "X4": {"name": "X4", "group": "Ungrouped variables", "definition": "20 + random(-3..3)", "description": "", "templateType": "anything", "can_override": false}, "CE": {"name": "CE", "group": "solution", "definition": "(x2+x3+x4)/dperp", "description": "<p>\\dfrac{\\var{q(x2+x3+x4)}}{\\var{q(dperp)}}</p>", "templateType": "anything", "can_override": false}, "Bx": {"name": "Bx", "group": "solution", "definition": "ce sin(radians(theta))", "description": "", "templateType": "anything", "can_override": false}, "params": {"name": "params", "group": "Ungrouped variables", "definition": "[['x1',x1],['x2',x2],['x3',x3],['x4',x4],['y1',y1],['y2',y2]]", "description": "", "templateType": "anything", "can_override": false}, "theta": {"name": "theta", "group": "solution", "definition": "degrees(arctan(x3/y1))", "description": "", "templateType": "anything", "can_override": false}, "A": {"name": "A", "group": "solution", "definition": "bx * (y1+y2)/x1", "description": "<p>&nbsp;(\\var{d(bx)} P) \\left(dfrac{\\var{q(y1+y2)}}{\\var{q(x1)}}</p>", "templateType": "anything", "can_override": false}, "applet": {"name": "applet", "group": "Ungrouped variables", "definition": "geogebra_applet('vugb9fzx', params)", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "dperp > 0.1 // dperp = 0 is degenerate", "maxRuns": 100}, "ungrouped_variables": ["X1", "X2", "X3", "X4", "Y1", "Y2", "params", "applet"], "variable_groups": [{"name": "solution", "variables": ["theta", "dperp", "CE", "Bx", "A"]}], "functions": {"q": {"parameters": [["n", "number"]], "type": "number", "language": "jme", "definition": "qty(d(n),'mm')"}, "d": {"parameters": [["n", "number"]], "type": "number", "language": "jme", "definition": "siground(n,4)"}}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p style=\"padding-left: 40px;\">$A$ =&nbsp;[[0]]&nbsp;</p>", "gaps": [{"type": "jme", "useCustomName": true, "customName": "A", "marks": "20", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "answer": "{siground(A,4)} P", "showPreview": true, "checkingType": "reldiff", "checkingAccuracy": "0.002", "failureRate": 1, "vsetRangePoints": 5, "vsetRange": ["0", "1"], "checkVariableNames": false, "singleLetterVariables": false, "allowUnknownFunctions": true, "implicitFunctionComposition": false, "caseSensitive": false, "valuegenerators": [{"name": "p", "value": ""}]}], "sortAnswers": false}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}], "resources": ["question-resources/vice-grips-fbd.png", "question-resources/vicegrips2.vwx"]}]}], "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}]}