// Numbas version: finer_feedback_settings
{"name": "Vector Addition: no diagram", "extensions": ["quantities"], "custom_part_types": [{"source": {"pk": 24, "author": {"name": "William Haynes", "pk": 2530}, "edit_page": "/part_type/24/edit"}, "name": "Angle quantity", "short_name": "angle-quantity-from-reference", "description": "<p>Angle&nbsp;as a quantity in degrees.</p>", "help_url": "", "input_widget": "string", "input_options": {"correctAnswer": "plain_string(settings['correct_quantity'])", "hint": {"static": true, "value": ""}, "allowEmpty": {"static": true, "value": false}}, "can_be_gap": true, "can_be_step": true, "marking_script": "mark:\nswitch( \nright and good_units and right_sign and angle_in_range, add_credit(1.0,'Correct.'),\nright and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\nright and right_sign and not good_units, add_credit(settings['C2'],'Correct angle, but missing degree symbol.'),\nright and good_units and right_sign and not angle_in_range,add_credit(settings['C1'],'Angle is out of range.'),\nclose and good_units, add_credit(settings['C1'],'Close.'),\nclose and not good_units, add_credit(settings['C3'],'Answer is close, but missing degree symbol.'),\nincorrect('Wrong answer.')\n)\n\ninterpreted_answer:\nqty(student_scalar, student_units)\n\ncorrect_scalar:\nscalar(correct_quantity)\n  \n\ncorrect_quantity:\nsettings['correct_quantity']\n\ncorrect_units:\nunits(correct_quantity)\n\nallowed_notation_styles:\n[\"plain\",\"en\"]\n\nmatch_student_number:\nmatchnumber(studentAnswer,allowed_notation_styles)\n\nstudent_scalar:\nmatch_student_number[1]\n\nstudent_units:\njoin(\nsplit(studentAnswer[len(match_student_number[0])..len(studentAnswer)]\n,\"\u00b0\"),\" deg\")\n\n\n\ngood_units:\ntry(\nkind(quantity(1, student_units))= kind(correct_quantity),\nmsg,\nfeedback(msg);false)\n\nstudent_quantity:\nswitch(not good_units, \nstudent_scalar * correct_units, \nnot right_sign,\n-quantity(student_scalar, student_units),\nquantity(student_scalar,student_units)\n)\n\nright_sign:\nsign(student_scalar) = sign(correct_quantity)\n\nangle_in_range:\nif(settings['restrict_angle'], abs(student_scalar) <= 90, true)\n\nright:\nwithinTolerance(abs(student_scalar), abs(correct_scalar), settings['right'])\n\nclose:\nwithinTolerance(student_scalar, correct_scalar, settings['close'])", "marking_notes": [{"name": "mark", "description": "This is the main marking note. It should award credit and provide feedback based on the student's answer.", "definition": "switch( \nright and good_units and right_sign and angle_in_range, add_credit(1.0,'Correct.'),\nright and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\nright and right_sign and not good_units, add_credit(settings['C2'],'Correct angle, but missing degree symbol.'),\nright and good_units and right_sign and not angle_in_range,add_credit(settings['C1'],'Angle is out of range.'),\nclose and good_units, add_credit(settings['C1'],'Close.'),\nclose and not good_units, add_credit(settings['C3'],'Answer is close, but missing degree symbol.'),\nincorrect('Wrong answer.')\n)"}, {"name": "interpreted_answer", "description": "A value representing the student's answer to this part.", "definition": "qty(student_scalar, student_units)"}, {"name": "correct_scalar", "description": "", "definition": "scalar(correct_quantity)\n  "}, {"name": "correct_quantity", "description": "", "definition": "settings['correct_quantity']"}, {"name": "correct_units", "description": "", "definition": "units(correct_quantity)"}, {"name": "allowed_notation_styles", "description": "", "definition": "[\"plain\",\"en\"]"}, {"name": "match_student_number", "description": "", "definition": "matchnumber(studentAnswer,allowed_notation_styles)"}, {"name": "student_scalar", "description": "", "definition": "match_student_number[1]"}, {"name": "student_units", "description": "", "definition": "join(\nsplit(studentAnswer[len(match_student_number[0])..len(studentAnswer)]\n,\"\u00b0\"),\" deg\")\n\n"}, {"name": "good_units", "description": "", "definition": "try(\nkind(quantity(1, student_units))= kind(correct_quantity),\nmsg,\nfeedback(msg);false)"}, {"name": "student_quantity", "description": "", "definition": "switch(not good_units, \nstudent_scalar * correct_units, \nnot right_sign,\n-quantity(student_scalar, student_units),\nquantity(student_scalar,student_units)\n)"}, {"name": "right_sign", "description": "", "definition": "sign(student_scalar) = sign(correct_quantity)"}, {"name": "angle_in_range", "description": "", "definition": "if(settings['restrict_angle'], abs(student_scalar) <= 90, true)"}, {"name": "right", "description": "<p>Will check for correct sign elswhere.</p>", "definition": "withinTolerance(abs(student_scalar), abs(correct_scalar), settings['right'])"}, {"name": "close", "description": "<p>Must have correct sign to be close.</p>", "definition": "withinTolerance(student_scalar, correct_scalar, settings['close'])\n"}], "settings": [{"name": "correct_quantity", "label": "Correct Angle as quantity ", "help_url": "", "hint": "", "input_type": "code", "default_value": "qty(45,'deg')", "evaluate": true}, {"name": "right", "label": "Accuracy for right.", "help_url": "", "hint": "Question will be considered correct if the scalar part of the student's answer is within &plusmn; this amount from the correct value.", "input_type": "code", "default_value": "0.1", "evaluate": true}, {"name": "restrict_angle", "label": "Less than 90\u00b0", "help_url": "", "hint": "When checked, angle must be between -90&deg; and +90&deg;.", "input_type": "checkbox", "default_value": true}, {"name": "C1", "label": "Close with units.", "help_url": "", "hint": "Partial Credit for close value with appropriate units.", "input_type": "percent", "default_value": "75"}, {"name": "close", "label": " Accuracy for close.", "help_url": "", "hint": "Question will be considered close if the scalar part of the student's answer is within &plusmn; this amount from the correct value.", "input_type": "code", "default_value": "0.5", "evaluate": true}, {"name": "C2", "label": "No units or wrong sign", "help_url": "", "hint": "Partial credit for forgetting units or using wrong sign.", "input_type": "percent", "default_value": "50"}, {"name": "C3", "label": "Close, no units.", "help_url": "", "hint": "Partial Credit for close value without units.", "input_type": "percent", "default_value": "25"}], "public_availability": "restricted", "published": false, "extensions": ["quantities"]}, {"source": {"pk": 19, "author": {"name": "William Haynes", "pk": 2530}, "edit_page": "/part_type/19/edit"}, "name": "Engineering Accuracy with units", "short_name": "engineering-answer", "description": "<p>A value with units marked right if within an adjustable % error of the correct value. &nbsp;Marked close if within a wider margin of error.</p>", "help_url": "", "input_widget": "string", "input_options": {"correctAnswer": "siground(settings['correctAnswer'],4)", "hint": {"static": true, "value": ""}, "allowEmpty": {"static": true, "value": true}}, "can_be_gap": true, "can_be_step": true, "marking_script": "mark:\nswitch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)\n\ninterpreted_answer:\nqty(student_scalar, student_units)\n\n\n\ncorrect_quantity:\nsettings[\"correctAnswer\"]\n\n\n\ncorrect_units:\nunits(correct_quantity)\n\n\nallowed_notation_styles:\n[\"plain\",\"en\"]\n\nmatch_student_number:\nmatchnumber(studentAnswer,allowed_notation_styles)\n\nstudent_scalar:\nmatch_student_number[1]\n\nstudent_units:\nreplace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")\n\ngood_units:\ntry(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n\n\nstudent_quantity:\nswitch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n\n\npercent_error:\ntry(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   \n\nright:\npercent_error <= settings['right']\n\n\nclose:\nright_sign and percent_error <= settings['close']\n\nright_sign:\nsign(student_scalar) = sign(correct_quantity)", "marking_notes": [{"name": "mark", "description": "This is the main marking note. It should award credit and provide feedback based on the student's answer.", "definition": "switch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)"}, {"name": "interpreted_answer", "description": "A value representing the student's answer to this part.", "definition": "qty(student_scalar, student_units)\n\n"}, {"name": "correct_quantity", "description": "", "definition": "settings[\"correctAnswer\"]\n\n"}, {"name": "correct_units", "description": "", "definition": "units(correct_quantity)\n"}, {"name": "allowed_notation_styles", "description": "", "definition": "[\"plain\",\"en\"]"}, {"name": "match_student_number", "description": "", "definition": "matchnumber(studentAnswer,allowed_notation_styles)"}, {"name": "student_scalar", "description": "", "definition": "match_student_number[1]"}, {"name": "student_units", "description": "<p>Modify&nbsp;the&nbsp;unit portion of&nbsp;the student's answer by</p>\n<p>1. replacing \"ohms\" with \"ohm\" &nbsp;case insensitive</p>\n<p>2.&nbsp;replacing '-'&nbsp;with ' '&nbsp;</p>\n<p>3. replacing '&deg;' with ' deg'&nbsp;</p>\n<p>to allow answers like 10 ft-lb and 30&deg;</p>", "definition": "replace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")"}, {"name": "good_units", "description": "", "definition": "try(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n"}, {"name": "student_quantity", "description": "<p>This fixes the student answer for two common errors. &nbsp;</p>\n<p>If student_units are wrong&nbsp;&nbsp;- replace with correct units</p>\n<p>If student_scalar has the wrong sign - replace with right sign</p>\n<p>If student makes both errors, only one gets fixed.</p>", "definition": "switch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n"}, {"name": "percent_error", "description": "", "definition": "try(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   "}, {"name": "right", "description": "", "definition": "percent_error <= settings['right']\n"}, {"name": "close", "description": "<p>Only marked close if the student actually has the right sign.</p>", "definition": "right_sign and percent_error <= settings['close']"}, {"name": "right_sign", "description": "", "definition": "sign(student_scalar) = sign(correct_quantity) "}], "settings": [{"name": "correctAnswer", "label": "Correct Quantity.", "help_url": "", "hint": "The correct answer given as a JME quantity.", "input_type": "code", "default_value": "", "evaluate": true}, {"name": "right", "label": "% Accuracy for right.", "help_url": "", "hint": "Question will be considered correct if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "0.2", "evaluate": true}, {"name": "close", "label": "% Accuracy for close.", "help_url": "", "hint": "Question will be considered close if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "1.0", "evaluate": true}, {"name": "C1", "label": "Close with units.", "help_url": "", "hint": "Partial Credit for close value with appropriate units. &nbsp;if correct answer is 100 N and close is &plusmn;1%,<br />99 &nbsp;N is accepted.", "input_type": "percent", "default_value": "75"}, {"name": "C2", "label": "No units or wrong sign", "help_url": "", "hint": "Partial credit for forgetting units or using wrong sign.<br />If the correct answer is 100 N, both 100 and -100 N are accepted.", "input_type": "percent", "default_value": "50"}, {"name": "C3", "label": "Close, no units.", "help_url": "", "hint": "Partial Credit for close value but forgotten units.<br />This value would be close&nbsp;if the expected units were provided. &nbsp;If the correct answer is 100 N, and close is &plusmn;1%,<br />99 is accepted.", "input_type": "percent", "default_value": "25"}], "public_availability": "always", "published": true, "extensions": ["quantities"]}], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Vector Addition: no diagram", "tags": ["Mechanics", "mechanics", "Statics", "statics", "Vector addition", "vector addition", "Vector Addition"], "metadata": {"description": "<p>Sum three force vectors&nbsp;based on a written description of the situation.</p>", "licence": "Creative Commons Attribution-ShareAlike 4.0 International"}, "statement": "<p><strong>Given: Three&nbsp;forces act&nbsp;on particle $P$.</strong></p>\n<ul>\n<li>Force $\\textbf{A}$&nbsp;has a magnitude of {qty(A,unit)} and acts {theta}&deg; CCW from the {dirA[0]}.&nbsp;</li>\n<li>Force $\\textbf{B}$ has a magnitude of {qty(B,unit)} and acts along a ray from point $P$&nbsp;with a slope of $\\dfrac{\\var{rise}}{\\var{run}}$.&nbsp;</li>\n<li>Force $\\textbf{C}$ has a magnitude of {qty(C,unit)} and acts straight {dirC[0]}.&nbsp;</li>\n</ul>\n<ol>\n<li>Find the scalar components of $\\textbf{A}, \\textbf{B}$, and $\\textbf{C}$.</li>\n<li>Draw a neat, labeled, scaled digram representing the tip-to-tail vector addition $\\textbf{R} = &nbsp;\\textbf{A} + \\textbf{B} + \\textbf{C}$.&nbsp;</li>\n<li>Find the scalar components of the resultant.</li>\n<li>Find the magnitude of the resultant $\\textbf{R}$ and the angle it makes with the positive $x$-axis.</li>\n</ol>\n<p style=\"padding-left: 30px;\"></p>", "advice": "<div style=\"padding-left: 100px;\">\n<p><strong>Given</strong></p>\n<p style=\"padding-left: 30px;\">$\\mathbf{A}$ = {A} {unit} &nbsp;at {alpha}&deg; &nbsp;</p>\n<p style=\"padding-left: 30px;\">$\\mathbf{B}$ = {B}&nbsp;{unit} at {precround(beta,2)}&deg;</p>\n<p style=\"padding-left: 30px;\">$\\mathbf{C}$ = {C}&nbsp;{unit} at {gamma}&deg;</p>\n<p><strong>Find Scalar Components</strong></p>\n<table style=\"padding-left: 30px' width: 352.328125px; float: left;\" border=\"1px solid black\">\n<tbody>\n<tr style=\"height: 17px;\">\n<td style=\"height: 17px; text-align: center; padding-left: 30px; width: 35px;\"></td>\n<td style=\"height: 17px; text-align: center; width: 123px;\">$F_x$</td>\n<td style=\"height: 17px; text-align: center; width: 147.328125px;\">$F_y$</td>\n</tr>\n<tr style=\"height: 18px;\">\n<td style=\"height: 18px; text-align: right; width: 35px;\">$\\mathbf{A}$</td>\n<td style=\"height: 18px; text-align: center; width: 123px;\">{precround(Ax,2)}&nbsp;{unit}&nbsp;</td>\n<td style=\"height: 18px; text-align: center; width: 147.328125px;\">{precround(Ay,2)} {unit}</td>\n</tr>\n<tr style=\"height: 17px;\">\n<td style=\"height: 17px; text-align: right; width: 35px;\">$\\mathbf{B}$</td>\n<td style=\"height: 17px; text-align: center; width: 123px;\">{precround(Bx,2)} {unit}</td>\n<td style=\"height: 17px; text-align: center; width: 147.328125px;\">{precround(By,2)} {unit}</td>\n</tr>\n<tr style=\"height: 17px;\">\n<td style=\"height: 17px; text-align: right; width: 35px;\">$\\mathbf{C}$</td>\n<td style=\"height: 17px; text-align: center; width: 123px;\">{precround(Cx,2)} {unit}<span class=\"Apple-tab-span\"> </span></td>\n<td style=\"height: 17px; text-align: center; width: 147.328125px;\">{precround(Cy,2)} {unit}</td>\n</tr>\n</tbody>\n</table>\n<p style=\"clear: both;\"><strong>Find rectangular components of the Resultant</strong></p>\n<p style=\"padding-left: 30px;\">$R_x = A_x + B_x +C_x = \\var{precround(Rx,2)}$&nbsp;{unit}</p>\n<p style=\"padding-left: 30px;\">$R_y = A_y + B_y +C_y = \\var{precround(Ry,2)}$&nbsp;{unit}</p>\n<p><strong>Find the Magnitude of the Resultant</strong></p>\n<p style=\"padding-left: 30px;\">$R = &nbsp;\\sqrt{{R_x}^2 + {R_y}^2} = \\var{precround(R,2)}$ &nbsp;{unit}</p>\n<p><strong>Find the Direction of the Resultant</strong></p>\n<p style=\"padding-left: 30px;\">The angle in a right triangle with sides of length $R_x$ and $R_y$ is</p>\n<p style=\"padding-left: 30px;\">$\\alpha = \\arctan{\\left|\\dfrac{R_y}{R_x}\\right|} = \\var{siground(degrees(arctan(abs(Ry/Rx))),4)}$&deg;</p>\n<p style=\"padding-left: 30px;\">Since $R_x = \\var{siground(qty(Rx,unit),3)}$ and $R_y = \\var{siground(qty(Ry,unit),3)}$ force $\\textbf{R}$ points {if(Rx &gt; 0, \"right\", \"left\")} and {if(Ry &gt; 0, \"up\", \"down\")}, therefore the angle between $\\textbf{R}$ and the positive $x$ axis is:</p>\n<p style=\"padding-left: 30px;\">$\\theta = &nbsp;\\var{(if(Rx&gt;0,if(Ry &gt; 0, \"$\\\\alpha$\",\"$-\\\\alpha$\"),if(Ry &gt; 0, \"$180&deg;-\\\\alpha$\", \"$-180&deg; + \\\\alpha$\")))} &nbsp;&nbsp; = \\var{siground(theta_r,4)} $&deg;</p>\n</div>", "rulesets": {}, "extensions": ["quantities"], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"A": {"name": "A", "group": "Ungrouped variables", "definition": "random(10..60#5)\n\n", "description": "<p>Magnitude of force A</p>\n<p></p>", "templateType": "anything", "can_override": false}, "B": {"name": "B", "group": "Ungrouped variables", "definition": "random(30..80#10)", "description": "", "templateType": "anything", "can_override": false}, "unit": {"name": "unit", "group": "Ungrouped variables", "definition": "random(['N','lb'])", "description": "", "templateType": "anything", "can_override": false}, "C": {"name": "C", "group": "Ungrouped variables", "definition": "random(10..60#5)", "description": "", "templateType": "anything", "can_override": false}, "dirC": {"name": "dirC", "group": "Ungrouped variables", "definition": "random(['left',180],['right',0],['up',90],['down',270])", "description": "<p>Description and angle for direction of B</p>", "templateType": "anything", "can_override": false}, "rise": {"name": "rise", "group": "Ungrouped variables", "definition": "random([-1,-3,-4,1,3,4])", "description": "", "templateType": "anything", "can_override": false}, "run": {"name": "run", "group": "Ungrouped variables", "definition": "random([-1,-3,-4,1,3,4])", "description": "", "templateType": "anything", "can_override": false}, "theta": {"name": "theta", "group": "Ungrouped variables", "definition": "random([30,60])", "description": "", "templateType": "anything", "can_override": false}, "dirA": {"name": "dirA", "group": "Ungrouped variables", "definition": "random([\n    [\"positive $x$ axis\",0+theta],\n     [\"positive $y$ axis\",90+theta],\n     [\"negative $x$ axis\",180+theta],\n     [\"negative $y$ axis\",270+theta]])\n    \n    \n    ", "description": "<p>Description and angle for direction of&nbsp;A</p>", "templateType": "anything", "can_override": false}, "alpha": {"name": "alpha", "group": "Ungrouped variables", "definition": "dirA[1]", "description": "", "templateType": "anything", "can_override": false}, "beta": {"name": "beta", "group": "Ungrouped variables", "definition": "degrees(atan2(rise,run))", "description": "", "templateType": "anything", "can_override": false}, "gamma": {"name": "gamma", "group": "Ungrouped variables", "definition": "dirC[1]", "description": "", "templateType": "anything", "can_override": false}, "ax": {"name": "ax", "group": "answers", "definition": "A cos(radians(alpha))", "description": "", "templateType": "anything", "can_override": false}, "bx": {"name": "bx", "group": "answers", "definition": "B cos(radians(beta))", "description": "", "templateType": "anything", "can_override": false}, "cx": {"name": "cx", "group": "answers", "definition": "precround(C cos(radians(gamma)),6)", "description": "", "templateType": "anything", "can_override": false}, "ay": {"name": "ay", "group": "answers", "definition": "A sin(radians(alpha))", "description": "", "templateType": "anything", "can_override": false}, "by": {"name": "by", "group": "answers", "definition": "B sin(radians(beta))", "description": "", "templateType": "anything", "can_override": false}, "cy": {"name": "cy", "group": "answers", "definition": "precround(C sin(radians(gamma)),5)", "description": "<p>precround to avoid precision errors &nbsp;0 != 0.00000000000000213</p>", "templateType": "anything", "can_override": false}, "rx": {"name": "rx", "group": "answers", "definition": "ax+bx+cx", "description": "", "templateType": "anything", "can_override": false}, "ry": {"name": "ry", "group": "answers", "definition": "ay+by+cy", "description": "", "templateType": "anything", "can_override": false}, "R": {"name": "R", "group": "answers", "definition": "sqrt(rx^2+ry^2)", "description": "", "templateType": "anything", "can_override": false}, "theta_R": {"name": "theta_R", "group": "answers", "definition": "degrees(atan2(ry,rx))", "description": "", "templateType": "anything", "can_override": false}, "debug": {"name": "debug", "group": "Ungrouped variables", "definition": "false", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "abs(rise) <> abs(run) and\nRx <>0 and\nRy <>0", "maxRuns": 100}, "ungrouped_variables": ["A", "B", "unit", "C", "rise", "run", "theta", "dirA", "dirC", "alpha", "beta", "gamma", "debug"], "variable_groups": [{"name": "answers", "variables": ["ax", "bx", "cx", "ay", "by", "cy", "rx", "ry", "R", "theta_R"]}], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": true, "customName": "Rectangular Components", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<table style=\"float: left;\">\n<tbody>\n<tr>\n<td style=\"text-align: center;\"><strong>Force</strong></td>\n<td style=\"text-align: center;\">$x$</td>\n<td style=\"text-align: center;\">$y$</td>\n</tr>\n<tr>\n<td style=\"text-align: center;\">$\\textbf{A}$</td>\n<td>[[0]]</td>\n<td style=\"text-align: left;\">[[1]]</td>\n</tr>\n<tr>\n<td style=\"text-align: center;\">$\\textbf{B}$</td>\n<td>[[2]]</td>\n<td>[[3]]</td>\n</tr>\n<tr>\n<td style=\"text-align: center;\">$\\textbf{C}$</td>\n<td>[[4]]</td>\n<td>[[5]]</td>\n</tr>\n</tbody>\n</table>\n<p></p>\n<p></p>\n<p></p>\n<p></p>\n<div data-jme-visible=\"debug\">\n<table>\n<tbody>\n<tr style=\"height: 17px;\">\n<td style=\"height: 17px;\"></td>\n<td style=\"height: 17px;\">x</td>\n<td style=\"height: 17px;\">y</td>\n</tr>\n<tr style=\"height: 18px;\">\n<td style=\"height: 18px;\">A</td>\n<td style=\"height: 18px;\">{precround(Ax,2)}&nbsp;{unit}&nbsp;</td>\n<td style=\"height: 18px;\">{precround(Ay,2)} {unit}</td>\n</tr>\n<tr style=\"height: 17px;\">\n<td style=\"height: 17px;\">B</td>\n<td style=\"height: 17px;\">{precround(Bx,2)} {unit}</td>\n<td style=\"height: 17px;\">{precround(By,2)} {unit}</td>\n</tr>\n<tr style=\"height: 17px;\">\n<td style=\"height: 17px;\">C&nbsp;</td>\n<td style=\"height: 17px;\">{precround(Cx,2)} {unit}<span class=\"Apple-tab-span\"> </span></td>\n<td style=\"height: 17px;\">{precround(Cy,2)} {unit}</td>\n</tr>\n<tr style=\"height: 17px;\">\n<td style=\"height: 17px;\">R</td>\n<td style=\"height: 17px;\">{precround(Rx,2)} {unit}</td>\n<td style=\"height: 17px;\">{precround(Ry,2)} {unit}</td>\n</tr>\n</tbody>\n</table>\n</div>", "gaps": [{"type": "engineering-answer", "useCustomName": true, "customName": "$A_x$", "marks": "3", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(Ax,unit)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "$A_y$", "marks": "3", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(Ay,unit)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "$B_x$", "marks": "3", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(Bx,unit)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "$B_y$", "marks": "3", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(By,unit)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "$C_x$", "marks": "3", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(Cx,unit)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "$C_y$", "marks": "3", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(Cy,unit)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}], "sortAnswers": false}, {"type": "gapfill", "useCustomName": true, "customName": "Vector Sum", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "alwaysreplace", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p style=\"padding-left: 40px;\">$R_x =$ [[0]]&nbsp;</p>\n<p style=\"padding-left: 40px;\">$R_y = $ [[1]]</p>\n<p style=\"padding-left: 40px;\">$R=$ [[2]] &nbsp;&nbsp;</p>\n<p style=\"padding-left: 40px;\">$\\theta =$&nbsp;[[3]]&nbsp;&nbsp;$(-180&deg; \\le \\theta \\le 180&deg;)$</p>", "gaps": [{"type": "engineering-answer", "useCustomName": true, "customName": "$R_x$", "marks": "3", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": false, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(Rx,unit)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "$R_y$", "marks": "3", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(Ry,unit)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "$R$", "marks": "3", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(R,unit)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "angle-quantity-from-reference", "useCustomName": true, "customName": "$\\theta$", "marks": "3", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correct_quantity": "precround(qty(theta_r,'deg'),2)", "right": "0.1", "restrict_angle": false, "C1": "75", "close": "0.5", "C2": "50", "C3": "25"}}], "sortAnswers": false}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}], "resources": []}]}], "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}]}