// Numbas version: finer_feedback_settings
{"name": "Equilibrium of a particle: two suspended loads", "extensions": ["geogebra", "quantities"], "custom_part_types": [{"source": {"pk": 19, "author": {"name": "William Haynes", "pk": 2530}, "edit_page": "/part_type/19/edit"}, "name": "Engineering Accuracy with units", "short_name": "engineering-answer", "description": "<p>A value with units marked right if within an adjustable % error of the correct value. &nbsp;Marked close if within a wider margin of error.</p>", "help_url": "", "input_widget": "string", "input_options": {"correctAnswer": "siground(settings['correctAnswer'],4)", "hint": {"static": true, "value": ""}, "allowEmpty": {"static": true, "value": true}}, "can_be_gap": true, "can_be_step": true, "marking_script": "mark:\nswitch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)\n\ninterpreted_answer:\nqty(student_scalar, student_units)\n\n\n\ncorrect_quantity:\nsettings[\"correctAnswer\"]\n\n\n\ncorrect_units:\nunits(correct_quantity)\n\n\nallowed_notation_styles:\n[\"plain\",\"en\"]\n\nmatch_student_number:\nmatchnumber(studentAnswer,allowed_notation_styles)\n\nstudent_scalar:\nmatch_student_number[1]\n\nstudent_units:\nreplace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")\n\ngood_units:\ntry(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n\n\nstudent_quantity:\nswitch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n\n\npercent_error:\ntry(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   \n\nright:\npercent_error <= settings['right']\n\n\nclose:\nright_sign and percent_error <= settings['close']\n\nright_sign:\nsign(student_scalar) = sign(correct_quantity)", "marking_notes": [{"name": "mark", "description": "This is the main marking note. It should award credit and provide feedback based on the student's answer.", "definition": "switch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)"}, {"name": "interpreted_answer", "description": "A value representing the student's answer to this part.", "definition": "qty(student_scalar, student_units)\n\n"}, {"name": "correct_quantity", "description": "", "definition": "settings[\"correctAnswer\"]\n\n"}, {"name": "correct_units", "description": "", "definition": "units(correct_quantity)\n"}, {"name": "allowed_notation_styles", "description": "", "definition": "[\"plain\",\"en\"]"}, {"name": "match_student_number", "description": "", "definition": "matchnumber(studentAnswer,allowed_notation_styles)"}, {"name": "student_scalar", "description": "", "definition": "match_student_number[1]"}, {"name": "student_units", "description": "<p>Modify&nbsp;the&nbsp;unit portion of&nbsp;the student's answer by</p>\n<p>1. replacing \"ohms\" with \"ohm\" &nbsp;case insensitive</p>\n<p>2.&nbsp;replacing '-'&nbsp;with ' '&nbsp;</p>\n<p>3. replacing '&deg;' with ' deg'&nbsp;</p>\n<p>to allow answers like 10 ft-lb and 30&deg;</p>", "definition": "replace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")"}, {"name": "good_units", "description": "", "definition": "try(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n"}, {"name": "student_quantity", "description": "<p>This fixes the student answer for two common errors. &nbsp;</p>\n<p>If student_units are wrong&nbsp;&nbsp;- replace with correct units</p>\n<p>If student_scalar has the wrong sign - replace with right sign</p>\n<p>If student makes both errors, only one gets fixed.</p>", "definition": "switch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n"}, {"name": "percent_error", "description": "", "definition": "try(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   "}, {"name": "right", "description": "", "definition": "percent_error <= settings['right']\n"}, {"name": "close", "description": "<p>Only marked close if the student actually has the right sign.</p>", "definition": "right_sign and percent_error <= settings['close']"}, {"name": "right_sign", "description": "", "definition": "sign(student_scalar) = sign(correct_quantity) "}], "settings": [{"name": "correctAnswer", "label": "Correct Quantity.", "help_url": "", "hint": "The correct answer given as a JME quantity.", "input_type": "code", "default_value": "", "evaluate": true}, {"name": "right", "label": "% Accuracy for right.", "help_url": "", "hint": "Question will be considered correct if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "0.2", "evaluate": true}, {"name": "close", "label": "% Accuracy for close.", "help_url": "", "hint": "Question will be considered close if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "1.0", "evaluate": true}, {"name": "C1", "label": "Close with units.", "help_url": "", "hint": "Partial Credit for close value with appropriate units. &nbsp;if correct answer is 100 N and close is &plusmn;1%,<br />99 &nbsp;N is accepted.", "input_type": "percent", "default_value": "75"}, {"name": "C2", "label": "No units or wrong sign", "help_url": "", "hint": "Partial credit for forgetting units or using wrong sign.<br />If the correct answer is 100 N, both 100 and -100 N are accepted.", "input_type": "percent", "default_value": "50"}, {"name": "C3", "label": "Close, no units.", "help_url": "", "hint": "Partial Credit for close value but forgotten units.<br />This value would be close&nbsp;if the expected units were provided. &nbsp;If the correct answer is 100 N, and close is &plusmn;1%,<br />99 is accepted.", "input_type": "percent", "default_value": "25"}], "public_availability": "always", "published": true, "extensions": ["quantities"]}], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Equilibrium of a particle: two suspended loads", "tags": ["2-d", "2-D", "equilibrium", "Equilibrium", "mechanics", "Mechanics", "particle", "Particle", "statics", "Statics"], "metadata": {"description": "<p>Find the tensions is a system of cables supporting two loads.</p>", "licence": "Creative Commons Attribution-NonCommercial 4.0 International"}, "statement": "<p>Two loads $W_1 = \\var{w1}$, and&nbsp;$W_2 = \\var{w2}$ are suspended from a system of cables as shown. Determine the tensions in the cables.</p>\n<table style=\"margin-left: auto; margin-right: auto;\">\n<tbody>\n<tr>\n<td>{geogebra_applet('vzqtfrbq',[['a1',a1+\"&deg;\"],['a2',a2+\"&deg;\"],['a3',a3+\"&deg;\"]])}</td>\n</tr>\n</tbody>\n</table>\n<p></p>", "advice": "<p>Begin by drawing a free body diagram of particle <em>A</em>, then apply the equations of equilibrium.</p>\n<p>$\\begin{align}\\\\ \\Sigma F_x = 0 &nbsp;\\\\-{T_B}_x + {T_C}_x = 0 \\end{align} \\\\ T_B \\cos(\\var{a1}&deg;) = T_C \\sin(\\var{90-a2}&deg;) \\qquad(1)$</p>\n<p>$\\begin{align}\\\\ \\Sigma F_y = 0 &nbsp;\\\\{T_B}_y + {T_C}_y&nbsp;- W_1 = 0 \\end{align} \\\\ T_B&nbsp;\\sin(\\var{a1}&deg;) +&nbsp;T_C \\cos(\\var{90-a2})&deg; = W_1 \\qquad(2)$</p>\n<p>Solve (1) and (2) simultaneously for $T_B$ and $T_C$.</p>\n<p>$T_B = \\var{siground(TB,4)}, \\quad T_C = \\var{siground(TC,4)}$</p>\n<hr/>\n<p>With $T_C$ known, use a similar approach on particle $C$ to find $T_D$ and $T_E$. &nbsp;</p>\n<p>Since cable<em> CD</em> is horizontal $T_D$&nbsp;has no <em>y</em>-component, so&nbsp;the $\\Sigma F_y = 0$ equation only has one unknown&nbsp;and&nbsp;simultaneous equations are unnecessary in this case.</p>\n<p>$T_D&nbsp;= \\var{siground(TD,4)}, \\quad T_E = \\var{siground(TE,4)}$</p>", "rulesets": {}, "extensions": ["geogebra", "quantities"], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"forceC": {"name": "forceC", "group": "check", "definition": "vector(cos(radians(a2)),sin(radians(a2))) scalar(TC)\n", "description": "", "templateType": "anything", "can_override": false}, "forceE": {"name": "forceE", "group": "check", "definition": "vector(cos(radians(a3)),sin(radians(a3))) scalar(TE)", "description": "", "templateType": "anything", "can_override": false}, "TB": {"name": "TB", "group": "magnitudes", "definition": "w1 cos(radians(a2))/sin(radians(a1+a2))", "description": "", "templateType": "anything", "can_override": false}, "forceD": {"name": "forceD", "group": "check", "definition": "vector(scalar(TD),0)", "description": "", "templateType": "anything", "can_override": false}, "a2": {"name": "a2", "group": "input", "definition": "random(10..70#5)", "description": "<p>a2 is the angle from the horizontal to AC. &nbsp;Note angle shown on diagram is the complement of a2</p>", "templateType": "anything", "can_override": false}, "debug": {"name": "debug", "group": "input", "definition": "false", "description": "", "templateType": "anything", "can_override": false}, "TE": {"name": "TE", "group": "magnitudes", "definition": "(W2 + TC sin(radians(A2)))/sin(radians(A3))", "description": "", "templateType": "anything", "can_override": false}, "a1": {"name": "a1", "group": "input", "definition": "random(10..80#5)", "description": "<p>angle of rope AB from horizontal, as shown on diagram.</p>", "templateType": "anything", "can_override": false}, "a3": {"name": "a3", "group": "input", "definition": "random(20..70#5)", "description": "<p>angle of CE as shown on diagram.</p>", "templateType": "anything", "can_override": false}, "w2": {"name": "w2", "group": "input", "definition": "qty(random(10..100#5),units)", "description": "", "templateType": "anything", "can_override": false}, "TC": {"name": "TC", "group": "magnitudes", "definition": "w1 cos(radians(a1))/sin(radians(a1+a2))\n\n", "description": "", "templateType": "anything", "can_override": false}, "w1": {"name": "w1", "group": "input", "definition": "qty(random(10..100#5),units)", "description": "", "templateType": "anything", "can_override": false}, "forceB": {"name": "forceB", "group": "check", "definition": "vector(-cos(radians(A1)),sin(radians(A1))) scalar(TB)", "description": "", "templateType": "anything", "can_override": false}, "TD": {"name": "TD", "group": "magnitudes", "definition": "TC cos(radians(a2))-TE cos(radians(a3))", "description": "", "templateType": "anything", "can_override": false}, "units": {"name": "units", "group": "input", "definition": "'lb'", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "a1 <> a2 and a3>a2+10 and scalar(TD) > 0", "maxRuns": 100}, "ungrouped_variables": [], "variable_groups": [{"name": "input", "variables": ["a1", "a2", "a3", "w1", "w2", "units", "debug"]}, {"name": "magnitudes", "variables": ["TC", "TE", "TD", "TB"]}, {"name": "check", "variables": ["forceB", "forceC", "forceE", "forceD"]}], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p style=\"padding-left: 40px;\">$T_{AB}$ =&nbsp;[[0]] <span data-jme-visible=\"debug\">{siground(TB,4)}</span></p>\n<p style=\"padding-left: 40px;\">$T_{AC}$ =&nbsp;[[1]] <span data-jme-visible=\"debug\">{siground(TC,4)}</span></p>\n<p style=\"padding-left: 40px;\">$T_{CE}$ =&nbsp;[[2]] <span data-jme-visible=\"debug\">{siground(TE,4)}</span></p>\n<p style=\"padding-left: 40px;\">$T_{CD}$ =&nbsp;[[3]] <span data-jme-visible=\"debug\">{siground(TD,4)}</span></p>", "gaps": [{"type": "engineering-answer", "useCustomName": true, "customName": "TB", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "TB", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "TC", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "TC", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "TE", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "TE", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "TD", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "TD", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}], "sortAnswers": false}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}], "resources": []}]}], "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}]}