// Numbas version: finer_feedback_settings
{"name": "VM8: Arbitrary loading function", "extensions": ["linear-algebra", "jsxgraph", "polynomials", "quantities"], "custom_part_types": [{"source": {"pk": 19, "author": {"name": "William Haynes", "pk": 2530}, "edit_page": "/part_type/19/edit"}, "name": "Engineering Accuracy with units", "short_name": "engineering-answer", "description": "<p>A value with units marked right if within an adjustable % error of the correct value. &nbsp;Marked close if within a wider margin of error.</p>", "help_url": "", "input_widget": "string", "input_options": {"correctAnswer": "siground(settings['correctAnswer'],4)", "hint": {"static": true, "value": ""}, "allowEmpty": {"static": true, "value": true}}, "can_be_gap": true, "can_be_step": true, "marking_script": "mark:\nswitch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)\n\ninterpreted_answer:\nqty(student_scalar, student_units)\n\n\n\ncorrect_quantity:\nsettings[\"correctAnswer\"]\n\n\n\ncorrect_units:\nunits(correct_quantity)\n\n\nallowed_notation_styles:\n[\"plain\",\"en\"]\n\nmatch_student_number:\nmatchnumber(studentAnswer,allowed_notation_styles)\n\nstudent_scalar:\nmatch_student_number[1]\n\nstudent_units:\nreplace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")\n\ngood_units:\ntry(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n\n\nstudent_quantity:\nswitch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n\n\npercent_error:\ntry(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   \n\nright:\npercent_error <= settings['right']\n\n\nclose:\nright_sign and percent_error <= settings['close']\n\nright_sign:\nsign(student_scalar) = sign(correct_quantity)", "marking_notes": [{"name": "mark", "description": "This is the main marking note. It should award credit and provide feedback based on the student's answer.", "definition": "switch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)"}, {"name": "interpreted_answer", "description": "A value representing the student's answer to this part.", "definition": "qty(student_scalar, student_units)\n\n"}, {"name": "correct_quantity", "description": "", "definition": "settings[\"correctAnswer\"]\n\n"}, {"name": "correct_units", "description": "", "definition": "units(correct_quantity)\n"}, {"name": "allowed_notation_styles", "description": "", "definition": "[\"plain\",\"en\"]"}, {"name": "match_student_number", "description": "", "definition": "matchnumber(studentAnswer,allowed_notation_styles)"}, {"name": "student_scalar", "description": "", "definition": "match_student_number[1]"}, {"name": "student_units", "description": "<p>Modify&nbsp;the&nbsp;unit portion of&nbsp;the student's answer by</p>\n<p>1. replacing \"ohms\" with \"ohm\" &nbsp;case insensitive</p>\n<p>2.&nbsp;replacing '-'&nbsp;with ' '&nbsp;</p>\n<p>3. replacing '&deg;' with ' deg'&nbsp;</p>\n<p>to allow answers like 10 ft-lb and 30&deg;</p>", "definition": "replace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")"}, {"name": "good_units", "description": "", "definition": "try(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n"}, {"name": "student_quantity", "description": "<p>This fixes the student answer for two common errors. &nbsp;</p>\n<p>If student_units are wrong&nbsp;&nbsp;- replace with correct units</p>\n<p>If student_scalar has the wrong sign - replace with right sign</p>\n<p>If student makes both errors, only one gets fixed.</p>", "definition": "switch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n"}, {"name": "percent_error", "description": "", "definition": "try(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   "}, {"name": "right", "description": "", "definition": "percent_error <= settings['right']\n"}, {"name": "close", "description": "<p>Only marked close if the student actually has the right sign.</p>", "definition": "right_sign and percent_error <= settings['close']"}, {"name": "right_sign", "description": "", "definition": "sign(student_scalar) = sign(correct_quantity) "}], "settings": [{"name": "correctAnswer", "label": "Correct Quantity.", "help_url": "", "hint": "The correct answer given as a JME quantity.", "input_type": "code", "default_value": "", "evaluate": true}, {"name": "right", "label": "% Accuracy for right.", "help_url": "", "hint": "Question will be considered correct if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "0.2", "evaluate": true}, {"name": "close", "label": "% Accuracy for close.", "help_url": "", "hint": "Question will be considered close if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "1.0", "evaluate": true}, {"name": "C1", "label": "Close with units.", "help_url": "", "hint": "Partial Credit for close value with appropriate units. &nbsp;if correct answer is 100 N and close is &plusmn;1%,<br />99 &nbsp;N is accepted.", "input_type": "percent", "default_value": "75"}, {"name": "C2", "label": "No units or wrong sign", "help_url": "", "hint": "Partial credit for forgetting units or using wrong sign.<br />If the correct answer is 100 N, both 100 and -100 N are accepted.", "input_type": "percent", "default_value": "50"}, {"name": "C3", "label": "Close, no units.", "help_url": "", "hint": "Partial Credit for close value but forgotten units.<br />This value would be close&nbsp;if the expected units were provided. &nbsp;If the correct answer is 100 N, and close is &plusmn;1%,<br />99 is accepted.", "input_type": "percent", "default_value": "25"}], "public_availability": "always", "published": true, "extensions": ["quantities"]}], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "VM8: Arbitrary loading function", "tags": ["Mechanics", "mechanics", "shear and bending moment", "Statics", "statics"], "metadata": {"description": "<p>Shear and bending moment diagram for a beam loaded with an arbitrary parabolic distributed load with a loading function &nbsp;given. &nbsp;Students must integrate to find reactions and equations for shear and bending moment. &nbsp;Problem sometimes fails when simplifying up the equations.</p>", "licence": "Creative Commons Attribution-NonCommercial 4.0 International"}, "statement": "<p>A simply-supported beam of length $L = \\var {qty(L, units[0])}$ is subjected to a parabolic loading $w(x)$ &nbsp;in {units[3]} given by the function:</p>\n<p>\\[w(x) = \\simplify[fractionNumbers,canonicalOrder]{{w}}\\] &nbsp;</p>\n<p>The load, shear and bending moment diagrams are shown below.</p>\n<p>{loadBoard}</p>\n<p>{shearBoard}</p>\n<p>{momentBoard}</p>\n<p>Find the reactions and derive the shear and bending moment as functions of $x$.</p>", "advice": "<p>Given $w(x)$ and beam length $L$,</p>\n<p>\\[ w(x) = \\simplify[fractionNumbers,canonicalOrder]{{w}} \\]</p>\n<p>1. &nbsp;Find load $W= \\int_0^L w(x)\\; dx$</p>\n<p style=\"padding-left: 40px;\">\\[\\begin{align} W &amp; = \\int w(x) \\text{d}x \\\\ &amp; = \\simplify[fractionNumbers,canonicalOrder]{defint({w}, x, 0, {L})} \\\\ &amp;= \\left[\\simplify[fractionNumbers,canonicalOrder]{{intw}}\\right]_0^\\var{L}\\\\ &amp;= \\var{siground(qty(load,units[1]),5)} = \\var{siground(qty(abs(load),units[1]),5)} &nbsp;\\var{arrow(-load)}\\end{align}\\]</p>\n<p>2. &nbsp;Find $Q_y = \\int_0^L x w(x)\\;dx$</p>\n<p style=\"padding-left: 40px;\">\\[\\begin{align} &nbsp;Q_y &amp;= \\int_0^L x\\ w(x) \\text{d}x\\\\&amp;= \\simplify[fractionNumbers,canonicalOrder]{defint({dqy},x,0,L)}\\\\&amp; = \\left[\\simplify[fractionNumbers,canonicalOrder]{{intdqy}} \\right]_0^\\var{L}\\\\ &amp;= \\var{siground(qty(Qy,units[0] + \" \" + units[1]),4)} \\end{align}\\]</p>\n<p>3. &nbsp;Find centroid $\\bar{x}=\\frac{Q_y}{W}$</p>\n<p style=\"padding-left: 40px;\">\\[\\bar{x} = \\dfrac{Q_y}{W} = \\dfrac{\\var{siground(Qy,4)}}{\\var{siground(load,4)}} = \\var{siground(qty(xbar,units[0]),4)}\\]</p>\n<p>4. &nbsp;Take moments about $A$ and $B$ to find reactions.</p>\n<p style=\"padding-left: 40px;\">\\[\\begin{align}<br/>\\sum M_B &amp;= 0 &amp; \\sum M_A &amp;=0 \\\\<br/>&nbsp; R_A L &amp;= &nbsp;\\simplify{{sign(load)} W }(L-\\bar{x}) &amp; + R_B L &amp;= \\simplify{{sign(load)} W }\\bar{x} &amp; &nbsp;\\\\<br/>R_A&amp;= \\var{force(abs(R_A))} \\var{arrow(R_A)} &amp;&nbsp;&nbsp;R_B &amp;= \\var{force(abs(R_B))} \\var{arrow(R_B)} <br/>&nbsp;\\end{align}\\]</p>\n<p>5. &nbsp;Integrate $V(x) = \\int w(x)\\; dx$ to get the shear function.</p>\n<p style=\"padding-left: 40px;\">Note: <em>Positive values of</em> $w(x)$ <em>indicate downward forces, so the integrations is over</em> $-w(x)$.</p>\n<p style=\"padding-left: 40px;\">\\[\\begin{align} V(x) &amp;= \\int - w(x) dx \\\\<br/>&amp;= \\simplify[fractionNumbers, canonicalOrder]{int({-w},x)} \\\\<br/>&amp;= \\simplify[fractionNumbers,canonicalOrder]{{v -v[0]}}+C_3 &nbsp;&amp;&amp; v(0) = R_A &nbsp;\\therefore C_3 = \\var{force(R_A)}<br/>\\\\ &amp;= \\simplify[fractionNumbers,canonicalOrder]{{v}} &amp;&amp; [\\var{units[1]}] &amp;\\end{align}\\]</p>\n<p>6. &nbsp;Integrate $M(x) = \\int V(x)\\; dx$ to get moment function. &nbsp;</p>\n<p style=\"padding-left: 30px;\">\\[\\begin{align}<br/>M(x) &amp;= \\int V(x) dx \\\\<br/>&amp;= \\simplify[fractionNumbers, canonicalOrder]{int({v},x)}\\\\<br/>&amp;= \\simplify[fractionNumbers,canonicalOrder]{{m}}+C_4 &nbsp;&amp;&amp; M(0) =0 &nbsp;\\therefore C_4 = 0 \\\\<br/>&amp;= \\simplify[fractionNumbers,canonicalOrder]{{m}} &amp;&amp; [\\var{units[1]}\\text{-}\\var{units[0]} ]&amp;<br/>\\end{align}\\]</p>\n<p style=\"padding-left: 30px;\"></p>\n<p style=\"padding-left: 30px;\"><span data-jme-visible=\"false\">Maximum moment occurs where $dM/dx = V = 0$</span></p>\n<p style=\"padding-left: 30px;\"><span data-jme-visible=\"false\">Evaluate $M(x)$ at point of maximum bending moment.</span></p>\n<p style=\"padding-left: 30px;\"><span data-jme-visible=\"false\">Maximum shear occurs where $d V&nbsp;/d x = -w(x) = 0$</span></p>\n<p style=\"padding-left: 30px;\"><span data-jme-visible=\"false\">Evaluate $V(x)$ at point of maximum shear.</span></p>\n<p style=\"padding-left: 40px;\"></p>\n<p style=\"padding-left: 40px;\"><span data-jme-visible=\"false\">To find maximum shear, determine where $\\simplify{diff(v,x,1)} = -w(x) = 0$ using the quadratic equation.</span></p>\n<p style=\"padding-left: 40px;\"><span data-jme-visible=\"false\">$x_{max} = &nbsp;\\frac{- b \\pm \\sqrt{b^2 - 4 a c}}{2a}$ where, $a= \\var[fractionNumbers]{-w[2]}, b=\\var[fractionNumbers]{-w[1]}$ and $c=\\var[fractionNumbers]{-w[0]}$.</span></p>\n<p style=\"padding-left: 40px;\"><span data-jme-visible=\"false\">Any real roots of $w(x)$ between the beam's endpoints plus the endpoints themselves are locations of local maximums or minimums.</span></p>\n<p style=\"padding-left: 40px;\"><span data-jme-visible=\"false\">{v_extremes}</span></p>\n<p style=\"padding-left: 40px;\"><span data-jme-visible=\"false\">The maximum of these local maximums is the the one we want</span></p>\n<p style=\"padding-left: 40px;\"><span data-jme-visible=\"false\">$V_{max} = \\var{show(qty(v_max, units[1]))}$</span></p>\n<p style=\"padding-left: 40px;\"></p>\n<p style=\"padding-left: 40px;\"></p>", "rulesets": {}, "extensions": ["jsxgraph", "linear-algebra", "polynomials", "quantities"], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"L": {"name": "L", "group": "Unnamed group", "definition": "20", "description": "", "templateType": "anything", "can_override": false}, "units": {"name": "units", "group": "Unnamed group", "definition": "random(['ft','lb','ft*lb', 'lb/ft'],['m','kN', 'kN*m', 'kN/m'])", "description": "", "templateType": "anything", "can_override": false}, "w": {"name": "w", "group": "functions", "definition": "polynomial(x,load_coefficients)", "description": "<p>loading function w(x)</p>\n<p></p>", "templateType": "anything", "can_override": false}, "v": {"name": "v", "group": "functions", "definition": "-polynomial(x, [-R_A, w[0], w[1]/2, w[2]/3, w[3]/4])", "description": "<p>This is the shear function V(x) = - \\int w(x) dx</p>\n<p>Constant C is reaction at x=0</p>", "templateType": "anything", "can_override": false}, "m": {"name": "m", "group": "functions", "definition": "polynomial(x,moment_coefficients)", "description": "", "templateType": "anything", "can_override": false}, "load": {"name": "load", "group": "Reactions", "definition": "eval(intW,L) - eval(intW,0)", "description": "<p>downward force on beam = $\\simplify{defint(w,x,0,{beamlen})$</p>\n<p></p>\n<p>Positive = down</p>", "templateType": "anything", "can_override": false}, "dQy": {"name": "dQy", "group": "Reactions", "definition": "polynomial(x) w", "description": "", "templateType": "anything", "can_override": false}, "intdQy": {"name": "intdQy", "group": "Reactions", "definition": "polynomial(x, [0, dQy[0], dQy[1]/2, dQy[2]/3, dQy[3]/4, dQy[4]/5])", "description": "<p>integral of dQy &nbsp;&nbsp;</p>", "templateType": "anything", "can_override": false}, "Qy": {"name": "Qy", "group": "Reactions", "definition": "eval(intdqy,L)", "description": "<p>definite integral of x dw from 0 to L</p>", "templateType": "anything", "can_override": false}, "xBar": {"name": "xBar", "group": "Reactions", "definition": "Qy/load", "description": "", "templateType": "anything", "can_override": false}, "intW": {"name": "intW", "group": "Reactions", "definition": "polynomial(x, [0, w[0], w[1]/2, w[2]/3, w[3]/4])", "description": "", "templateType": "anything", "can_override": false}, "R_B": {"name": "R_B", "group": "Reactions", "definition": "xbar load / L", "description": "<p>reaction at B &nbsp;Positive is up</p>", "templateType": "anything", "can_override": false}, "R_A": {"name": "R_A", "group": "Reactions", "definition": "(L-xbar) load/ L", "description": "", "templateType": "anything", "can_override": false}, "check": {"name": "check", "group": "Reactions", "definition": "withintolerance(R_A+R_B,load,10^10)", "description": "<p>check sometimes has roundoff error.</p>", "templateType": "anything", "can_override": false}, "roots": {"name": "roots", "group": "maxshear", "definition": "sort(\n  if(d>=0, \n    filter(x>=0 and x <=L, x,\n      let([a: w[2], b: w[1], c: w[0]], \n     [(-b + sqrt(d))/(2 a), (-b - sqrt(d))/(2 a)]\n)),[]) + [0,L])", "description": "<p>use quadratic equation to find where w(x) is zero, filter out any roots not on interval, or imaginary.</p>\n<p>The maximum shear will occur at one of these points</p>", "templateType": "anything", "can_override": false}, "d": {"name": "d", "group": "Ungrouped variables", "definition": "w[1]^2 - 4 w[2] * w[0]", "description": "<p>descriminant for quadratic equation w(x)</p>", "templateType": "anything", "can_override": false}, "v_extremes": {"name": "v_extremes", "group": "maxshear", "definition": "map(vector(x,eval(v,x)),x,roots)", "description": "<p>extreme points</p>\n<p>list of shear force v(x) at ends and at roots of w(x) -- potential maximums or minimums</p>", "templateType": "anything", "can_override": false}, "v_max": {"name": "v_max", "group": "functions", "definition": "max([v_maxmin[0],-v_maxmin[1]])", "description": "<p>filters the list of v_extremes to find the one with the greatest absolute value of shear</p>", "templateType": "anything", "can_override": false}, "debug": {"name": "debug", "group": "Ungrouped variables", "definition": "false", "description": "", "templateType": "anything", "can_override": false}, "loadBoard": {"name": "loadBoard", "group": "Ungrouped variables", "definition": "let(max, w_maxmin,\n  jsxgraph(\n    500,250,[-2, max(1.2 max[0],-max[1]/2),  L+2,  min(1.2 max[1], -max[0]/2)],\n    merge(\n    arrows,\n    [\"$w(x)$\": ['functiongraph', [load_function, 0, L], [name: \"a\", withLabel: true, strokeWidth: 3, strokeColor: \"#CC0000\", fillOpacity: 0.1,label:[fixed: false, fontSize: 14]]]],\n    [\"beam\":   ['segment', [[0,0],[L,0]], [strokeWidth: 5, color: \"black\"]] ]\n     \n\n      )\n))\n", "description": "", "templateType": "anything", "can_override": false}, "load_function": {"name": "load_function", "group": "functions", "definition": "expr(w)", "description": "", "templateType": "anything", "can_override": false}, "reactionoptions": {"name": "reactionoptions", "group": "Ungrouped variables", "definition": "[withLabel: true, \"fixed\": true, \"color\": \"#CC0000\", \"strokewidth\": 3, \"highlight\": false, \"lastArrow\": [ \"type\": 4 ] ]", "description": "", "templateType": "anything", "can_override": false}, "drawArrow": {"name": "drawArrow", "group": "Ungrouped variables", "definition": " (x) ->  ['segment', [[x, eval(w,x)],[x,0]],[ fixed: true, color: \"black\",strokewidth: 1.5,highlight: false, lastArrow: [ type: 4 ]]]", "description": "", "templateType": "anything", "can_override": false}, "arrows": {"name": "arrows", "group": "Ungrouped variables", "definition": "dict(map((i)->([(\"A\"+i),drawArrow(i)]), 1..L))", "description": "", "templateType": "anything", "can_override": false}, "w_max": {"name": "w_max", "group": "functions", "definition": "max([abs(r_a),abs(r_b)])", "description": "<p>hack to scale load diagram</p>", "templateType": "anything", "can_override": false}, "w_maxmin": {"name": "w_maxmin", "group": "functions", "definition": "[max(map(x->eval(w,x),0..l)) |> round(), round(min(map(x->eval(w,x),0..l)))]", "description": "", "templateType": "anything", "can_override": false}, "shearBoard": {"name": "shearBoard", "group": "Ungrouped variables", "definition": "\n  jsxgraph(\n    500,250,[-2, 1.2 v_max, L+2, -v_max],\n    [   \n        \"$V(x)$\": ['functiongraph', [shear_function, 0, L], [withLabel: true, strokeWidth: 3, strokeColor: \"#CC0000\", label:[fixed: false, fontSize: 14]]],\n\n    ]\n)", "description": "", "templateType": "anything", "can_override": false}, "shear_function": {"name": "shear_function", "group": "functions", "definition": "expr(v)", "description": "", "templateType": "anything", "can_override": false}, "v_maxmin": {"name": "v_maxmin", "group": "functions", "definition": "[max(map(x->eval(v,x),0..l)) |> round(), round(min(map(x->eval(v,x),0..l)))]", "description": "", "templateType": "anything", "can_override": false}, "moment_function": {"name": "moment_function", "group": "functions", "definition": "expr(m)", "description": "", "templateType": "anything", "can_override": false}, "m_maxmin": {"name": "m_maxmin", "group": "functions", "definition": "[max(map(x->eval(m,x),0..l)) |> round(), round(min(map(x->eval(m,x),0..l)))]", "description": "", "templateType": "anything", "can_override": false}, "m_max": {"name": "m_max", "group": "functions", "definition": "max([m_maxmin[0],-m_maxmin[1]])", "description": "", "templateType": "anything", "can_override": false}, "momentBoard": {"name": "momentBoard", "group": "Ungrouped variables", "definition": "  let(max, m_maxmin[0], min, m_maxmin[1],\n  jsxgraph(\n    500,250,[-2, 1.2 max(1.2 max, -min/2), L+2, min(1.2 min, -max/2)],\n    [\"$M(x)$\": ['functiongraph', [moment_function, 0, L], [withLabel: true, strokeWidth: 3, strokeColor: \"#3070AD\", label:[fixed: false, fontSize: 14] ]],\n ]\n  )  )", "description": "<p>max(1.2 max[0],-max[1]/2)</p>", "templateType": "anything", "can_override": false}, "force": {"name": "force", "group": "Ungrouped variables", "definition": "(f)->siground(qty(f,units[1]),4)", "description": "", "templateType": "anything", "can_override": false}, "load_coefficients": {"name": "load_coefficients", "group": "functions", "definition": "random(coefficients)", "description": "<p>get coefficients puts curve through three points, A,B,C &nbsp;siground tries to make the coefficients simpler fractions</p>", "templateType": "anything", "can_override": false}, "moment_coefficients": {"name": "moment_coefficients", "group": "functions", "definition": "[0,v[0],v[1]/2,v[2]/3,v[3]/4,v[4]/5]", "description": "", "templateType": "anything", "can_override": false}, "shear_coefficients": {"name": "shear_coefficients", "group": "functions", "definition": "[-R_A, w[0], w[1]/2, w[2]/3, w[3]/4]", "description": "", "templateType": "anything", "can_override": false}, "random_frac": {"name": "random_frac", "group": "Ungrouped variables", "definition": "() -> random(-9..9 except 0)/random(1..9)", "description": "", "templateType": "anything", "can_override": false}, "coefficients": {"name": "coefficients", "group": "functions", "definition": "[[ 0, -10, 3, -0.1 ],\n[ 0, -10, 4, -0.1 ],\n[ 0, -10, 4, -0.1 ],\n[ 0, -10, 4, -0.2 ],\n[ 0, -20, 4, -0.2 ],\n[ 0, -20, 5, -0.2 ],\n[ 0, -3, 0.3, -0.006 ],\n[ 0, -4, 0.04, 0.009 ],\n[ 0, -4, 0.07, 0.006 ],\n[ 0, -4, 0.3, -0.007 ],\n[ 0, -4, 0.4, -0.008 ],\n[ 0, -5, 0.5, -0.01 ],\n[ 0, -5, 0.6, -0.02 ],\n[ 0, -5, 0.7, -0.02 ],\n[ 0, -6, 0.5, -0.01 ],\n[ 0, -6, 0.8, -0.03 ],\n[ 0, -7, 0.9, -0.03 ],\n[ 0, -8, 1, -0.03 ],\n[ 0, -8, 1, -0.04 ],\n[ 0, -9, 1, -0.04 ],\n[ 0, 2, -0.04, -0.004 ],\n[ 0, 2, -0.05, -0.004 ],\n[ 0, 2, -0.1, 0 ],\n[ 0, 20, -5, 0.2 ],\n[ 0, 3, -0.03, -0.007 ],\n[ 0, 3, -0.05, -0.005 ],\n[ 0, 3, -0.1, 0 ],\n[ 0, 3, -0.3, 0.006 ],\n[ 0, 4, -0.03, -0.007 ],\n[ 0, 4, -0.04, -0.008 ],\n[ 0, 4, -0.04, -0.009 ],\n[ 0, 4, -0.07, -0.006 ],\n[ 0, 4, -0.2, 0 ],\n[ 0, 5, -0.05, -0.01 ],\n[ 0, 5, -0.7, 0.02 ],\n[ 0, 6, -0.05, -0.01 ],\n[ 0, 6, -0.06, -0.01 ],\n[ 0, 6, -0.8, 0.03 ],\n[ 0, 8, -1, 0.03 ],\n[ 0, -3, 0.3, -0.007 ]]", "description": "<p>These coefficients do not cause \"Simplifier is stuck in a loop\" errors. &nbsp;The were generated by giving three random points to getCoefficients() extension function</p>", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["d", "debug", "loadBoard", "shearBoard", "momentBoard", "reactionoptions", "drawArrow", "arrows", "force", "random_frac"], "variable_groups": [{"name": "Unnamed group", "variables": ["L", "units"]}, {"name": "functions", "variables": ["w", "v", "m", "coefficients", "load_coefficients", "shear_coefficients", "moment_coefficients", "load_function", "shear_function", "moment_function", "w_maxmin", "v_maxmin", "m_maxmin", "w_max", "v_max", "m_max"]}, {"name": "Reactions", "variables": ["load", "dQy", "intdQy", "Qy", "xBar", "intW", "R_B", "R_A", "check"]}, {"name": "maxshear", "variables": ["roots", "v_extremes"]}], "functions": {"arrow": {"parameters": [["F", "?"]], "type": "anything", "language": "jme", "definition": "if(sign(F) >= 0 ,latex(\"\\\\uparrow\"), latex(\"\\\\downarrow\"))"}, "getCoefficients": {"parameters": [["A", "list of integer"], ["B", "list of integer"], ["C", "list of integer"]], "type": "list", "language": "javascript", "definition": "    const jme = Numbas.jme;\n    const extension = Numbas.extensions.polynomial\n    const [x0, y1] = A;\n    const [x2, y2] = B;\n    const [x3, y3] = C;\n\n    // Sanity check: first point must be at x=0\n    if (x0 !== 0) {\n        throw new Error(\"First point A must have x = 0\");\n    }\n\n    const d = y1;\n\n    // Right-hand side: subtract y1 to isolate a,b,c terms\n    const r1 = y2 - y1;\n    const r2 = y3 - y1;\n\n    // Matrix coefficients\n    const a11 = Math.pow(x2, 3);\n    const a12 = Math.pow(x2, 2);\n    const a13 = x2;\n\n    const a21 = Math.pow(x3, 3);\n    const a22 = Math.pow(x3, 2);\n    const a23 = x3;\n\n    // Solve the system using Cramer's Rule or elimination\n    // We'll eliminate c first\n\n    // Multiply row 1 by x3 and row 2 by x2 to align c terms\n    const m1 = x3, m2 = x2;\n\n    const A1 = m1 * a11 - m2 * a21;\n    const B1 = m1 * a12 - m2 * a22;\n    const R1 = m1 * r1 - m2 * r2;\n\n    // Now solve A1 * a + B1 * b = R1\n    // We'll solve for b in terms of a, then back-substitute\n\n    // Solve for a and b\n    // System: A1 * a + B1 * b = R1\n    //         a13 * a + a12 * b + c = r1\n\n    // Choose to solve directly (2x2 linear system)\n    const det = A1 * (-a12) - B1 * (-a13);\n    if (Math.abs(det) < 1e-12) {\n        throw new Error(\"Points are degenerate or lead to singular system.\");\n    }\n\n    // Use Cramer's Rule\n    const detA = R1 * (-a12) - B1 * (r1 - a13 * 0); // assuming a = 0 temporarily\n    const detB = A1 * (r1 - a13 * 0) - R1 * (-a13);\n\n    const a = detA / det;\n    const b = detB / det;\n\n    // Solve for c using first equation\n    const c = (r1 - a * a11 - b * a12) / a13;\n\n    // Return function  a * x ** 3 + b * x ** 2 + c * x + d;\n    return [d,c,b,a]\n    \n"}, "show": {"parameters": [["A", "number"]], "type": "anything", "language": "jme", "definition": "if(abs(if(type(A)=type(qty(1,'m')),scalar(A),A))<1,\n  precround(A,4),siground(A,4))"}}, "preamble": {"js": "question.signals.on('adviceDisplayed',function() {\n   try{\n   }\n  catch(err){}  \n})\n\n\n", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": true, "customName": "Reactions", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Use integration to replace the distributed load with an equivalent concentrated force, then find the reactions at the end suppports. &nbsp;Use negative values to indicate downward forces.</p>\n<p style=\"padding-left: 40px;\">$R_A = $ [[0]]</p>\n<p style=\"padding-left: 40px;\">$ R_B= $ [[1]]</p>", "gaps": [{"type": "engineering-answer", "useCustomName": true, "customName": "A", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(R_A,units[1])", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "B", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(R_B,units[1])", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}], "sortAnswers": false}, {"type": "gapfill", "useCustomName": true, "customName": "Shear Function", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Use integration to determine the equation of the shear function $V (x) = \\int -w(x)\\ dx$</p>\n<p><span data-jme-visible=\"debug\">$V\\!(x) = \\simplify[fractionNumbers,canonicalOrder]{{v}} = &nbsp;\\simplify[canonicalOrder]{{v}} $</span></p>\n<p style=\"padding-left: 40px;\">$V\\!(x) = $[[0]]</p>", "gaps": [{"type": "jme", "useCustomName": true, "customName": "$V(x)$", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "answer": "{expr(v)}", "answerSimplification": "fractionNumbers,canonicalOrder", "showPreview": true, "checkingType": "sigfig", "checkingAccuracy": "2", "failureRate": 1, "vsetRangePoints": 5, "vsetRange": [0, "1"], "checkVariableNames": false, "singleLetterVariables": false, "allowUnknownFunctions": true, "implicitFunctionComposition": false, "caseSensitive": false, "valuegenerators": []}], "sortAnswers": false}, {"type": "gapfill", "useCustomName": true, "customName": "Bending Moment", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Use integration to find the bending moment function $M(x) =\\int V(x)\\ dx$.</p>\n<p><span data-jme-visible=\"debug\">$M(x) = \\simplify[fractionNumbers,canonicalOrder]{{m}}$&nbsp;<br/></span></p>\n<p><span data-jme-visible=\"debug\">$=\\simplify[canonicalOrder]{{m}}$</span></p>\n<p style=\"padding-left: 40px;\">$M(x) = $ [[0]]</p>", "gaps": [{"type": "jme", "useCustomName": true, "customName": "$M(x)$", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "answer": "{expr(m)}", "answerSimplification": "fractionNumbers,canonicalOrder", "showPreview": true, "checkingType": "sigfig", "checkingAccuracy": "2", "failureRate": 1, "vsetRangePoints": 5, "vsetRange": [0, 1], "checkVariableNames": false, "singleLetterVariables": false, "allowUnknownFunctions": true, "implicitFunctionComposition": false, "caseSensitive": false, "valuegenerators": []}], "sortAnswers": false}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}], "resources": []}]}], "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}]}