// Numbas version: finer_feedback_settings
{"name": "Truss: Solve a joint", "extensions": ["geogebra", "quantities"], "custom_part_types": [{"source": {"pk": 19, "author": {"name": "William Haynes", "pk": 2530}, "edit_page": "/part_type/19/edit"}, "name": "Engineering Accuracy with units", "short_name": "engineering-answer", "description": "<p>A value with units marked right if within an adjustable % error of the correct value. &nbsp;Marked close if within a wider margin of error.</p>", "help_url": "", "input_widget": "string", "input_options": {"correctAnswer": "siground(settings['correctAnswer'],4)", "hint": {"static": true, "value": ""}, "allowEmpty": {"static": true, "value": true}}, "can_be_gap": true, "can_be_step": true, "marking_script": "mark:\nswitch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)\n\ninterpreted_answer:\nqty(student_scalar, student_units)\n\n\n\ncorrect_quantity:\nsettings[\"correctAnswer\"]\n\n\n\ncorrect_units:\nunits(correct_quantity)\n\n\nallowed_notation_styles:\n[\"plain\",\"en\"]\n\nmatch_student_number:\nmatchnumber(studentAnswer,allowed_notation_styles)\n\nstudent_scalar:\nmatch_student_number[1]\n\nstudent_units:\nreplace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")\n\ngood_units:\ntry(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n\n\nstudent_quantity:\nswitch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n\n\npercent_error:\ntry(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   \n\nright:\npercent_error <= settings['right']\n\n\nclose:\nright_sign and percent_error <= settings['close']\n\nright_sign:\nsign(student_scalar) = sign(correct_quantity)", "marking_notes": [{"name": "mark", "description": "This is the main marking note. It should award credit and provide feedback based on the student's answer.", "definition": "switch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)"}, {"name": "interpreted_answer", "description": "A value representing the student's answer to this part.", "definition": "qty(student_scalar, student_units)\n\n"}, {"name": "correct_quantity", "description": "", "definition": "settings[\"correctAnswer\"]\n\n"}, {"name": "correct_units", "description": "", "definition": "units(correct_quantity)\n"}, {"name": "allowed_notation_styles", "description": "", "definition": "[\"plain\",\"en\"]"}, {"name": "match_student_number", "description": "", "definition": "matchnumber(studentAnswer,allowed_notation_styles)"}, {"name": "student_scalar", "description": "", "definition": "match_student_number[1]"}, {"name": "student_units", "description": "<p>Modify&nbsp;the&nbsp;unit portion of&nbsp;the student's answer by</p>\n<p>1. replacing \"ohms\" with \"ohm\" &nbsp;case insensitive</p>\n<p>2.&nbsp;replacing '-'&nbsp;with ' '&nbsp;</p>\n<p>3. replacing '&deg;' with ' deg'&nbsp;</p>\n<p>to allow answers like 10 ft-lb and 30&deg;</p>", "definition": "replace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")"}, {"name": "good_units", "description": "", "definition": "try(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n"}, {"name": "student_quantity", "description": "<p>This fixes the student answer for two common errors. &nbsp;</p>\n<p>If student_units are wrong&nbsp;&nbsp;- replace with correct units</p>\n<p>If student_scalar has the wrong sign - replace with right sign</p>\n<p>If student makes both errors, only one gets fixed.</p>", "definition": "switch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n"}, {"name": "percent_error", "description": "", "definition": "try(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   "}, {"name": "right", "description": "", "definition": "percent_error <= settings['right']\n"}, {"name": "close", "description": "<p>Only marked close if the student actually has the right sign.</p>", "definition": "right_sign and percent_error <= settings['close']"}, {"name": "right_sign", "description": "", "definition": "sign(student_scalar) = sign(correct_quantity) "}], "settings": [{"name": "correctAnswer", "label": "Correct Quantity.", "help_url": "", "hint": "The correct answer given as a JME quantity.", "input_type": "code", "default_value": "", "evaluate": true}, {"name": "right", "label": "% Accuracy for right.", "help_url": "", "hint": "Question will be considered correct if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "0.2", "evaluate": true}, {"name": "close", "label": "% Accuracy for close.", "help_url": "", "hint": "Question will be considered close if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "1.0", "evaluate": true}, {"name": "C1", "label": "Close with units.", "help_url": "", "hint": "Partial Credit for close value with appropriate units. &nbsp;if correct answer is 100 N and close is &plusmn;1%,<br />99 &nbsp;N is accepted.", "input_type": "percent", "default_value": "75"}, {"name": "C2", "label": "No units or wrong sign", "help_url": "", "hint": "Partial credit for forgetting units or using wrong sign.<br />If the correct answer is 100 N, both 100 and -100 N are accepted.", "input_type": "percent", "default_value": "50"}, {"name": "C3", "label": "Close, no units.", "help_url": "", "hint": "Partial Credit for close value but forgotten units.<br />This value would be close&nbsp;if the expected units were provided. &nbsp;If the correct answer is 100 N, and close is &plusmn;1%,<br />99 is accepted.", "input_type": "percent", "default_value": "25"}], "public_availability": "always", "published": true, "extensions": ["quantities"]}], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Truss: Solve a joint", "tags": ["equilibrium of a particle", "Mechanics", "mechanics", "method of joints", "Statics", "statics", "truss"], "metadata": {"description": "<p>This just an equilibrium of a particle problem, but posed in context of a joint as a warm-up for the method of joints.</p>\n<p>Advice is brief.</p>", "licence": "Creative Commons Attribution-ShareAlike 4.0 International"}, "statement": "<p>The diagram shows the free body diagram of a joint in a truss which is supporting a load, $F$. In the FBD the direction (sense) of forces $A$ and $B$ are assumed, so may not be correct.</p>\n<table style=\"margin-left: auto; margin-right: auto;\">\n<tbody>\n<tr>\n<td>{applet}</td>\n</tr>\n</tbody>\n</table>\n<p>Given: $F = \\var{qty(magF,unit)}$ and $C = \\var{(qty(abs(magC),unit))}$</p>\n<ol>\n<li>Draw&nbsp;draw a neat, labeled free body diagram representing the situation.</li>\n<li>Write two equilibrium equations, symbolically, based on your free body diagram.</li>\n<li>Solve your equations to determine the magnitudes of forces $A$ and $B$ necessary for equilibrium.</li>\n<li>Indicate whether forces $A$ and $B$ are in tension or compression.</li>\n</ol>\n<p></p>", "advice": "<p>Draw&nbsp;a free body diagram, then write two equilibrium equations:</p>\n<p>$\\Sigma F_x = 0&nbsp;\\qquad\\Sigma F_y = 0$</p>\n<p>Solve them simultaneously for unknown magnitudes $A$ and $B$.</p>\n<p></p>\n<p>Answers:</p>\n<p>$A = \\var{qty(precround(abs(magA),1),unit)}$ {if(magA=0,'',if(magA&gt;0,'Tension','Compression'))}</p>\n<p>$B = \\var{qty(precround(abs(magB),1),unit)}$&nbsp;{if(magB=0,'',if(magB&gt;0,'Tension','Compression'))}</p>\n<p></p>", "rulesets": {}, "extensions": ["geogebra", "quantities"], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"a": {"name": "a", "group": "Ungrouped variables", "definition": "random(30..270#30 except [90,270]) ", "description": "<p>dirction of force A</p>", "templateType": "anything", "can_override": false}, "b": {"name": "b", "group": "Ungrouped variables", "definition": "random([45,135,225,315])", "description": "<p>direction of force b</p>", "templateType": "anything", "can_override": false}, "c": {"name": "c", "group": "Ungrouped variables", "definition": "random(0..360#30)", "description": "<p>direction of force c</p>", "templateType": "anything", "can_override": false}, "f": {"name": "f", "group": "Ungrouped variables", "definition": "random(0..270#90)", "description": "<p>direction of force f</p>", "templateType": "anything", "can_override": false}, "magF": {"name": "magF", "group": "Ungrouped variables", "definition": "random(100..500#25)", "description": "", "templateType": "anything", "can_override": false}, "MagC": {"name": "MagC", "group": "Ungrouped variables", "definition": "random(50..245#50 except 0)", "description": "", "templateType": "anything", "can_override": false}, "unit": {"name": "unit", "group": "Ungrouped variables", "definition": "'kN'", "description": "", "templateType": "anything", "can_override": false}, "c1": {"name": "c1", "group": "solution", "definition": "-(magC cos(radians(c)) + magF cos(radians(f)))", "description": "", "templateType": "anything", "can_override": false}, "c2": {"name": "c2", "group": "solution", "definition": "-(magC sin(radians(c)) + magF sin(radians(f)))", "description": "", "templateType": "anything", "can_override": false}, "magB": {"name": "magB", "group": "solution", "definition": "(c2-c1 tan(radians(a)))/(-cos(radians(b))tan(radians(a)) + sin(radians(b)))\n", "description": "", "templateType": "anything", "can_override": false}, "magA": {"name": "magA", "group": "solution", "definition": "(c1-magB cos(radians(b)))/cos(radians(a))", "description": "", "templateType": "anything", "can_override": false}, "check_x": {"name": "check_x", "group": "solution", "definition": "precround(magA cos(radians(a)) + magB cos(radians(b)) + magC cos(radians(c)) + magF cos(radians(f)),6)", "description": "", "templateType": "anything", "can_override": false}, "check_y": {"name": "check_y", "group": "solution", "definition": "precround( magA sin(radians(a)) + magB sin(radians(b)) + magC sin(radians(c)) + magF sin(radians(f)),6)", "description": "", "templateType": "anything", "can_override": false}, "applet": {"name": "applet", "group": "Unnamed group", "definition": "geogebra_applet('yaqwabnr', params)\n\n", "description": "", "templateType": "anything", "can_override": false}, "params": {"name": "params", "group": "Unnamed group", "definition": "[a: a+'\u00b0', b: b+'\u00b0', c: c+'\u00b0', f: f+'\u00b0',\n  hideA: 'false', hideB: 'false', hideC: 'false', hideF: 'false']\n", "description": "", "templateType": "anything", "can_override": false}, "FF": {"name": "FF", "group": "alternate solution", "definition": "vector(cos(radians(f))\n  ,sin(radians(f))) magF", "description": "<p>Vector F</p>", "templateType": "anything", "can_override": false}, "FC": {"name": "FC", "group": "alternate solution", "definition": "vector(cos(radians(c)),sin(radians(c))) magC", "description": "<p>Force C</p>", "templateType": "anything", "can_override": false}, "FR": {"name": "FR", "group": "alternate solution", "definition": "FF + FC", "description": "<p>R = F + C</p>", "templateType": "anything", "can_override": false}, "AmatrixInverse": {"name": "AmatrixInverse", "group": "alternate solution", "definition": "matrix([sin(radians(b)), - cos(radians(b))],[-sin(radians(a)),cos(radians(a))])/det(Amatrix)", "description": "<p>used for matrix solution [X] = A^{-1} [B]</p>", "templateType": "anything", "can_override": false}, "Bmatrix": {"name": "Bmatrix", "group": "alternate solution", "definition": "matrix([[-FR[0]],[-FR[1]]])", "description": "", "templateType": "anything", "can_override": false}, "Amatrix": {"name": "Amatrix", "group": "alternate solution", "definition": "matrix([[cos(radians(a)),cos(radians(b))],[sin(radians(a)),sin(radians(b))]])", "description": "<p>used for matric solution&nbsp;</p>", "templateType": "anything", "can_override": false}, "xMatrix": {"name": "xMatrix", "group": "alternate solution", "definition": "Amatrixinverse Bmatrix", "description": "<p>[A,B]</p>", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "a <> b and a <> c and a <> f and b<> c and b<> f and c<> f", "maxRuns": 100}, "ungrouped_variables": ["a", "b", "c", "f", "magF", "MagC", "unit"], "variable_groups": [{"name": "solution", "variables": ["c1", "c2", "magB", "magA", "check_x", "check_y"]}, {"name": "Unnamed group", "variables": ["applet", "params"]}, {"name": "alternate solution", "variables": ["FF", "FC", "FR", "AmatrixInverse", "Bmatrix", "Amatrix", "xMatrix"]}], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p style=\"padding-left: 40px;\"><br/>$A=$ [[0]]&nbsp;[[1]]<br/><br/>$B=$ [[2]]&nbsp;[[3]] <br/><br/></p>", "gaps": [{"type": "engineering-answer", "useCustomName": true, "customName": "Magnitude of A", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "abs(qty(magA,unit))", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": true, "customName": "T or C", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "choices": ["Tension", "Compression", "Neither"], "matrix": ["if(sign(magA) > 0,5,0)", "if(sign(magA) < 0,5,0)", "if(sign(magA) = 0,5,0)"], "distractors": ["", "", ""]}, {"type": "engineering-answer", "useCustomName": true, "customName": "Magnitude of B", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "abs(qty(magB,unit))", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": true, "customName": "T or C", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "choices": ["Tension", "Compression", "Neither"], "matrix": ["if(sign(magB) > 0,5,0)", "if(sign(magB) < 0,5,0)", "if(sign(magB) = 0,5,0)"], "distractors": ["", "", ""]}], "sortAnswers": false}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}], "resources": []}]}], "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}]}