// Numbas version: finer_feedback_settings
{"name": "Distributed Loads: trapezoidal loading", "extensions": ["jsxgraph", "shear-and-bending-moment-diagrams", "quantities"], "custom_part_types": [{"source": {"pk": 19, "author": {"name": "William Haynes", "pk": 2530}, "edit_page": "/part_type/19/edit"}, "name": "Engineering Accuracy with units", "short_name": "engineering-answer", "description": "<p>A value with units marked right if within an adjustable % error of the correct value. &nbsp;Marked close if within a wider margin of error.</p>", "help_url": "", "input_widget": "string", "input_options": {"correctAnswer": "siground(settings['correctAnswer'],4)", "hint": {"static": true, "value": ""}, "allowEmpty": {"static": true, "value": true}}, "can_be_gap": true, "can_be_step": true, "marking_script": "mark:\nswitch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)\n\ninterpreted_answer:\nqty(student_scalar, student_units)\n\n\n\ncorrect_quantity:\nsettings[\"correctAnswer\"]\n\n\n\ncorrect_units:\nunits(correct_quantity)\n\n\nallowed_notation_styles:\n[\"plain\",\"en\"]\n\nmatch_student_number:\nmatchnumber(studentAnswer,allowed_notation_styles)\n\nstudent_scalar:\nmatch_student_number[1]\n\nstudent_units:\nreplace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")\n\ngood_units:\ntry(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n\n\nstudent_quantity:\nswitch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n\n\npercent_error:\ntry(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   \n\nright:\npercent_error <= settings['right']\n\n\nclose:\nright_sign and percent_error <= settings['close']\n\nright_sign:\nsign(student_scalar) = sign(correct_quantity)", "marking_notes": [{"name": "mark", "description": "This is the main marking note. It should award credit and provide feedback based on the student's answer.", "definition": "switch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)"}, {"name": "interpreted_answer", "description": "A value representing the student's answer to this part.", "definition": "qty(student_scalar, student_units)\n\n"}, {"name": "correct_quantity", "description": "", "definition": "settings[\"correctAnswer\"]\n\n"}, {"name": "correct_units", "description": "", "definition": "units(correct_quantity)\n"}, {"name": "allowed_notation_styles", "description": "", "definition": "[\"plain\",\"en\"]"}, {"name": "match_student_number", "description": "", "definition": "matchnumber(studentAnswer,allowed_notation_styles)"}, {"name": "student_scalar", "description": "", "definition": "match_student_number[1]"}, {"name": "student_units", "description": "<p>Modify&nbsp;the&nbsp;unit portion of&nbsp;the student's answer by</p>\n<p>1. replacing \"ohms\" with \"ohm\" &nbsp;case insensitive</p>\n<p>2.&nbsp;replacing '-'&nbsp;with ' '&nbsp;</p>\n<p>3. replacing '&deg;' with ' deg'&nbsp;</p>\n<p>to allow answers like 10 ft-lb and 30&deg;</p>", "definition": "replace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")"}, {"name": "good_units", "description": "", "definition": "try(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n"}, {"name": "student_quantity", "description": "<p>This fixes the student answer for two common errors. &nbsp;</p>\n<p>If student_units are wrong&nbsp;&nbsp;- replace with correct units</p>\n<p>If student_scalar has the wrong sign - replace with right sign</p>\n<p>If student makes both errors, only one gets fixed.</p>", "definition": "switch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n"}, {"name": "percent_error", "description": "", "definition": "try(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   "}, {"name": "right", "description": "", "definition": "percent_error <= settings['right']\n"}, {"name": "close", "description": "<p>Only marked close if the student actually has the right sign.</p>", "definition": "right_sign and percent_error <= settings['close']"}, {"name": "right_sign", "description": "", "definition": "sign(student_scalar) = sign(correct_quantity) "}], "settings": [{"name": "correctAnswer", "label": "Correct Quantity.", "help_url": "", "hint": "The correct answer given as a JME quantity.", "input_type": "code", "default_value": "", "evaluate": true}, {"name": "right", "label": "% Accuracy for right.", "help_url": "", "hint": "Question will be considered correct if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "0.2", "evaluate": true}, {"name": "close", "label": "% Accuracy for close.", "help_url": "", "hint": "Question will be considered close if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "1.0", "evaluate": true}, {"name": "C1", "label": "Close with units.", "help_url": "", "hint": "Partial Credit for close value with appropriate units. &nbsp;if correct answer is 100 N and close is &plusmn;1%,<br />99 &nbsp;N is accepted.", "input_type": "percent", "default_value": "75"}, {"name": "C2", "label": "No units or wrong sign", "help_url": "", "hint": "Partial credit for forgetting units or using wrong sign.<br />If the correct answer is 100 N, both 100 and -100 N are accepted.", "input_type": "percent", "default_value": "50"}, {"name": "C3", "label": "Close, no units.", "help_url": "", "hint": "Partial Credit for close value but forgotten units.<br />This value would be close&nbsp;if the expected units were provided. &nbsp;If the correct answer is 100 N, and close is &plusmn;1%,<br />99 is accepted.", "input_type": "percent", "default_value": "25"}], "public_availability": "always", "published": true, "extensions": ["quantities"]}], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Distributed Loads: trapezoidal loading", "tags": ["beam", "distributed load", "Equilibrium", "equilibrium", "Mechanics", "mechanics", "Statics", "statics"], "metadata": {"description": "<p>Find the reactions of a simply-supported beam with a trapazoidal distributed load.</p>", "licence": "Creative Commons Attribution-ShareAlike 4.0 International"}, "statement": "<p>$\\var{distance(L)}$ long beam $AB$ is subjected to a distributed load that varies as shown.</p>\n<p>{loadDiagram}</p>", "advice": "<p></p>\n<ol>\n<li>Determine locations of the centroids and the 'area' under the loading curve for each part.\n<p>$\\begin{align}\\\\ L &amp;= \\var{distance(L)}\\\\W_R &amp;= \\var{force(xr)} &amp; x_R &amp;= \\var{distance(xr)} \\\\<br/>&nbsp;W_T &amp;= \\var{force(wt)} &amp;x_T &amp;= \\var{distance(xt)}\\end{align}$</p>\n</li>\n<li>\n<p>Draw a free body diagram of the beam, replacing the distributed loads with equivalent concentrated loads at the centroids of the triangle and rectangle.</p>\n<p>{fbd}</p>\n</li>\n<li>Find the reactions\n<p>$\\begin{align}\\\\ <br/>\\Sigma M_A &amp;= 0 &nbsp;\\\\<br/>W_R (x_r) + W_T (x_T) &amp;=&nbsp;&nbsp;B (L )\\\\<br/>B &amp;= \\frac{W_R( x_r) + W_T( x_T)}{L} &nbsp;\\\\<br/>&amp;= \\var{force(rb)} \\\\ \\\\ <br/>\\Sigma F_y &amp;= 0 \\\\ A + B &amp;= W_T + W_R \\\\ A &amp;= W_T + W_R - B \\\\ <br/>&amp;= \\var{force(ra)} \\end{align}$</p>\n</li>\n<li>You may check your answer by finding $A$ using $\\Sigma M_B = 0$. &nbsp;You should get the same result either way.</li>\n</ol>\n<p></p>", "rulesets": {}, "extensions": ["jsxgraph", "quantities", "shear-and-bending-moment-diagrams"], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"w": {"name": "w", "group": "Ungrouped variables", "definition": "random(50..150#10)", "description": "<p>Max value of loading &nbsp;in &nbsp;[force/distance]</p>", "templateType": "anything", "can_override": false}, "WT": {"name": "WT", "group": "solution", "definition": "triangle_base w /2", "description": "<p>weight of the triangular load</p>", "templateType": "anything", "can_override": false}, "WR": {"name": "WR", "group": "solution", "definition": "rectangle_base w", "description": "<p>weight of the rectangular load</p>", "templateType": "anything", "can_override": false}, "WTot": {"name": "WTot", "group": "solution", "definition": "WT + WR", "description": "<p>Total weight of load</p>", "templateType": "anything", "can_override": false}, "xR": {"name": "xR", "group": "solution", "definition": "rectangle_base/2 + rectangle[0]", "description": "<p>location of centroid of rect</p>", "templateType": "anything", "can_override": false}, "xT": {"name": "xT", "group": "solution", "definition": "if(type=0,triangle_base/3, 2 triangle_base/3) + triangle[0]\n ", "description": "<p>Centroid of triangle</p>", "templateType": "anything", "can_override": false}, "L": {"name": "L", "group": "Ungrouped variables", "definition": "random([16,20,24])", "description": "<p>length of beam</p>", "templateType": "anything", "can_override": false}, "units": {"name": "units", "group": "Ungrouped variables", "definition": "random([['ft','lb', 'lb/ft'],['m','kN', 'kN/m']])", "description": "", "templateType": "anything", "can_override": false}, "RB": {"name": "RB", "group": "solution", "definition": "(WR xR + WT xT)/ L", "description": "", "templateType": "anything", "can_override": false}, "RA": {"name": "RA", "group": "solution", "definition": "(WT * (L-xT) + WR * (L-xR ))/ L", "description": "", "templateType": "anything", "can_override": false}, "check": {"name": "check", "group": "solution", "definition": "RA + RB = Wtot", "description": "", "templateType": "anything", "can_override": false}, "properties": {"name": "properties", "group": "diagram", "definition": "[beamLength: L,  loads: loads, symbols: symbols ,  forces: [], moments: [], reactions: []]", "description": "", "templateType": "anything", "can_override": false}, "loads": {"name": "loads", "group": "diagram", "definition": "[\n[ [x: rectangle[0], value: w, label: w + \" \" + units[2], visible: true], \n  [x: rectangle[1], value: w, label: \"\" ] ],\n[ [x: triangle[0], value: if(type=0,w,0), label: \"\" , visible: true], \n  [x: triangle[1], value: if(type=0,0,w), label: \"\" ]\n]\n]\n", "description": "<p>choses between triangle-rectangle and rectangle triangle</p>", "templateType": "anything", "can_override": false}, "symbols": {"name": "symbols", "group": "diagram", "definition": "[[type: \"pin\", x: 0, label: \"$A$\", visible: true], [type: \"roller\", x: L, label: \"$B$\", visible: true]]", "description": "", "templateType": "anything", "can_override": false}, "LoadDiagram": {"name": "LoadDiagram", "group": "diagram", "definition": "vmloaddiagram(properties)", "description": "", "templateType": "anything", "can_override": false}, "uniquePoint": {"name": "uniquePoint", "group": "Ungrouped variables", "definition": "sort(shuffle(0..L)[0..3])", "description": "", "templateType": "anything", "can_override": false}, "type": {"name": "type", "group": "Ungrouped variables", "definition": "random(0,1)", "description": "<p>0 &nbsp;= rect then triangle</p>\n<p>1 = triangle then rect</p>", "templateType": "anything", "can_override": false}, "Rectangle": {"name": "Rectangle", "group": "Ungrouped variables", "definition": "if(type=0,\n  [load_start, load_start+ rectangle_base],\n  [load_start+triangle_base, load_start+triangle_base+rectangle_base])", "description": "", "templateType": "anything", "can_override": false}, "triangle": {"name": "triangle", "group": "Ungrouped variables", "definition": "if(type=0,\n  [load_start+rectangle_base, load_start+ rectangle_base + triangle_base],\n  [load_start, load_start+triangle_base])", "description": "", "templateType": "anything", "can_override": false}, "Triangle_base": {"name": "Triangle_base", "group": "Ungrouped variables", "definition": "6", "description": "", "templateType": "anything", "can_override": false}, "rectangle_base": {"name": "rectangle_base", "group": "Ungrouped variables", "definition": "random(4..L-triangle_base)", "description": "", "templateType": "anything", "can_override": false}, "Load_start": {"name": "Load_start", "group": "Ungrouped variables", "definition": "random(0..L-triangle_base-rectangle_base)", "description": "", "templateType": "anything", "can_override": false}, "distance": {"name": "distance", "group": "quantities", "definition": "d -> qty(d, units[0])", "description": "", "templateType": "anything", "can_override": false}, "force": {"name": "force", "group": "quantities", "definition": "f -> siground(qty(f,units[1]),4)", "description": "", "templateType": "anything", "can_override": false}, "load": {"name": "load", "group": "quantities", "definition": "w -> qty(w,units[2])", "description": "", "templateType": "anything", "can_override": false}, "FBDProperties": {"name": "FBDProperties", "group": "diagram", "definition": "[beamLength: L,  reactions: fbdReactions, loads: fbdloads, symbols: fbdsymbols, moments: [], forces: fbdEquivalentloads]", "description": "", "templateType": "anything", "can_override": false}, "FBDSymbols": {"name": "FBDSymbols", "group": "diagram", "definition": "[[type: \"dot\", x: xt, label: xt, visible: true], [type: \"dot\", x: xr, label: xr, visible: true]]", "description": "", "templateType": "anything", "can_override": false}, "fbd": {"name": "fbd", "group": "diagram", "definition": "vmloaddiagram(fbdproperties)", "description": "", "templateType": "anything", "can_override": false}, "fbdReactions": {"name": "fbdReactions", "group": "diagram", "definition": "[[x: 0, value: ra, label: \"$A$\", visible: true], [x: L, value: rb, label: \"$B$\", visible: true]]\n", "description": "", "templateType": "anything", "can_override": false}, "fbdEquivalentLoads": {"name": "fbdEquivalentLoads", "group": "diagram", "definition": "[[x: xr, value: wr, label: \"$W_R$\", visible: true], [x: xt, value: wt, label: \"$W_T$\", visible: true]]", "description": "", "templateType": "anything", "can_override": false}, "fbdLoads": {"name": "fbdLoads", "group": "diagram", "definition": "[\n[ [x: rectangle[0], value: w, label: \"\" , visible: false], \n  [x: rectangle[1], value: w, label: \"\" ] ],\n[ [x: triangle[0], value: if(type=0,w,0), label: \"\" , visible: false], \n  [x: triangle[1], value: if(type=0,0,w), label: \"\" ]\n]\n]\n", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["w", "L", "units", "uniquePoint", "type", "Rectangle", "triangle", "Triangle_base", "rectangle_base", "Load_start"], "variable_groups": [{"name": "solution", "variables": ["WT", "WR", "WTot", "xR", "xT", "RB", "RA", "check"]}, {"name": "quantities", "variables": ["distance", "force", "load"]}, {"name": "diagram", "variables": ["properties", "loads", "symbols", "FBDProperties", "FBDSymbols", "fbdReactions", "fbdEquivalentLoads", "fbdLoads", "fbd", "LoadDiagram"]}], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Determine the reactions at $A$ and $B$.</p>\n<p style=\"padding-left: 40px;\">$A = $&nbsp;[[0]]&nbsp; <br/><br/>$B = $&nbsp;[[1]]&nbsp;&nbsp;</p>", "gaps": [{"type": "engineering-answer", "useCustomName": true, "customName": "$A$", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "force(ra)", "right": "0.2", "close": "1.0", "C1": "80", "C2": "80", "C3": "65"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "$B$", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "force(rb)", "right": "0.2", "close": "1.0", "C1": "80", "C2": "80", "C3": "65"}}], "sortAnswers": false}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}], "resources": []}]}], "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}]}