// Numbas version: finer_feedback_settings
{"name": "Truss: Method of Joints", "extensions": ["geogebra", "quantities"], "custom_part_types": [{"source": {"pk": 19, "author": {"name": "William Haynes", "pk": 2530}, "edit_page": "/part_type/19/edit"}, "name": "Engineering Accuracy with units", "short_name": "engineering-answer", "description": "<p>A value with units marked right if within an adjustable % error of the correct value. &nbsp;Marked close if within a wider margin of error.</p>", "help_url": "", "input_widget": "string", "input_options": {"correctAnswer": "siground(settings['correctAnswer'],4)", "hint": {"static": true, "value": ""}, "allowEmpty": {"static": true, "value": true}}, "can_be_gap": true, "can_be_step": true, "marking_script": "mark:\nswitch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)\n\ninterpreted_answer:\nqty(student_scalar, student_units)\n\n\n\ncorrect_quantity:\nsettings[\"correctAnswer\"]\n\n\n\ncorrect_units:\nunits(correct_quantity)\n\n\nallowed_notation_styles:\n[\"plain\",\"en\"]\n\nmatch_student_number:\nmatchnumber(studentAnswer,allowed_notation_styles)\n\nstudent_scalar:\nmatch_student_number[1]\n\nstudent_units:\nreplace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")\n\ngood_units:\ntry(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n\n\nstudent_quantity:\nswitch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n\n\npercent_error:\ntry(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   \n\nright:\npercent_error <= settings['right']\n\n\nclose:\nright_sign and percent_error <= settings['close']\n\nright_sign:\nsign(student_scalar) = sign(correct_quantity)", "marking_notes": [{"name": "mark", "description": "This is the main marking note. It should award credit and provide feedback based on the student's answer.", "definition": "switch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)"}, {"name": "interpreted_answer", "description": "A value representing the student's answer to this part.", "definition": "qty(student_scalar, student_units)\n\n"}, {"name": "correct_quantity", "description": "", "definition": "settings[\"correctAnswer\"]\n\n"}, {"name": "correct_units", "description": "", "definition": "units(correct_quantity)\n"}, {"name": "allowed_notation_styles", "description": "", "definition": "[\"plain\",\"en\"]"}, {"name": "match_student_number", "description": "", "definition": "matchnumber(studentAnswer,allowed_notation_styles)"}, {"name": "student_scalar", "description": "", "definition": "match_student_number[1]"}, {"name": "student_units", "description": "<p>Modify&nbsp;the&nbsp;unit portion of&nbsp;the student's answer by</p>\n<p>1. replacing \"ohms\" with \"ohm\" &nbsp;case insensitive</p>\n<p>2.&nbsp;replacing '-'&nbsp;with ' '&nbsp;</p>\n<p>3. replacing '&deg;' with ' deg'&nbsp;</p>\n<p>to allow answers like 10 ft-lb and 30&deg;</p>", "definition": "replace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")"}, {"name": "good_units", "description": "", "definition": "try(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n"}, {"name": "student_quantity", "description": "<p>This fixes the student answer for two common errors. &nbsp;</p>\n<p>If student_units are wrong&nbsp;&nbsp;- replace with correct units</p>\n<p>If student_scalar has the wrong sign - replace with right sign</p>\n<p>If student makes both errors, only one gets fixed.</p>", "definition": "switch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n"}, {"name": "percent_error", "description": "", "definition": "try(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   "}, {"name": "right", "description": "", "definition": "percent_error <= settings['right']\n"}, {"name": "close", "description": "<p>Only marked close if the student actually has the right sign.</p>", "definition": "right_sign and percent_error <= settings['close']"}, {"name": "right_sign", "description": "", "definition": "sign(student_scalar) = sign(correct_quantity) "}], "settings": [{"name": "correctAnswer", "label": "Correct Quantity.", "help_url": "", "hint": "The correct answer given as a JME quantity.", "input_type": "code", "default_value": "", "evaluate": true}, {"name": "right", "label": "% Accuracy for right.", "help_url": "", "hint": "Question will be considered correct if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "0.2", "evaluate": true}, {"name": "close", "label": "% Accuracy for close.", "help_url": "", "hint": "Question will be considered close if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "1.0", "evaluate": true}, {"name": "C1", "label": "Close with units.", "help_url": "", "hint": "Partial Credit for close value with appropriate units. &nbsp;if correct answer is 100 N and close is &plusmn;1%,<br />99 &nbsp;N is accepted.", "input_type": "percent", "default_value": "75"}, {"name": "C2", "label": "No units or wrong sign", "help_url": "", "hint": "Partial credit for forgetting units or using wrong sign.<br />If the correct answer is 100 N, both 100 and -100 N are accepted.", "input_type": "percent", "default_value": "50"}, {"name": "C3", "label": "Close, no units.", "help_url": "", "hint": "Partial Credit for close value but forgotten units.<br />This value would be close&nbsp;if the expected units were provided. &nbsp;If the correct answer is 100 N, and close is &plusmn;1%,<br />99 is accepted.", "input_type": "percent", "default_value": "25"}], "public_availability": "always", "published": true, "extensions": ["quantities"]}], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Truss: Method of Joints", "tags": ["equilibrium", "Equilibrium", "Mechanics", "mechanics", "method of joints", "rigid body", "Rigid Body", "Statics", "statics", "truss"], "metadata": {"description": "<p>Find the forces in the members of a truss using the method of joints.</p>", "licence": "Creative Commons Attribution-NonCommercial 4.0 International"}, "statement": "<p>Determine the forces in all members of the truss shown when loaded with two forces: &nbsp;$F_B = \\var{magB}$ and $F_C = \\var{magC}$. &nbsp;</p>\n<table style=\"margin-left: auto; margin-right: auto; width: 61px;\">\n<tbody>\n<tr>\n<td style=\"width: 61.328125px;\">{applet}</td>\n</tr>\n</tbody>\n</table>\n<p>Determine the reactions, and the forces in all five members (indicate tension or compression).</p>", "advice": "<p><strong>Find the reactions:</strong></p>\n<ol>\n<li>Draw a free body diagram of the entire truss.</li>\n<li>Get scalar components of loads $\\mathbf{B}$ and&nbsp;$\\mathbf{C}$&nbsp;for later use.\n<p style=\"padding-left: 30px;\">$\\begin{align}\\\\ B_x &amp;= \\var{display(qty(vecB[0],units))} &amp;&nbsp;B_y &amp;= \\var{display(qty(vecB[1],units))}\\\\C_x &amp;= \\var{display(qty(vecC[0],units))} &amp; C_y &amp;= \\var{display(qty(vecC[1],units))}\\end{align}$</p>\n</li>\n<li>Take moments at $A$&nbsp;to find reactions at&nbsp;$D$. &nbsp;Note 1-2-$\\sqrt{3}$ triangle. &nbsp;\n<p style=\"padding-left: 30px;\">$\\begin{align}\\\\ \\Sigma M_A &amp;= 0\\\\&nbsp; 2\\sqrt{3} \\,D &amp;= \\var{if(sign(vecB[0])=-1,'&ndash; ','')} 1\\, B_x + &nbsp;\\var{if(sign(vecB[1])=1,'&ndash; ','')} \\sqrt{3}\\, B_y +&nbsp;\\sqrt{3}\\,C \\\\&nbsp;&nbsp;D &amp;= &nbsp;\\dfrac{ \\var{if(sign(vecB[0])=-1,'&ndash; ','')} B_x}{2\\sqrt{3}} &nbsp; + \\dfrac{\\var{if(sign(vecB[1])=1,'&ndash; ','')}\\, B_y }{2} + \\dfrac{C }{2}\\\\D&amp;= \\dfrac{1}{2}\\left(\\dfrac{ \\var{precround(vecB[0],1)}&nbsp;}{\\sqrt{3}} + \\var{precround( -&nbsp;vecB[1],1)&nbsp;}&nbsp;+&nbsp;\\var{scalar(magC)} \\right)&amp;= \\var{display(magD)}\\text{ up.}\\end{align}$&nbsp;</p>\n</li>\n<li>Apply $\\Sigma F_x$ and $\\Sigma F_y $ equations to get the components of the reaction force at $A$<em>.</em>\n<p style=\"padding-left: 30px;\">$\\begin{align} \\\\\\Sigma F_x &amp;= 0\\\\ A_x &amp;= B_x \\\\ &amp;= \\var{display(qty(abs(vecA[0]),units))}\\quad \\var{if(sign(vecA[0])=1,'to the right.','to the left.')} \\end{align}$&nbsp;</p>\n<p style=\"padding-left: 30px;\">$\\begin{align} \\\\\\Sigma F_y &amp;= 0\\\\&nbsp;A_y &amp;= C_y -&nbsp;D_y + B_y &nbsp;\\\\ A_y &amp;= \\simplify[!collectNumbers]{{siground(-VecC[1],5)} + {siground(-VecD[1],5)} + {siground(-vecB[1],5)}} \\\\ &amp;= \\var{display(qty(abs(vecA[1]),units))} \\quad \\var{if(sign(vecA[1])=1,'up.','down.')} \\end{align}$&nbsp;</p>\n</li>\n</ol>\n<p><strong>Analyze the joints:</strong></p>\n<ol>\n<li>Draw a FBDs of the joints showing your assumed directions for the forces in the members. &nbsp;Members in tension pull away from the joint, compression push towards the joint.\n<p>&nbsp; &nbsp;{geogebra_applet('menecfet',[['&alpha;',alpha+'&deg;'],['FB','\"'+\"B = \" +&nbsp;string(magB)+'\"'],['FC', '\"'+\"C = \" + string(magC)+'\"'],['FD', '\"'+\"D = \" +&nbsp;display(magD)+'\"'],['FAx', '\"'+\"A_x = \" + display(qty(-vecA[0],units))+'\"'],['FAy', '\"'+\"A_y = \" + display(qty(vecA[1],units))+ '\"']])}</p>\n</li>\n<li>Pick a joint with two unknowns, joint $D$ for example. &nbsp;Joint $A$ would also work.</li>\n<li>Apply equations of equilibrium to joint $D$ to find forces in members $BD$ and $CD$.\n<p style=\"padding-left: 30px;\">$\\begin{align}\\\\D: \\Sigma F_y &nbsp;&amp;= 0\\\\&nbsp;BD_y &amp;= D\\\\&nbsp;BD &amp;= \\dfrac{D}{\\sin 30&deg;}\\\\ &amp;= \\dfrac{\\var{display(magD)}}{0.5}\\\\ &amp;= \\var{display(BD)} \\quad \\text{Compression.}\\end{align}$</p>\n<p style=\"padding-left: 30px;\">$\\begin{align}\\\\D: \\Sigma F_x &amp;= 0\\\\ CD&nbsp;&amp;=BD_x\\\\&nbsp;&amp;= BD&nbsp;\\cos 30&deg; \\\\&amp;= &nbsp; \\var{display(BD)} &nbsp;\\cos 30&deg; \\\\&amp;= \\var{display(CD)} \\quad \\text{Tension.}\\end{align} $</p>\n</li>\n<li>Move to an adjacent joint with two unknowns, joint $C$ in this case. &nbsp;Joint $C$ is a special case.\n<p style=\"padding-left: 30px;\">$\\begin{align}\\\\&nbsp;C: \\Sigma F_y &amp;= 0\\\\BC&nbsp;&amp;= C\\\\ &amp;= \\var{display(BC)} \\quad \\text{Tension.}\\end{align}$&nbsp;</p>\n<p style=\"padding-left: 30px;\">$\\begin{align}\\\\ C: \\Sigma F_x &amp;= 0\\\\AC &amp;= CD\\\\ &amp;= \\var{display(AC)} \\quad \\text{Tension.}\\end{align}$</p>\n</li>\n<li>\n<p>Move to joint $B$ and solve for&nbsp;$AB$.</p>\n<p style=\"padding-left: 30px;\">$B: \\Sigma F_y=0$</p>\n<p style=\"padding-left: 30px;\">$-AB_y &nbsp;-&nbsp;BC + BD_y \\, \\mathbf{\\var{if(sign(vecB[1])=1,'+','&ndash;')}}\\,&nbsp;B_y = 0$</p>\n<p style=\"padding-left: 30px;\">$AB_y=&nbsp;\\var{vecC[1]}&nbsp;+&nbsp;\\var{precround(vecD[1],1)} \\, \\var{if(sign(vecB[1])=-1,'','+')} \\, \\var{precround(vecB[1],1)}= \\var{display(-ABy)}$</p>\n<p style=\"padding-left: 30px;\">$\\begin{align}AB &amp;= \\dfrac{AB_y}{\\sin 30&deg;}\\\\ &amp;= \\var{display(-2ABy)}\\\\&nbsp;&amp;= \\var{display(AB)} \\var{if(sign(ABy)=-1,' Tension.', ' Compression.')}\\end{align}$</p>\n</li>\n</ol>\n<p><strong>Check your work:</strong></p>\n<p>At this point you should check your work by verifying that joint $A$ is in equilibrium. &nbsp;If it is not you have made a mistake somewhere above.</p>\n<p style=\"padding-left: 30px;\">$A: \\Sigma F_x \\stackrel{?}{=}&nbsp; 0$</p>\n<p style=\"padding-left: 30px;\">$-A_x + AB_x + AC =&nbsp;\\var{precround(vecA[0],1)}+ \\var{precround(-vecAB[0],1) } + \\var{precround(-vecCD[0],1) } = &nbsp;\\var{vecA[0] -vecAB[0] - vecCD[0]}$</p>\n<p style=\"padding-left: 30px;\">{if((vecA[0] -vecAB[0] - vecCD[0]) = 0,'Check!', if((vecA[0] -vecAB[0] - vecCD[0]) &lt; 0.00001, 'Close Enough', 'Something is wrong.'))}</p>\n<p style=\"padding-left: 30px;\">$A: \\Sigma F_y \\stackrel{?}{=}&nbsp; 0$</p>\n<p style=\"padding-left: 30px;\">$A_y + AB_y&nbsp;=&nbsp;\\var{precround(vecA[1],1)}+ \\var{precround(-vecAB[1],1) } = \\var{vecA[1] -vecAB[1]}$</p>\n<p style=\"padding-left: 30px;\">{if((vecA[1] -vecAB[1]) = 0,'Check!', if((vecA[1] -vecAB[1]) &lt; 0.00001, 'Close Enough', 'Something is wrong.'))}</p>", "rulesets": {}, "extensions": ["geogebra", "quantities"], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"ABy": {"name": "ABy", "group": "force in members", "definition": "qty(vecAB[1],units)", "description": "", "templateType": "anything", "can_override": false}, "vecBD": {"name": "vecBD", "group": "force in members", "definition": "siground(scalar(BD) vector(sqrt(3/2),-0.5),5)", "description": "", "templateType": "anything", "can_override": false}, "r_ab": {"name": "r_ab", "group": "vectors", "definition": "vector(sqrt(3),1,0)", "description": "", "templateType": "anything", "can_override": false}, "vecD": {"name": "vecD", "group": "vectors", "definition": "scalar(magD) vector(0,1,0)", "description": "", "templateType": "anything", "can_override": false}, "BC": {"name": "BC", "group": "force in members", "definition": "magC", "description": "", "templateType": "anything", "can_override": false}, "units": {"name": "units", "group": "Inputs", "definition": "random('N', 'kN', 'lb')", "description": "", "templateType": "anything", "can_override": false}, "vecC": {"name": "vecC", "group": "vectors", "definition": "scalar(magC)vector(0,-1,0)", "description": "", "templateType": "anything", "can_override": false}, "debug": {"name": "debug", "group": "Inputs", "definition": "false", "description": "", "templateType": "anything", "can_override": false}, "magAx": {"name": "magAx", "group": "quantities", "definition": "qty(abs(vecA[0]),units)", "description": "", "templateType": "anything", "can_override": false}, "magAy": {"name": "magAy", "group": "quantities", "definition": "qty(abs(vecA[1]),units)", "description": "", "templateType": "anything", "can_override": false}, "vecAB": {"name": "vecAB", "group": "force in members", "definition": "-(vecB + vecC + scalar(BD) vector(-cos(radians(30)),sin(radians(30))))", "description": "", "templateType": "anything", "can_override": false}, "vecB": {"name": "vecB", "group": "vectors", "definition": "scalar(magB) vector(cos(radians(alpha)),sin(radians(alpha)),0)", "description": "", "templateType": "anything", "can_override": false}, "AB": {"name": "AB", "group": "force in members", "definition": "qty(abs(vecAB),units)", "description": "", "templateType": "anything", "can_override": false}, "CD": {"name": "CD", "group": "force in members", "definition": "BD cos(radians(30))", "description": "", "templateType": "anything", "can_override": false}, "r_ad": {"name": "r_ad", "group": "vectors", "definition": "2 r_ac", "description": "", "templateType": "anything", "can_override": false}, "r_ac": {"name": "r_ac", "group": "vectors", "definition": "vector(sqrt(3),0,0)", "description": "", "templateType": "anything", "can_override": false}, "magC": {"name": "magC", "group": "Inputs", "definition": "qty(random(100..800#25),units)", "description": "<p>Magnitude of force C, always points straight down</p>", "templateType": "anything", "can_override": false}, "BD": {"name": "BD", "group": "force in members", "definition": " magD/sin(radians(30)) ", "description": "", "templateType": "anything", "can_override": false}, "magD": {"name": "magD", "group": "quantities", "definition": "qty(-(cross(r_ab,vecB)+cross(r_ac, vecC))[2]/abs(r_ad),units)", "description": "", "templateType": "anything", "can_override": false}, "AC": {"name": "AC", "group": "force in members", "definition": "CD", "description": "", "templateType": "anything", "can_override": false}, "Abs": {"name": "Abs", "group": "force in members", "definition": "qty(vecAB[0],units)", "description": "", "templateType": "anything", "can_override": false}, "check": {"name": "check", "group": "vectors", "definition": "vecA+vecB+vecC+vecD", "description": "", "templateType": "anything", "can_override": false}, "magB": {"name": "magB", "group": "Inputs", "definition": "qty(random(100..800#25),units)", "description": "", "templateType": "anything", "can_override": false}, "alpha": {"name": "alpha", "group": "Inputs", "definition": "random(-110..0#10 except -90) ", "description": "<p>direction of force B &nbsp;- standard angle</p>", "templateType": "anything", "can_override": false}, "vecA": {"name": "vecA", "group": "vectors", "definition": "-(vecB+VecC+vecD)", "description": "", "templateType": "anything", "can_override": false}, "vecCD": {"name": "vecCD", "group": "force in members", "definition": "scalar(CD) vector(-1,0)", "description": "", "templateType": "anything", "can_override": false}, "applet": {"name": "applet", "group": "Inputs", "definition": "geogebra_applet('h7m3fvsj',params)", "description": "", "templateType": "anything", "can_override": false}, "params": {"name": "params", "group": "Inputs", "definition": "['\u03b1': radians(alpha),\nF_B: scalar(MagB),\u00a0\nF_C: scalar(MagC),\nunits: '\"{units}\"' \n]", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "not (alpha >= 30 and alpha <= 150) and // keep force outside the truss\nscalar(magD) > 0 // so roller doesn't need to pull down", "maxRuns": 100}, "ungrouped_variables": [], "variable_groups": [{"name": "Inputs", "variables": ["magC", "magB", "units", "alpha", "debug", "applet", "params"]}, {"name": "vectors", "variables": ["vecB", "vecC", "r_ab", "r_ac", "r_ad", "vecD", "vecA", "check"]}, {"name": "quantities", "variables": ["magD", "magAx", "magAy"]}, {"name": "force in members", "variables": ["BD", "CD", "BC", "AC", "vecAB", "AB", "ABy", "Abs", "vecCD", "vecBD"]}], "functions": {"display": {"parameters": [["q", "quantity"]], "type": "string", "language": "jme", "definition": "string(precround(q,2))"}}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": true, "customName": "Reactions", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Determine the components of the reactions at $A$ and $D$.</p>\n<p style=\"padding-left: 40px;\">$A_x$ =&nbsp;[[0]]&nbsp;[[1]]</p>\n<p style=\"padding-left: 40px;\">$A_y$ =&nbsp;[[2]]&nbsp;[[3]]</p>\n<p style=\"padding-left: 40px;\">$D_x$ =&nbsp;[[4]]&nbsp;[[5]]</p>\n<p style=\"padding-left: 40px;\">$D_y$ = [[6]]&nbsp;[[7]]</p>\n<p><span data-jme-visible=\"debug\">{precround(vecA,1)} {units}</span> <span data-jme-visible=\"debug\">{vecD}</span></p>", "gaps": [{"type": "engineering-answer", "useCustomName": false, "customName": "", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "magAx", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "choices": ["Left", "Right", "Neither"], "matrix": "map(if(sign(vecA[0])=s,1,0),s,[-1,1,0])"}, {"type": "engineering-answer", "useCustomName": false, "customName": "", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "magAy", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "choices": ["Up", "Down", "Neither"], "matrix": "map(if(sign(vecA[1])=s,1,0),s,[1,-1,0])"}, {"type": "engineering-answer", "useCustomName": false, "customName": "", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(0,units)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "choices": ["Up", "Down", "Neither"], "matrix": "[0,0,1]"}, {"type": "engineering-answer", "useCustomName": false, "customName": "", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "magD", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "choices": ["Up", "Down", "Neither"], "matrix": "map(if(sign(vecD[1])=s,1,0),s,[1,-1,0])"}], "sortAnswers": false}, {"type": "gapfill", "useCustomName": true, "customName": "Joint D", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Analyze joint&nbsp;$D$ to find the forces in members $BD$ and&nbsp;$CD$.</p>\n<p style=\"padding-left: 40px;\">$BD$ =&nbsp;[[0]]&nbsp;[[1]]&nbsp;</p>\n<p style=\"padding-left: 40px;\">$CD$ =&nbsp;[[2]]&nbsp;[[3]]&nbsp;</p>", "gaps": [{"type": "engineering-answer", "useCustomName": false, "customName": "", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "BD", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "choices": ["Tension", "Compression", "Neither"], "matrix": [0, "1", "if(scalar(BD)=0,1,0)"], "distractors": ["", "", ""]}, {"type": "engineering-answer", "useCustomName": false, "customName": "", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "CD", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "choices": ["Tension", "Compression", "Neither"], "matrix": ["1", "0", "if(scalar(CD)=0,1,0)"], "distractors": ["", "", ""]}], "sortAnswers": false}, {"type": "gapfill", "useCustomName": true, "customName": "Joint C", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Analyze joint $C$ to find the forces in members $AC$&nbsp;and $BC$.</p>\n<p style=\"padding-left: 40px;\">$AC$ =&nbsp;[[0]]&nbsp;[[1]]&nbsp;</p>\n<p style=\"padding-left: 40px;\">$BC$ =&nbsp;[[2]]&nbsp;[[3]]&nbsp;</p>\n<p><span data-jme-visible=\"debug\">{display(AC)} {display(BC)}</span></p>", "gaps": [{"type": "engineering-answer", "useCustomName": false, "customName": "", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "AC", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "choices": ["Tension", "Compression", "Neither"], "matrix": ["1", "0", "if(scalar(AC)=0,1,0)"], "distractors": ["", "", ""]}, {"type": "engineering-answer", "useCustomName": false, "customName": "", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "BC", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "choices": ["Tension", "Compression", "Neither"], "matrix": ["1", "0", "if(scalar(BC)=0,1,0)"], "distractors": ["", "", ""]}], "sortAnswers": false}, {"type": "gapfill", "useCustomName": true, "customName": "Joint B", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Analyze joint $B$ to find the force in member $AB$.</p>\n<p style=\"padding-left: 40px;\">$AB$ =&nbsp;[[0]]&nbsp;[[1]]&nbsp; <span data-jme-visible=\"debug\">{display(AB)} {if(sign(ABy)=-1,' Tension.', ' Compression.')}</span></p>\n<p></p>", "gaps": [{"type": "engineering-answer", "useCustomName": false, "customName": "", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "AB", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "choices": ["Tension", "Compression", "Neither"], "matrix": ["if(sign(ABy)=-1,1,0)", "if(sign(ABy)=1,1,0)", "if(sign(ABy)=0,1,0)"], "distractors": ["", "", ""]}], "sortAnswers": false}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}], "resources": []}]}], "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}]}