// Numbas version: finer_feedback_settings
{"name": "Moment of inertia of a composite shape: U or T", "extensions": ["geogebra", "quantities"], "custom_part_types": [{"source": {"pk": 19, "author": {"name": "William Haynes", "pk": 2530}, "edit_page": "/part_type/19/edit"}, "name": "Engineering Accuracy with units", "short_name": "engineering-answer", "description": "<p>A value with units marked right if within an adjustable % error of the correct value. &nbsp;Marked close if within a wider margin of error.</p>", "help_url": "", "input_widget": "string", "input_options": {"correctAnswer": "siground(settings['correctAnswer'],4)", "hint": {"static": true, "value": ""}, "allowEmpty": {"static": true, "value": true}}, "can_be_gap": true, "can_be_step": true, "marking_script": "mark:\nswitch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)\n\ninterpreted_answer:\nqty(student_scalar, student_units)\n\n\n\ncorrect_quantity:\nsettings[\"correctAnswer\"]\n\n\n\ncorrect_units:\nunits(correct_quantity)\n\n\nallowed_notation_styles:\n[\"plain\",\"en\"]\n\nmatch_student_number:\nmatchnumber(studentAnswer,allowed_notation_styles)\n\nstudent_scalar:\nmatch_student_number[1]\n\nstudent_units:\nreplace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")\n\ngood_units:\ntry(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n\n\nstudent_quantity:\nswitch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n\n\npercent_error:\ntry(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   \n\nright:\npercent_error <= settings['right']\n\n\nclose:\nright_sign and percent_error <= settings['close']\n\nright_sign:\nsign(student_scalar) = sign(correct_quantity)", "marking_notes": [{"name": "mark", "description": "This is the main marking note. It should award credit and provide feedback based on the student's answer.", "definition": "switch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)"}, {"name": "interpreted_answer", "description": "A value representing the student's answer to this part.", "definition": "qty(student_scalar, student_units)\n\n"}, {"name": "correct_quantity", "description": "", "definition": "settings[\"correctAnswer\"]\n\n"}, {"name": "correct_units", "description": "", "definition": "units(correct_quantity)\n"}, {"name": "allowed_notation_styles", "description": "", "definition": "[\"plain\",\"en\"]"}, {"name": "match_student_number", "description": "", "definition": "matchnumber(studentAnswer,allowed_notation_styles)"}, {"name": "student_scalar", "description": "", "definition": "match_student_number[1]"}, {"name": "student_units", "description": "<p>Modify&nbsp;the&nbsp;unit portion of&nbsp;the student's answer by</p>\n<p>1. replacing \"ohms\" with \"ohm\" &nbsp;case insensitive</p>\n<p>2.&nbsp;replacing '-'&nbsp;with ' '&nbsp;</p>\n<p>3. replacing '&deg;' with ' deg'&nbsp;</p>\n<p>to allow answers like 10 ft-lb and 30&deg;</p>", "definition": "replace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")"}, {"name": "good_units", "description": "", "definition": "try(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n"}, {"name": "student_quantity", "description": "<p>This fixes the student answer for two common errors. &nbsp;</p>\n<p>If student_units are wrong&nbsp;&nbsp;- replace with correct units</p>\n<p>If student_scalar has the wrong sign - replace with right sign</p>\n<p>If student makes both errors, only one gets fixed.</p>", "definition": "switch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n"}, {"name": "percent_error", "description": "", "definition": "try(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   "}, {"name": "right", "description": "", "definition": "percent_error <= settings['right']\n"}, {"name": "close", "description": "<p>Only marked close if the student actually has the right sign.</p>", "definition": "right_sign and percent_error <= settings['close']"}, {"name": "right_sign", "description": "", "definition": "sign(student_scalar) = sign(correct_quantity) "}], "settings": [{"name": "correctAnswer", "label": "Correct Quantity.", "help_url": "", "hint": "The correct answer given as a JME quantity.", "input_type": "code", "default_value": "", "evaluate": true}, {"name": "right", "label": "% Accuracy for right.", "help_url": "", "hint": "Question will be considered correct if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "0.2", "evaluate": true}, {"name": "close", "label": "% Accuracy for close.", "help_url": "", "hint": "Question will be considered close if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "1.0", "evaluate": true}, {"name": "C1", "label": "Close with units.", "help_url": "", "hint": "Partial Credit for close value with appropriate units. &nbsp;if correct answer is 100 N and close is &plusmn;1%,<br />99 &nbsp;N is accepted.", "input_type": "percent", "default_value": "75"}, {"name": "C2", "label": "No units or wrong sign", "help_url": "", "hint": "Partial credit for forgetting units or using wrong sign.<br />If the correct answer is 100 N, both 100 and -100 N are accepted.", "input_type": "percent", "default_value": "50"}, {"name": "C3", "label": "Close, no units.", "help_url": "", "hint": "Partial Credit for close value but forgotten units.<br />This value would be close&nbsp;if the expected units were provided. &nbsp;If the correct answer is 100 N, and close is &plusmn;1%,<br />99 is accepted.", "input_type": "percent", "default_value": "25"}], "public_availability": "always", "published": true, "extensions": ["quantities"]}], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Moment of inertia of a composite shape: U or T", "tags": ["mechanics", "Mechanics", "moment of inerta", "parallel axis theorem", "second moment of area", "Statics", "statics"], "metadata": {"description": "<p>Calculate the moment of inertia of a composite shape consiting of two rectangles about the <em>x</em> or <em>y</em>-axis. &nbsp;Parallel axis theorem is often required.</p>", "licence": "Creative Commons Attribution-NonCommercial 4.0 International"}, "statement": "<p>A composite shape is drawn to scale below. &nbsp;Grid units are [{units}]. Determine the total area of the shape, &nbsp;and the moment of inertia and radius of gyration with respect to the <em>{axis}</em>-axis.</p>\n<table style=\"margin-left: auto; margin-right: auto; width: 397px;\" height=\"48\">\n<tbody>\n<tr>\n<td style=\"width: 380px;\">{geogebra_applet('ez5xyd5w',[['A',A],['B',B]])}</td>\n</tr>\n</tbody>\n</table>\n<div data-jme-visible=\"debug\">\n<p><strong>Dimensions:</strong></p>\n<p>Rectangle 1: $b_1 = \\var{b1} \\qquad h_1 = \\var{h1}$</p>\n<p>Rectangle 2: &nbsp;$b_2 = \\var{b2}\\qquad h_2 = \\var{h2}$&nbsp;<strong>{if(B[1]&lt;A[1],'Removed','')}</strong></p>\n<p><strong>Area:</strong></p>\n<p>$A_1 = b_1 h_1 = \\var{A1}$ <strong>&nbsp;</strong></p>\n<p>$A_2 = b_2 h_2 = \\var{A2}$<strong>&nbsp;{if(B[1]&lt;A[1],'Removed','')}</strong></p>\n<p>$A_t = A_1 + A_2 = \\var{A1} + &nbsp;\\var{if(B[1]&lt;A[1],'&ndash;','')} \\var{A2} = \\var{A_t}$</p>\n<p><strong>For the <em>x</em>-axis:</strong></p>\n<p>$I_{1x} = \\frac{1}{3}b_1 h_1^3 = \\var{siground(I1x,4)}$</p>\n<p>$I_{2x} = \\bar{I} + A d^2 =&nbsp;&nbsp;\\frac{1}{12} (\\var{B2}) (\\var{H2})^3 +(\\var{A2}) (\\var{dx})^2&nbsp;= \\var{siground(I2x,4)}&nbsp;$ &nbsp;<strong>{if(B[1]&lt;A[1],'Removed','')}</strong></p>\n<p>$I_x = I_{1x} +&nbsp;I_{2x} = \\var{siground(I1x,4)} +&nbsp;&nbsp;\\var{if(B[1]&lt;A[1],'&ndash;','')} &nbsp;\\var{siground(I2x,4)} = \\var{siground(I_tx,4)}$</p>\n<p>$k_x = \\sqrt{\\frac{I_x}{A}} &nbsp;= \\var{siground(kx,4)}$</p>\n<p><strong>For the centroidal&nbsp;<em>y'</em>-axis:</strong></p>\n<p>$I_{1y'} = \\frac{1}{12}h_1 b_1^3 &nbsp;= \\var{siground(I1y,4)}$</p>\n<p>$I_{2y'} = \\frac{1}{12}h_2 b_2^3 &nbsp;=&nbsp;\\var{siground(I2y,4)} $ <strong>{if(B[1]&lt;A[1],'Removed','')}</strong></p>\n<p>$\\bar{I}_{y'} = I_{1y'} + I_{2y'} = \\var{siground(I1y,4)} + \\var{if(B[1]&lt;A[1],'&ndash;','')} \\var{siground(I2y,4)} = \\var{siground(Ibar_y',4)}$</p>\n<p><strong>For the <em>y</em>-axis, using the parallel axis theorem:</strong></p>\n<p>$I_y = \\bar{I}_{y'}+ A d^2 = \\var{siground(Ibar_y',4)} + (\\var{A_t}) (\\var{dy})^2 = &nbsp;\\var{siground(I_ty,4)}$</p>\n<p>$k_y = \\sqrt{\\frac{I_y}{A}} &nbsp;= \\var{siground(ky,4)}$</p>\n</div>", "advice": "<p><strong>Dimensions:</strong></p>\n<p style=\"padding-left: 40px;\">$b_1 = \\var{b1} \\qquad h_1 = \\var{h1}\\\\b_2 = \\var{b2}\\qquad h_2 = \\var{h2}$</p>\n<p><strong>Area:</strong></p>\n<p style=\"padding-left: 40px;\">$A_1 = b_1 h_1 = \\var{A1}$ <strong>&nbsp;</strong></p>\n<p style=\"padding-left: 40px;\">$A_2 = b_2 h_2 = \\var{A2}$<strong>&nbsp;{if(B[1]&lt;A[1],'Removed','')}</strong></p>\n<p style=\"padding-left: 40px;\">$A_t = A_1 + A_2 = \\var{A1} + &nbsp;\\var{if(B[1]&lt;A[1],'&ndash;','')} \\var{A2} = \\var{A_t}$</p>\n<p></p>\n<div data-jme-visible=\"axis='x'\">\n<p><strong>For the <em>x</em>-axis:</strong></p>\n<p style=\"padding-left: 40px;\">$I_{1x} = \\frac{1}{3}b_1 h_1^3 = \\var{siground(I1x,4)}$</p>\n<p style=\"padding-left: 40px;\">$I_{2x} = \\bar{I} + A d^2 =&nbsp;&nbsp;\\frac{1}{12} (\\var{B2}) (\\var{H2})^3 +(\\var{A2}) (\\var{dx})^2&nbsp;= \\var{siground(I2x,4)}&nbsp;$ &nbsp;<strong>{if(B[1]&lt;A[1],'Removed','')}</strong></p>\n<p style=\"padding-left: 40px;\">$I_x = I_{1x} +&nbsp;I_{2x} = \\var{siground(I1x,4)} +&nbsp;&nbsp;\\var{if(B[1]&lt;A[1],'&ndash;','')} &nbsp;\\var{siground(I2x,4)} = \\var{siground(I_tx,4)}$</p>\n<p style=\"padding-left: 40px;\">$k_x = \\sqrt{\\frac{I_x}{A}} &nbsp;= \\var{siground(kx,4)}$</p>\n</div>\n<div data-jme-visible=\"axis='y'\">\n<p><strong>For the centroidal&nbsp;<em>y'</em>-axis:</strong></p>\n<p style=\"padding-left: 40px;\">$I_{1y'} = \\frac{1}{12}h_1 b_1^3 &nbsp;= \\var{siground(I1y,4)}$</p>\n<p style=\"padding-left: 40px;\">$I_{2y'} = \\frac{1}{12}h_2 b_2^3 &nbsp;=&nbsp;\\var{siground(I2y,4)} $ <strong>{if(B[1]&lt;A[1],'Removed','')}</strong></p>\n<p style=\"padding-left: 40px;\">$\\bar{I}_{y'} = I_{1y'} + I_{2y'} = \\var{siground(I1y,4)} + \\var{if(B[1]&lt;A[1],'&ndash;','')} \\var{siground(I2y,4)} = \\var{siground(Ibar_y',4)}$</p>\n<p><strong>For the <em>y</em>-axis, using the parallel axis theorem:</strong></p>\n<p style=\"padding-left: 40px;\">$I_y = \\bar{I}_{y'}+ A d^2 = \\var{siground(Ibar_y',4)} + (\\var{A_t}) (\\var{dy})^2 = &nbsp;\\var{siground(I_ty,4)}$</p>\n<p style=\"padding-left: 40px;\">$k_y = \\sqrt{\\frac{I_y}{A}} &nbsp;= \\var{siground(ky,4)}$</p>\n</div>", "rulesets": {}, "extensions": ["geogebra", "quantities"], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"A": {"name": "A", "group": "inputs", "definition": "vector(random(0.5..6#0.5),random(1..9))", "description": "<p>This point is at the half width of rectangle 1</p>", "templateType": "anything", "can_override": false}, "b1": {"name": "b1", "group": "Ungrouped variables", "definition": "qty(2 A[0],units)", "description": "<p>base of rectangle 1</p>", "templateType": "anything", "can_override": false}, "I2y": {"name": "I2y", "group": "Ungrouped variables", "definition": " h2 b2 b2 b2/12", "description": "", "templateType": "anything", "can_override": false}, "units": {"name": "units", "group": "inputs", "definition": "random('in','ft','mm','cm','m')", "description": "", "templateType": "anything", "can_override": false}, "debug": {"name": "debug", "group": "inputs", "definition": "false", "description": "", "templateType": "anything", "can_override": false}, "A_t": {"name": "A_t", "group": "Ungrouped variables", "definition": "A1+ if(B[1]>A[1],1,-1) A2", "description": "", "templateType": "anything", "can_override": false}, "b2": {"name": "b2", "group": "Ungrouped variables", "definition": "qty(2 B[0],units)\n", "description": "<p>base of rectangle 2</p>", "templateType": "anything", "can_override": false}, "dy": {"name": "dy", "group": "Ungrouped variables", "definition": "qty(max(A[0],B[0]),units)", "description": "<p>half width of shape - distance from centroid to y axis.</p>", "templateType": "anything", "can_override": false}, "ky": {"name": "ky", "group": "Ungrouped variables", "definition": "qty(sqrt(scalar(I_ty/A_t)),units)", "description": "", "templateType": "anything", "can_override": false}, "I_tx": {"name": "I_tx", "group": "Ungrouped variables", "definition": "I1x + if(B[1]>A[1],1,-1) I2x", "description": "<p>total moment of inertia wrt</p>", "templateType": "anything", "can_override": false}, "Ibar_y'": {"name": "Ibar_y'", "group": "Ungrouped variables", "definition": "I1y +  if(B[1]>A[1],1,-1) I2y", "description": "", "templateType": "anything", "can_override": false}, "h2": {"name": "h2", "group": "Ungrouped variables", "definition": "qty(abs(A[1]-B[1]),units)", "description": "<p>height of rectangle 2</p>", "templateType": "anything", "can_override": false}, "I_ty": {"name": "I_ty", "group": "Ungrouped variables", "definition": "( I1y + if(B[1]>A[1],1,-1) I2y ) + A_t dy dy", "description": "<p>I_bar_y + A&nbsp;d^2</p>", "templateType": "anything", "can_override": false}, "dx": {"name": "dx", "group": "Ungrouped variables", "definition": "qty((A[1]+B[1])/2,units)", "description": "<p>distance to centroid of rectangle 2 for parallel axis theorem</p>", "templateType": "anything", "can_override": false}, "kx": {"name": "kx", "group": "Ungrouped variables", "definition": "qty(sqrt(scalar(I_tx/A_t)),units)", "description": "", "templateType": "anything", "can_override": false}, "B": {"name": "B", "group": "inputs", "definition": "vector(random(0.5..6#0.5),random(1..9))", "description": "<p>This point is at the half width of rectangle 2.</p>", "templateType": "anything", "can_override": false}, "A1": {"name": "A1", "group": "Ungrouped variables", "definition": "B1 H1", "description": "<p>area of rectangle 1</p>", "templateType": "anything", "can_override": false}, "I2x": {"name": "I2x", "group": "Ungrouped variables", "definition": "(b2 * h2 * h2 * h2)/12 + A2 dx dx\n", "description": "<p>distance up to centroid of rectangle 2</p>", "templateType": "anything", "can_override": false}, "axis": {"name": "axis", "group": "inputs", "definition": "random('x','y')", "description": "", "templateType": "anything", "can_override": false}, "h1": {"name": "h1", "group": "Ungrouped variables", "definition": "qty(A[1],units)", "description": "<p>height of rectangle 1</p>\n<p></p>", "templateType": "anything", "can_override": false}, "I1x": {"name": "I1x", "group": "Ungrouped variables", "definition": "b1 * h1 * h1 * h1/3", "description": "", "templateType": "anything", "can_override": false}, "I1y": {"name": "I1y", "group": "Ungrouped variables", "definition": "(h1 b1 b1 b1)/12", "description": "", "templateType": "anything", "can_override": false}, "A2": {"name": "A2", "group": "Ungrouped variables", "definition": "B2 H2 ", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "not(B[0]>A[0] and B[1]<A[1]) and \nnot(B[0]=A[0] or B[1] = A[1]) and\nmod(2 A[0],2) = mod(2 B[0],2)", "maxRuns": 100}, "ungrouped_variables": ["b1", "h1", "b2", "h2", "I1x", "I1y", "I2x", "I2y", "I_tx", "I_ty", "dx", "dy", "A1", "A2", "A_t", "kx", "ky", "Ibar_y'"], "variable_groups": [{"name": "inputs", "variables": ["A", "B", "units", "debug", "axis"]}], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p style=\"padding-left: 40px;\">$A$ =&nbsp;[[0]]&nbsp;&nbsp; <span data-jme-visible=\"debug\">{A_t}</span></p>\n<p style=\"padding-left: 40px;\">$I_\\var{axis} =$&nbsp;[[1]] &nbsp;<span data-jme-visible=\"debug\">{siground(if(axis='x',I_tx,I_ty),4)}</span></p>\n<p style=\"padding-left: 40px;\">$k_\\var{axis} =$&nbsp;[[2]]&nbsp;<span data-jme-visible=\"debug\"> {siground(if(axis='x',kx,ky),4)}</span></p>", "gaps": [{"type": "engineering-answer", "useCustomName": true, "customName": "A", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "abs(A_t)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "$I_\\var{axis}$", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "abs(A_t)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "$k_\\var{axis}$", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "abs(A_t)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}], "sortAnswers": false}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}], "resources": []}]}], "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}]}