// Numbas version: finer_feedback_settings
{"name": "Method of Joints: Symmetric Truss", "extensions": ["geogebra", "quantities"], "custom_part_types": [{"source": {"pk": 19, "author": {"name": "William Haynes", "pk": 2530}, "edit_page": "/part_type/19/edit"}, "name": "Engineering Accuracy with units", "short_name": "engineering-answer", "description": "<p>A value with units marked right if within an adjustable % error of the correct value. &nbsp;Marked close if within a wider margin of error.</p>", "help_url": "", "input_widget": "string", "input_options": {"correctAnswer": "siground(settings['correctAnswer'],4)", "hint": {"static": true, "value": ""}, "allowEmpty": {"static": true, "value": true}}, "can_be_gap": true, "can_be_step": true, "marking_script": "mark:\nswitch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)\n\ninterpreted_answer:\nqty(student_scalar, student_units)\n\n\n\ncorrect_quantity:\nsettings[\"correctAnswer\"]\n\n\n\ncorrect_units:\nunits(correct_quantity)\n\n\nallowed_notation_styles:\n[\"plain\",\"en\"]\n\nmatch_student_number:\nmatchnumber(studentAnswer,allowed_notation_styles)\n\nstudent_scalar:\nmatch_student_number[1]\n\nstudent_units:\nreplace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")\n\ngood_units:\ntry(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n\n\nstudent_quantity:\nswitch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n\n\npercent_error:\ntry(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   \n\nright:\npercent_error <= settings['right']\n\n\nclose:\nright_sign and percent_error <= settings['close']\n\nright_sign:\nsign(student_scalar) = sign(correct_quantity)", "marking_notes": [{"name": "mark", "description": "This is the main marking note. It should award credit and provide feedback based on the student's answer.", "definition": "switch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)"}, {"name": "interpreted_answer", "description": "A value representing the student's answer to this part.", "definition": "qty(student_scalar, student_units)\n\n"}, {"name": "correct_quantity", "description": "", "definition": "settings[\"correctAnswer\"]\n\n"}, {"name": "correct_units", "description": "", "definition": "units(correct_quantity)\n"}, {"name": "allowed_notation_styles", "description": "", "definition": "[\"plain\",\"en\"]"}, {"name": "match_student_number", "description": "", "definition": "matchnumber(studentAnswer,allowed_notation_styles)"}, {"name": "student_scalar", "description": "", "definition": "match_student_number[1]"}, {"name": "student_units", "description": "<p>Modify&nbsp;the&nbsp;unit portion of&nbsp;the student's answer by</p>\n<p>1. replacing \"ohms\" with \"ohm\" &nbsp;case insensitive</p>\n<p>2.&nbsp;replacing '-'&nbsp;with ' '&nbsp;</p>\n<p>3. replacing '&deg;' with ' deg'&nbsp;</p>\n<p>to allow answers like 10 ft-lb and 30&deg;</p>", "definition": "replace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")"}, {"name": "good_units", "description": "", "definition": "try(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n"}, {"name": "student_quantity", "description": "<p>This fixes the student answer for two common errors. &nbsp;</p>\n<p>If student_units are wrong&nbsp;&nbsp;- replace with correct units</p>\n<p>If student_scalar has the wrong sign - replace with right sign</p>\n<p>If student makes both errors, only one gets fixed.</p>", "definition": "switch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n"}, {"name": "percent_error", "description": "", "definition": "try(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   "}, {"name": "right", "description": "", "definition": "percent_error <= settings['right']\n"}, {"name": "close", "description": "<p>Only marked close if the student actually has the right sign.</p>", "definition": "right_sign and percent_error <= settings['close']"}, {"name": "right_sign", "description": "", "definition": "sign(student_scalar) = sign(correct_quantity) "}], "settings": [{"name": "correctAnswer", "label": "Correct Quantity.", "help_url": "", "hint": "The correct answer given as a JME quantity.", "input_type": "code", "default_value": "", "evaluate": true}, {"name": "right", "label": "% Accuracy for right.", "help_url": "", "hint": "Question will be considered correct if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "0.2", "evaluate": true}, {"name": "close", "label": "% Accuracy for close.", "help_url": "", "hint": "Question will be considered close if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "1.0", "evaluate": true}, {"name": "C1", "label": "Close with units.", "help_url": "", "hint": "Partial Credit for close value with appropriate units. &nbsp;if correct answer is 100 N and close is &plusmn;1%,<br />99 &nbsp;N is accepted.", "input_type": "percent", "default_value": "75"}, {"name": "C2", "label": "No units or wrong sign", "help_url": "", "hint": "Partial credit for forgetting units or using wrong sign.<br />If the correct answer is 100 N, both 100 and -100 N are accepted.", "input_type": "percent", "default_value": "50"}, {"name": "C3", "label": "Close, no units.", "help_url": "", "hint": "Partial Credit for close value but forgotten units.<br />This value would be close&nbsp;if the expected units were provided. &nbsp;If the correct answer is 100 N, and close is &plusmn;1%,<br />99 is accepted.", "input_type": "percent", "default_value": "25"}], "public_availability": "always", "published": true, "extensions": ["quantities"]}], "resources": [["question-resources/joan-problem_fs07hWN.ggb", "joan-problem_fs07hWN.ggb"], ["question-resources/Joint_B_6eszcTT.svg", "Joint_B_6eszcTT.svg"], ["question-resources/Joint_D_EOoCf9Z.svg", "Joint_D_EOoCf9Z.svg"], ["question-resources/fbds_jqa9zTW_0O5KNSm.afdesign", "fbds_jqa9zTW_0O5KNSm.afdesign"], ["question-resources/Joint_A_j2xOQ6P.svg", "Joint_A_j2xOQ6P.svg"], ["question-resources/joans-problem.ggb", "joans-problem_OHPlYhg.ggb"]], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Method of Joints: Symmetric Truss", "tags": ["Mechanics", "mechanics", "method of joints", "Statics", "statics", "truss"], "metadata": {"description": "<p>Solve Truss by the method of joints. &nbsp;The solution is simplified by recognizing symmetry.</p>", "licence": "Creative Commons Attribution-NonCommercial 4.0 International"}, "statement": "<p><b>Using the method of joints, </b>draw the necessary free-body diagrams and determine the force in all 13 members of the truss shown. &nbsp;Note the symmetry of the truss and load.</p>\n<table style=\"margin-left: auto; margin-right: auto;\">\n<tbody>\n<tr>\n<td>{applet(app)}</td>\n</tr>\n</tbody>\n</table>\n<p>Present your answers as a table giving the magnitude of the forces, and whether each member is in tension, compression or zero-force. Include any simplifying assumptions you make.&nbsp;</p>", "advice": "<p><strong>Given:</strong></p>\n<p style=\"padding-left: 40px;\">Truss geometry and these loads:</p>\n<p style=\"padding-left: 80px;\">$P_A = P_H = \\var{display(P_A)}$</p>\n<p style=\"padding-left: 80px;\">$P_B = P_F = \\var{display(P_B)}$</p>\n<p style=\"padding-left: 80px;\">$P_D = \\var{display(P_D)}$</p>\n<p><strong>Find reactions at $A$.</strong></p>\n<p style=\"padding-left: 40px;\">Since the truss and load are both symmetrical, the reactions at $A$ and $H$ are equal and equal to half the total load.</p>\n<p style=\"padding-left: 40px;\">$R_A = R_H = \\dfrac{2 \\cdot \\var{P_A} + 2 \\cdot\\var{P_B} + \\var{P_D}}{2} = \\var{display(R_A)}$</p>\n<p style=\"padding-left: 40px;\"></p>\n<p>For the following, unknown forces are shown in blue, known forces in red, and all members are assumed to be in tension. &nbsp;Negative results indicate the member is actually in compression.</p>\n<p></p>\n<p><strong>Find angle $\\theta$.</strong></p>\n<p style=\"padding-left: 40px;\">$\\theta = \\tan^{-1} \\left(\\dfrac{\\var{h}}{\\var{b}}\\right) = \\var{precround(theta,1)}$&deg;</p>\n<p></p>\n<p><strong>Solve Joint $A$.</strong></p>\n<p style=\"padding-left: 160px;\"><img src=\"resources/question-resources/Joint_A_j2xOQ6P.svg\" style=\"width: 251px; height: 250px;\"/></p>\n<p style=\"padding-left: 40px; text-align: left;\">\\begin{align} \\sum F_y &amp;= 0 \\\\ R_A - P_A + AB \\sin\\theta &amp;= 0\\\\ AB &amp;= &nbsp;\\dfrac{P_A - R_A}{\\sin\\theta} = \\var{display(AB)} \\end{align}</p>\n<p style=\"padding-left: 40px;\">\\begin{align} \\sum F_x &amp;= 0 \\\\ AB \\cos\\theta + AC &amp; = 0\\\\ AC &amp;= - AB \\cos \\theta = \\var{display(AC)} \\end{align}</p>\n<p><strong></strong></p>\n<p><strong>Solve Joint $C$</strong></p>\n<p style=\"padding-left: 40px;\">Note member $BC$ is a zero-force member, by inspection, so</p>\n<p style=\"padding-left: 40px;\">$CE = AC =\\var{display(CE)}$</p>\n<p><strong>Solve Joint $B$</strong></p>\n<p style=\"padding-left: 80px;\"><img src=\"resources/question-resources/Joint_B_6eszcTT.svg\" style=\"height: 220px; width: 362px;\"/></p>\n<p style=\"padding-left: 40px;\"></p>\n<p style=\"padding-left: 40px;\">\\begin{align} \\sum F_{y'} &amp;= 0 \\\\ -P_B \\cos\\theta &nbsp;- BE \\sin(2 \\theta) &amp;= 0\\\\ BE &amp;= - \\dfrac{P_B \\cos\\theta}{\\sin(2\\theta)} = \\var{display(BE)} \\end{align}</p>\n<p style=\"padding-left: 40px;\">\\begin{align} \\sum F_{x'} &amp;= 0 \\\\ &nbsp;- AB - P_B \\sin\\theta + BD + BE \\cos(2 \\theta) &amp; = 0\\\\ BD &amp;= AB + P_B \\sin\\theta - BE \\cos(2\\theta) = \\var{display(BD)} \\end{align}</p>\n<p><strong></strong></p>\n<p><strong>Solve Joint $D$</strong></p>\n<p style=\"padding-left: 120px;\"><img src=\"resources/question-resources/Joint_D_EOoCf9Z.svg\" style=\"height: 220px; width: 302px;\"/></p>\n<p style=\"padding-left: 40px;\">Note by symmetry $DF = BD$.</p>\n<p style=\"padding-left: 40px;\">\\begin{align} \\sum F_y &amp;= 0 \\\\ -DE &nbsp;- P_D &nbsp;- &nbsp;BD \\sin\\theta - DF \\sin\\theta &amp;= 0\\\\ DE &amp;= - P_D &nbsp;- 2 BD \\sin\\theta= \\var{display(DE)} \\end{align}</p>\n<p style=\"padding-left: 40px;\"></p>\n<p><strong>Find remaining forces by symmetry.</strong></p>\n<p style=\"padding-left: 40px;\">$EG = CE = \\var{display(EG)}$</p>\n<p style=\"padding-left: 40px;\">$EF = BE = \\var{display(EF)}$</p>\n<p style=\"padding-left: 40px;\">$FG = BC = $ Zero-force member<br/><br/>$FH = AB = \\var{display(FH)}$</p>\n<p style=\"padding-left: 40px;\">$GH = AC = \\var{display(GH)}$</p>\n<p></p>", "rulesets": {}, "extensions": ["geogebra", "quantities"], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"app": {"name": "app", "group": "Ungrouped variables", "definition": "//geogebra_file(\"joan-problem_fs07hWN.ggb\",params)\ngeogebra_applet(\"jvjupegg\",params)", "description": "", "templateType": "anything", "can_override": false}, "b": {"name": "b", "group": "Variables", "definition": "random(6,8,10,12, 16,18)", "description": "<p>width of panel</p>", "templateType": "anything", "can_override": false}, "h": {"name": "h", "group": "Variables", "definition": "max(2, b - random(1..5))", "description": "<p>height of panel</p>", "templateType": "anything", "can_override": false}, "params": {"name": "params", "group": "Ungrouped variables", "definition": "[b: b, h: h, \"F_A\": P_A, \"F_B\": P_B, \"F_D\": P_D, unitD:'\"'+units[0]+'\"', \"unitF\": '\"'+units[1]+'\"']", "description": "", "templateType": "anything", "can_override": false}, "P_A": {"name": "P_A", "group": "Variables", "definition": "random(50..500#50)", "description": "<p>Force at A and H</p>", "templateType": "anything", "can_override": false}, "P_B": {"name": "P_B", "group": "Variables", "definition": "random(50..500#50)", "description": "<p>Force at B and E</p>", "templateType": "anything", "can_override": false}, "P_D": {"name": "P_D", "group": "Variables", "definition": "random(50..500#50)", "description": "<p>Force at D</p>", "templateType": "anything", "can_override": false}, "units": {"name": "units", "group": "Variables", "definition": "random([\"ft\",\"lb\"],[\"m\",\"kN\"])", "description": "", "templateType": "anything", "can_override": false}, "R_A": {"name": "R_A", "group": "solutions", "definition": "(2 P_A + 2 P_B + P_D)/2", "description": "", "templateType": "anything", "can_override": false}, "theta": {"name": "theta", "group": "solutions", "definition": "degrees(arctan(h/b))", "description": "", "templateType": "anything", "can_override": false}, "AB": {"name": "AB", "group": "solutions", "definition": "(P_A - R_A)/sin(radians(theta))", "description": "", "templateType": "anything", "can_override": false}, "AC": {"name": "AC", "group": "solutions", "definition": "- AB cos(radians(theta))", "description": "", "templateType": "anything", "can_override": false}, "BE": {"name": "BE", "group": "solutions", "definition": "-(P_B cos(radians(theta)))/sin(2 radians(theta))", "description": "", "templateType": "anything", "can_override": false}, "BD": {"name": "BD", "group": "solutions", "definition": "AB + P_B sin(radians(theta)) - BE cos(2 radians(theta))", "description": "", "templateType": "anything", "can_override": false}, "DE": {"name": "DE", "group": "solutions", "definition": "-(P_D + 2 BD sin(radians(theta)))", "description": "", "templateType": "anything", "can_override": false}, "BC": {"name": "BC", "group": "solutions", "definition": "0", "description": "", "templateType": "anything", "can_override": false}, "CE": {"name": "CE", "group": "solutions", "definition": "AC", "description": "", "templateType": "anything", "can_override": false}, "GH": {"name": "GH", "group": "solutions", "definition": "AC", "description": "", "templateType": "anything", "can_override": false}, "EG": {"name": "EG", "group": "solutions", "definition": "AC", "description": "", "templateType": "anything", "can_override": false}, "DF": {"name": "DF", "group": "solutions", "definition": "BD", "description": "", "templateType": "anything", "can_override": false}, "EF": {"name": "EF", "group": "solutions", "definition": "BE", "description": "", "templateType": "anything", "can_override": false}, "FG": {"name": "FG", "group": "solutions", "definition": "BC", "description": "", "templateType": "anything", "can_override": false}, "FH": {"name": "FH", "group": "solutions", "definition": "AB", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["app", "params"], "variable_groups": [{"name": "Variables", "variables": ["b", "h", "P_B", "P_A", "P_D", "units"]}, {"name": "solutions", "variables": ["R_A", "theta", "AB", "AC", "BE", "BD", "DE", "BC", "CE", "GH", "EG", "DF", "EF", "FG", "FH"]}], "functions": {"applet": {"parameters": [["app", "ggbapplet"]], "type": "anything", "language": "javascript", "definition": "// See https://wiki.geogebra.org/en/Scripting_Commands and\n// https://wiki.geogebra.org/en/Reference:GeoGebra_Apps_API\napp.promise.then(function(d) {\n  d.app.setGridVisible(false)\n  //d.app.setAxesVisible(false,false);\n});\nreturn new Numbas.jme.types.ggbapplet(app);\n"}, "force": {"parameters": [["F", "number"]], "type": "anything", "language": "jme", "definition": "siground(qty(abs(f), units[1]),4)"}, "TC": {"parameters": [["f", "number"]], "type": "list", "language": "jme", "definition": "[if(sign(f)>0,1,0),if(sign(f)<0,1,0), if(sign(f)=0,1,0)]"}, "display": {"parameters": [["F", "number"]], "type": "anything", "language": "jme", "definition": "siground(qty(f,units[1]),4)"}}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<table height=\"585\" style=\"width: 650px; height: 283px;\">\n<tbody>\n<tr style=\"height: 53px; border-bottom: solid; background-color: #e6e6e6;\">\n<td style=\"width: 104.50367px; height: 53px; text-align: center;\"><strong>Member</strong></td>\n<td style=\"width: 93.841911px; height: 53px; text-align: center;\"><strong>Force</strong></td>\n<td style=\"width: 401.966888px; text-align: center; height: 53px;\"><strong>T or C</strong></td>\n</tr>\n<tr style=\"height: 23px;\">\n<td style=\"width: 104.50367px; text-align: center; height: 23px;\"><strong>AB</strong></td>\n<td style=\"width: 93.841911px; text-align: center; height: 23px;\">[[0]]</td>\n<td style=\"width: 401.966888px; text-align: center; height: 23px;\">[[1]]</td>\n</tr>\n<tr style=\"height: 25px;\">\n<td style=\"width: 104.50367px; text-align: center; height: 25px;\"><strong>AC</strong></td>\n<td style=\"width: 93.841911px; text-align: center; height: 25px;\">[[2]]</td>\n<td style=\"width: 401.966888px; text-align: center; height: 25px;\">[[3]]</td>\n</tr>\n<tr style=\"height: 22px;\">\n<td style=\"width: 104.50367px; text-align: center; height: 22px;\"><strong>&nbsp;BC</strong></td>\n<td style=\"width: 93.841911px; text-align: center; height: 22px;\">[[4]]</td>\n<td style=\"width: 401.966888px; text-align: center; height: 22px;\">[[5]]</td>\n</tr>\n<tr style=\"height: 22px;\">\n<td style=\"width: 104.50367px; text-align: center; height: 22px;\"><strong>&nbsp;BD</strong></td>\n<td style=\"width: 93.841911px; text-align: center; height: 22px;\">[[6]]</td>\n<td style=\"width: 401.966888px; text-align: center; height: 22px;\">[[7]]</td>\n</tr>\n<tr style=\"height: 24px;\">\n<td style=\"width: 104.50367px; text-align: center; height: 24px;\"><strong>&nbsp;BE</strong></td>\n<td style=\"width: 93.841911px; text-align: center; height: 24px;\">[[8]]</td>\n<td style=\"width: 401.966888px; text-align: center; height: 24px;\">[[9]]</td>\n</tr>\n<tr style=\"height: 22px;\">\n<td style=\"width: 104.50367px; text-align: center; height: 22px;\"><strong>&nbsp;CE</strong></td>\n<td style=\"width: 93.841911px; text-align: center; height: 22px;\">[[10]]</td>\n<td style=\"width: 401.966888px; text-align: center; height: 22px;\">[[11]]</td>\n</tr>\n<tr style=\"height: 22px;\">\n<td style=\"width: 104.50367px; text-align: center; height: 22px;\"><strong>DE</strong></td>\n<td style=\"width: 93.841911px; text-align: center; height: 22px;\">[[12]]</td>\n<td style=\"width: 401.966888px; text-align: center; height: 22px;\">[[13]]</td>\n</tr>\n<tr style=\"height: 17px;\">\n<td style=\"width: 104.50367px; text-align: center; height: 17px;\"><strong>DF</strong></td>\n<td style=\"width: 93.841911px; text-align: center; height: 17px;\">[[14]]</td>\n<td style=\"width: 401.966888px; text-align: center; height: 17px;\">[[15]]</td>\n</tr>\n<tr style=\"height: 19px;\">\n<td style=\"width: 104.50367px; text-align: center; height: 19px;\"><strong>EF</strong></td>\n<td style=\"width: 93.841911px; text-align: center; height: 19px;\">[[16]]</td>\n<td style=\"width: 401.966888px; text-align: center; height: 19px;\">[[17]]</td>\n</tr>\n<tr style=\"height: 17px;\">\n<td style=\"width: 104.50367px; text-align: center; height: 17px;\"><strong>&nbsp;EG</strong></td>\n<td style=\"width: 93.841911px; text-align: center; height: 17px;\">[[18]]</td>\n<td style=\"width: 401.966888px; text-align: center; height: 17px;\">[[19]]</td>\n</tr>\n<tr style=\"height: 17px;\">\n<td style=\"width: 104.50367px; text-align: center; height: 17px;\"><strong>FG&nbsp;</strong></td>\n<td style=\"width: 93.841911px; text-align: center; height: 17px;\">[[20]]</td>\n<td style=\"width: 401.966888px; text-align: center; height: 17px;\">[[21]]</td>\n</tr>\n<tr>\n<td style=\"width: 104.50367px; text-align: center;\"><strong>FH</strong></td>\n<td style=\"width: 93.841911px; text-align: center;\">[[22]]</td>\n<td style=\"width: 401.966888px; text-align: center;\">[[23]]</td>\n</tr>\n<tr>\n<td style=\"width: 104.50367px; text-align: center;\"><strong>GH</strong></td>\n<td style=\"width: 93.841911px; text-align: center;\">[[24]]</td>\n<td style=\"width: 401.966888px; text-align: center;\">[[25]]</td>\n</tr>\n</tbody>\n</table>\n<p>\\</p>", "gaps": [{"type": "engineering-answer", "useCustomName": true, "customName": "$AB$", "marks": "3", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "force(AB)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": true, "customName": "TCAB", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": false, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "radiogroup", "displayColumns": "3", "showBlankOption": true, "showCellAnswerState": true, "choices": ["T", "C", "Zero Force"], "matrix": "TC(AB)"}, {"type": "engineering-answer", "useCustomName": true, "customName": "$AC$", "marks": "2", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "force(AC)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": true, "customName": "TCAC", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": false, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "radiogroup", "displayColumns": "3", "showBlankOption": true, "showCellAnswerState": true, "choices": ["T", "C", "Zero Force"], "matrix": "TC(AC)"}, {"type": "engineering-answer", "useCustomName": true, "customName": "$BC$", "marks": "2", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "force(BC)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": true, "customName": "TCBC", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": false, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "radiogroup", "displayColumns": "3", "showBlankOption": true, "showCellAnswerState": true, "choices": ["T", "C", "Zero Force"], "matrix": "TC(BC)"}, {"type": "engineering-answer", "useCustomName": true, "customName": "$BD$", "marks": "2", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "force(BD)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": true, "customName": "TCBD", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "radiogroup", "displayColumns": "3", "showBlankOption": true, "showCellAnswerState": true, "choices": ["T", "C", "Zero Force"], "matrix": "TC(BD)"}, {"type": "engineering-answer", "useCustomName": true, "customName": "$BE$", "marks": "2", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "force(BE)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": true, "customName": "TCBE", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "radiogroup", "displayColumns": "3", "showBlankOption": true, "showCellAnswerState": true, "choices": ["T", "C", "Zero Force"], "matrix": "TC(BE)"}, {"type": "engineering-answer", "useCustomName": true, "customName": "$CE$", "marks": "2", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "force(CE)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": true, "customName": "TCCE", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "radiogroup", "displayColumns": "3", "showBlankOption": true, "showCellAnswerState": true, "choices": ["T", "C", "Zero Force"], "matrix": "TC(CE)"}, {"type": "engineering-answer", "useCustomName": true, "customName": "$DE$", "marks": "2", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "force(DE)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": true, "customName": "TCDE", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "radiogroup", "displayColumns": "3", "showBlankOption": true, "showCellAnswerState": true, "choices": ["T", "C", "Zero Force"], "matrix": "TC(DE)"}, {"type": "engineering-answer", "useCustomName": true, "customName": "$DF$", "marks": "2", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "force(DF)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": true, "customName": "TCDF", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "radiogroup", "displayColumns": "3", "showBlankOption": true, "showCellAnswerState": true, "choices": ["T", "C", "Zero Force"], "matrix": "TC(DF)"}, {"type": "engineering-answer", "useCustomName": true, "customName": "$EF$", "marks": "2", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "force(EF)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": true, "customName": "TCEF", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "radiogroup", "displayColumns": "3", "showBlankOption": true, "showCellAnswerState": true, "choices": ["T", "C", "Zero Force"], "matrix": "TC(EF)"}, {"type": "engineering-answer", "useCustomName": true, "customName": "$EG$", "marks": "2", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "force(EG)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": true, "customName": "TCEG", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "radiogroup", "displayColumns": "3", "showBlankOption": true, "showCellAnswerState": true, "choices": ["T", "C", "Zero Force"], "matrix": "TC(EG)"}, {"type": "engineering-answer", "useCustomName": true, "customName": "$FG$", "marks": "2", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "force(FG)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": true, "customName": "TCFG", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "radiogroup", "displayColumns": "3", "showBlankOption": true, "showCellAnswerState": true, "choices": ["T", "C", "Zero Force"], "matrix": "TC(FG)"}, {"type": "engineering-answer", "useCustomName": true, "customName": "$FH$", "marks": "2", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "force(FH)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": true, "customName": "TCFH", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "radiogroup", "displayColumns": "3", "showBlankOption": true, "showCellAnswerState": true, "choices": ["T", "C", "Zero Force"], "matrix": "TC(FH)"}, {"type": "engineering-answer", "useCustomName": true, "customName": "$GH$", "marks": "2", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "force(GH)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": true, "customName": "TCGH", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "radiogroup", "displayColumns": "3", "showBlankOption": true, "showCellAnswerState": true, "choices": ["T", "C", "Zero Force"], "matrix": "TC(GH)"}], "sortAnswers": false}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}], "resources": ["question-resources/joan-problem_fs07hWN.ggb", "question-resources/Joint_B_6eszcTT.svg", "question-resources/Joint_D_EOoCf9Z.svg", "question-resources/fbds_jqa9zTW_0O5KNSm.afdesign", "question-resources/Joint_A_j2xOQ6P.svg", "question-resources/joans-problem.ggb"]}]}], "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}]}