// Numbas version: finer_feedback_settings
{"name": "Moment of inertia: semicircle and hole", "extensions": ["geogebra", "quantities"], "custom_part_types": [{"source": {"pk": 19, "author": {"name": "William Haynes", "pk": 2530}, "edit_page": "/part_type/19/edit"}, "name": "Engineering Accuracy with units", "short_name": "engineering-answer", "description": "<p>A value with units marked right if within an adjustable % error of the correct value. &nbsp;Marked close if within a wider margin of error.</p>", "help_url": "", "input_widget": "string", "input_options": {"correctAnswer": "siground(settings['correctAnswer'],4)", "hint": {"static": true, "value": ""}, "allowEmpty": {"static": true, "value": true}}, "can_be_gap": true, "can_be_step": true, "marking_script": "mark:\nswitch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)\n\ninterpreted_answer:\nqty(student_scalar, student_units)\n\n\n\ncorrect_quantity:\nsettings[\"correctAnswer\"]\n\n\n\ncorrect_units:\nunits(correct_quantity)\n\n\nallowed_notation_styles:\n[\"plain\",\"en\"]\n\nmatch_student_number:\nmatchnumber(studentAnswer,allowed_notation_styles)\n\nstudent_scalar:\nmatch_student_number[1]\n\nstudent_units:\nreplace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")\n\ngood_units:\ntry(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n\n\nstudent_quantity:\nswitch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n\n\npercent_error:\ntry(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   \n\nright:\npercent_error <= settings['right']\n\n\nclose:\nright_sign and percent_error <= settings['close']\n\nright_sign:\nsign(student_scalar) = sign(correct_quantity)", "marking_notes": [{"name": "mark", "description": "This is the main marking note. It should award credit and provide feedback based on the student's answer.", "definition": "switch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)"}, {"name": "interpreted_answer", "description": "A value representing the student's answer to this part.", "definition": "qty(student_scalar, student_units)\n\n"}, {"name": "correct_quantity", "description": "", "definition": "settings[\"correctAnswer\"]\n\n"}, {"name": "correct_units", "description": "", "definition": "units(correct_quantity)\n"}, {"name": "allowed_notation_styles", "description": "", "definition": "[\"plain\",\"en\"]"}, {"name": "match_student_number", "description": "", "definition": "matchnumber(studentAnswer,allowed_notation_styles)"}, {"name": "student_scalar", "description": "", "definition": "match_student_number[1]"}, {"name": "student_units", "description": "<p>Modify&nbsp;the&nbsp;unit portion of&nbsp;the student's answer by</p>\n<p>1. replacing \"ohms\" with \"ohm\" &nbsp;case insensitive</p>\n<p>2.&nbsp;replacing '-'&nbsp;with ' '&nbsp;</p>\n<p>3. replacing '&deg;' with ' deg'&nbsp;</p>\n<p>to allow answers like 10 ft-lb and 30&deg;</p>", "definition": "replace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")"}, {"name": "good_units", "description": "", "definition": "try(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n"}, {"name": "student_quantity", "description": "<p>This fixes the student answer for two common errors. &nbsp;</p>\n<p>If student_units are wrong&nbsp;&nbsp;- replace with correct units</p>\n<p>If student_scalar has the wrong sign - replace with right sign</p>\n<p>If student makes both errors, only one gets fixed.</p>", "definition": "switch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n"}, {"name": "percent_error", "description": "", "definition": "try(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   "}, {"name": "right", "description": "", "definition": "percent_error <= settings['right']\n"}, {"name": "close", "description": "<p>Only marked close if the student actually has the right sign.</p>", "definition": "right_sign and percent_error <= settings['close']"}, {"name": "right_sign", "description": "", "definition": "sign(student_scalar) = sign(correct_quantity) "}], "settings": [{"name": "correctAnswer", "label": "Correct Quantity.", "help_url": "", "hint": "The correct answer given as a JME quantity.", "input_type": "code", "default_value": "", "evaluate": true}, {"name": "right", "label": "% Accuracy for right.", "help_url": "", "hint": "Question will be considered correct if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "0.2", "evaluate": true}, {"name": "close", "label": "% Accuracy for close.", "help_url": "", "hint": "Question will be considered close if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "1.0", "evaluate": true}, {"name": "C1", "label": "Close with units.", "help_url": "", "hint": "Partial Credit for close value with appropriate units. &nbsp;if correct answer is 100 N and close is &plusmn;1%,<br />99 &nbsp;N is accepted.", "input_type": "percent", "default_value": "75"}, {"name": "C2", "label": "No units or wrong sign", "help_url": "", "hint": "Partial credit for forgetting units or using wrong sign.<br />If the correct answer is 100 N, both 100 and -100 N are accepted.", "input_type": "percent", "default_value": "50"}, {"name": "C3", "label": "Close, no units.", "help_url": "", "hint": "Partial Credit for close value but forgotten units.<br />This value would be close&nbsp;if the expected units were provided. &nbsp;If the correct answer is 100 N, and close is &plusmn;1%,<br />99 is accepted.", "input_type": "percent", "default_value": "25"}], "public_availability": "always", "published": true, "extensions": ["quantities"]}], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Moment of inertia: semicircle and hole", "tags": ["first moment of area", "Mechanics", "mechanics", "moment of inertia", "parallel axis theorem", "semicircle", "Statics", "statics"], "metadata": {"description": "<p>Find moment of inertia of a shape which requires the use of the parallel axis theorem for a semicircle.</p>", "licence": "Creative Commons Attribution-NonCommercial 4.0 International"}, "statement": "<p>Determine the moment of inertia and radius of gyration of the composite shape with respect to the $x$- and $y$-axes knowing that $b$ = {b}, $h$ = {h}, $r_o$ = {r_o}, and $r_i$ = {display(r_i)}.</p>\n<table style=\"margin-left: auto; margin-right: auto;\">\n<tbody>\n<tr>\n<td>{geogebra_applet('fqywxznp',{ggbArgs})}</td>\n</tr>\n</tbody>\n</table>", "advice": "<p><strong>Number the parts:</strong> &nbsp;</p>\n<p style=\"padding-left: 30px;\">Part 1 rectangle: $b$ = {b}, $h$ = {h}.</p>\n<p style=\"padding-left: 30px;\">Part 2 semicircle: &nbsp;$r_o$ = {r_o}.</p>\n<p style=\"padding-left: 30px;\">Part 3 circular hole: $r_i$ = {display(r_i)}.</p>\n<p><strong>Calculate the areas:</strong></p>\n<p style=\"padding-left: 30px;\">$\\begin{align} A_1 &amp;= b h &amp; A_2 &amp;= &nbsp;\\tfrac{1}{2}\\pi r_o^2 &amp;A_3 &amp;= \\pi r_i^2&amp;A&amp;=A_1 +A_2 - A_3\\\\&amp;=\\var{display(A1)} &amp;&amp;=\\var{display(A2)} &amp;&amp;=\\var{display(A3)}&amp;&amp;=\\var{display(A)}\\\\&amp;&amp;\\end{align}$</p>\n<p><strong>Calculate the moment of inertia with respect to the <em>x</em>-axis.</strong></p>\n<p style=\"padding-left: 30px;\">$\\begin{align} \\\\{I_1}_x&nbsp;&amp;= \\tfrac{1}{3} b h^3 \\\\&amp;= &nbsp;\\tfrac{1}{3}( \\var{b})(\\var{h})^3 \\\\&amp;= \\var{display(Ix1)} \\\\\\\\{I_x}_2 &amp;= [\\bar{I} + A &nbsp;d^2]_2\\\\ &nbsp;&amp;= 0.1098 r_o^4 + \\left(\\frac{\\pi r_o^2}{2}\\right) \\left(h + \\frac{4 r_o}{3 \\pi}\\right)^2\\\\&amp;= \\var{display(Ibarx2)} + (\\var{display(A2)}) (\\var{display(d2)})^2\\\\&amp;=\\var{display(Ix2)}&nbsp;\\\\\\\\ {I_x}_3 &amp;= [\\bar{I} + A &nbsp;d^2]_3&nbsp;\\\\ &nbsp;&amp;= \\frac{\\pi&nbsp;r_i^4}{4} &nbsp;+ \\left(\\pi r_i^2\\right) \\left(h \\right)^2\\\\&amp;= \\var{display(Ibarx3)} + (\\var{display(A3)}) (\\var{display(h)})^2\\\\&amp;=\\var{display(Ix3)}\\\\\\\\I_x &amp;= {I_x}_1+{I_x}_2-{I_x}_3\\\\&amp;=\\var{display(Ix1)}+\\var{display(Ix2)}-\\var{display(Ix3)}\\\\&amp;=\\var{display(Ix)}\\end{align}$</p>\n<p><strong>Calculate the moment of inertia with respect to the <em>y</em>-axis.</strong></p>\n<p style=\"padding-left: 30px;\">$\\begin{align} \\\\{I_y}_1 &amp;= \\tfrac{1}{3} h&nbsp;b^3 \\\\&amp;= \\tfrac{1}{3}( \\var{h})(\\var{b})^3 \\\\&amp;= \\var{display(Iy1)} \\\\\\\\{I_y}_2 &amp;= [\\bar{I} + A d^2]_2\\\\ &amp;= \\frac{\\pi&nbsp;r_o^4}{8} + \\left(\\frac{\\pi r_o^2}{2}\\right) \\left(\\frac{b}{2}\\right)^2\\\\&amp;= \\var{display(Ibary2)} + (\\var{display(A2)}) (\\var{display(b/2)})^2\\\\&amp;=\\var{display(Iy2)} \\\\\\\\ {I_y}_3 &amp;= [\\bar{I} + A d^2]_3 \\\\ &amp;= \\frac{\\pi r_i^4}{4} + \\left(\\pi r_i^2\\right) \\left(\\frac{b}{2}&nbsp;\\right)^2\\\\&amp;= \\var{display(Ibary3)} + (\\var{display(A3)}) (\\var{display(b/2)})^2\\\\&amp;=\\var{display(Iy3)}\\\\\\\\I_y &amp;= {I_y}_1+{I_y}_2-{I_y}_3\\\\&amp;=\\var{display(Iy1)}+\\var{display(Iy2)}-\\var{display(Iy3)}\\\\&amp;=\\var{display(Iy)}\\end{align}$</p>\n<p><strong>Calculate the radius of gyration with respect to the <em>x</em>- and <em>y</em>- axes.</strong></p>\n<p style=\"padding-left: 60px;\">$\\begin{align}\\\\k_x &amp;= \\sqrt{\\dfrac{I_x}{A}}&amp;k_y &amp;= \\sqrt{\\dfrac{I_y}{A}}\\\\ &amp;= \\sqrt{\\dfrac{\\var{display(Ix)}}{\\var{display(A)}}}&amp;&amp;= \\sqrt{\\dfrac{\\var{display(Iy)}}{\\var{display(A)}}}\\\\ &amp;= \\var{display(kx)}&nbsp;&amp;&amp;=\\var{display(ky)} \\end{align}$</p>\n<p style=\"padding-left: 60px;\"></p>", "rulesets": {}, "extensions": ["geogebra", "quantities"], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"A3": {"name": "A3", "group": "Area", "definition": "3.14159 r_i r_i", "description": "", "templateType": "anything", "can_override": false}, "A2": {"name": "A2", "group": "Area", "definition": "3.14159 r_o r_o /2", "description": "", "templateType": "anything", "can_override": false}, "A1": {"name": "A1", "group": "Area", "definition": "b h", "description": "", "templateType": "anything", "can_override": false}, "r_i": {"name": "r_i", "group": "Inputs", "definition": "random(1/3, 0.5, 2/3,0.75) r_o ", "description": "", "templateType": "anything", "can_override": false}, "h": {"name": "h", "group": "Inputs", "definition": "qty(random(2..8),units)", "description": "", "templateType": "anything", "can_override": false}, "Iy": {"name": "Iy", "group": "Iy", "definition": "Iy1 + Iy2 - Iy3", "description": "", "templateType": "anything", "can_override": false}, "Ibary2": {"name": "Ibary2", "group": "Iy", "definition": "3.14159 r_o  r_o  r_o  r_o /8", "description": "", "templateType": "anything", "can_override": false}, "Iy2": {"name": "Iy2", "group": "Iy", "definition": "Ibary2 + A2 b/2 b/2", "description": "", "templateType": "anything", "can_override": false}, "Ix2": {"name": "Ix2", "group": "Ix", "definition": "Ibarx2 + A2 d2 d2", "description": "", "templateType": "anything", "can_override": false}, "Ix1": {"name": "Ix1", "group": "Ix", "definition": "b h h h /3", "description": "<p>rectangle</p>", "templateType": "anything", "can_override": false}, "kx": {"name": "kx", "group": "Ix", "definition": "qty(sqrt(scalar(Ix/A)),units)", "description": "", "templateType": "anything", "can_override": false}, "Ibary3": {"name": "Ibary3", "group": "Iy", "definition": "Ibarx3", "description": "", "templateType": "anything", "can_override": false}, "d2": {"name": "d2", "group": "Ix", "definition": "h + 4 r_o/3/3.14159", "description": "", "templateType": "anything", "can_override": false}, "b": {"name": "b", "group": "Inputs", "definition": "qty(random(4..12#2),units)", "description": "", "templateType": "anything", "can_override": false}, "ky": {"name": "ky", "group": "Iy", "definition": "qty(sqrt(scalar(Iy/A)),units)", "description": "", "templateType": "anything", "can_override": false}, "Ibarx2": {"name": "Ibarx2", "group": "Ix", "definition": "0.1098 r_o r_o r_o r_o", "description": "", "templateType": "anything", "can_override": false}, "Ibarx3": {"name": "Ibarx3", "group": "Ix", "definition": "3.14159 r_i r_i r_i r_i/4", "description": "", "templateType": "anything", "can_override": false}, "Iy3": {"name": "Iy3", "group": "Iy", "definition": "Ibary3 + A3 b/2 b/2", "description": "", "templateType": "anything", "can_override": false}, "ggbArgs": {"name": "ggbArgs", "group": "Inputs", "definition": "[['A',vector(scalar(b),scalar(h))],['R_o',scalar(r_o)],['R_i',scalar(r_i)]]", "description": "", "templateType": "anything", "can_override": false}, "A": {"name": "A", "group": "Area", "definition": "A1 + A2 - A3", "description": "", "templateType": "anything", "can_override": false}, "units": {"name": "units", "group": "Inputs", "definition": "random('in','cm')", "description": "", "templateType": "anything", "can_override": false}, "r_o": {"name": "r_o", "group": "Inputs", "definition": "qty(random(2..scalar(b)/2),units)", "description": "<p>r</p>", "templateType": "anything", "can_override": false}, "debug": {"name": "debug", "group": "Inputs", "definition": "false", "description": "", "templateType": "anything", "can_override": false}, "Ix3": {"name": "Ix3", "group": "Ix", "definition": "Ibarx3 + A3 h h", "description": "", "templateType": "anything", "can_override": false}, "Iy1": {"name": "Iy1", "group": "Iy", "definition": "h b b b /3", "description": "", "templateType": "anything", "can_override": false}, "Ix": {"name": "Ix", "group": "Ix", "definition": "Ix1 + Ix2 - Ix3", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "r_i < h", "maxRuns": 100}, "ungrouped_variables": [], "variable_groups": [{"name": "Inputs", "variables": ["b", "h", "r_o", "r_i", "units", "debug", "ggbArgs"]}, {"name": "Ix", "variables": ["Ix1", "Ibarx2", "Ix2", "d2", "Ibarx3", "Ix3", "Ix", "kx"]}, {"name": "Area", "variables": ["A1", "A2", "A3", "A"]}, {"name": "Iy", "variables": ["Iy1", "Ibary2", "Iy2", "Ibary3", "Iy3", "Iy", "ky"]}], "functions": {"display": {"parameters": [["q", "quantity"]], "type": "string", "language": "jme", "definition": "string(siground(q,4))"}}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>$I_x$ =&nbsp;[[0]]</p>\n<p>$k_x$ =&nbsp;[[1]]</p>\n<p>$I_y$ = [[2]]</p>\n<p>$k_y$ = [[3]]</p>\n<p><span data-jme-visible=\"debug\">{display(Ix)} &nbsp;</span> <span data-jme-visible=\"debug\">{display(kx)} {display(Iy)} {display(ky)}</span></p>", "gaps": [{"type": "engineering-answer", "useCustomName": true, "customName": "$I_x$", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "Ix", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "$k_x$", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "kx", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "$I_y$", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "Iy", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "$k_y$", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "ky", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}], "sortAnswers": false}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}], "resources": []}]}], "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}]}