// Numbas version: exam_results_page_options
{"name": "Equilibrium of a particle - equilibrium equation method", "extensions": ["/srv/numbas/media/user-extensions/extracted/41/geogebra", "/srv/numbas/media/user-extensions/extracted/55/quantities", "/srv/numbas/media/user-extensions/extracted/98/weh"], "custom_part_types": [{"can_be_step": true, "settings": [{"hint": "The correct answer given as a JME quantity.", "input_type": "code", "evaluate": true, "label": "Correct Quantity.", "name": "correctAnswer", "default_value": "", "help_url": ""}, {"hint": "Question will be considered correct if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "evaluate": true, "label": "% Accuracy for right.", "name": "right", "default_value": "0.2", "help_url": ""}, {"hint": "Question will be considered close if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "evaluate": true, "label": "% Accuracy for close.", "name": "close", "default_value": "1.0", "help_url": ""}, {"hint": "Partial Credit for close value with appropriate units. &nbsp;if correct answer is 100 N and close is &plusmn;1%,<br />99 &nbsp;N is accepted.", "input_type": "percent", "default_value": "75", "label": "Close with units.", "name": "C1", "help_url": ""}, {"hint": "Partial credit for forgetting units or using wrong sign.<br />If the correct answer is 100 N, both 100 and -100 N are accepted.", "input_type": "percent", "default_value": "50", "label": "No units or wrong sign", "name": "C2", "help_url": ""}, {"hint": "Partial Credit for close value but forgotten units.<br />This value would be close&nbsp;if the expected units were provided. &nbsp;If the correct answer is 100 N, and close is &plusmn;1%,<br />99 is accepted.", "input_type": "percent", "default_value": "25", "label": "Close, no units.", "name": "C3", "help_url": ""}], "marking_notes": [{"definition": "switch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)\n", "name": "mark", "description": "This is the main marking note. It should award credit and provide feedback based on the student's answer."}, {"definition": "qty(student_scalar, student_units)\n\n", "name": "interpreted_answer", "description": "A value representing the student's answer to this part."}, {"definition": "settings[\"correctAnswer\"]\n\n", "name": "correct_quantity", "description": ""}, {"definition": "units(correct_quantity)\n", "name": "correct_units", "description": ""}, {"definition": "[\"plain\",\"en\"]", "name": "allowed_notation_styles", "description": ""}, {"definition": "matchnumber(studentAnswer,allowed_notation_styles)", "name": "match_student_number", "description": ""}, {"definition": "match_student_number[1]", "name": "student_scalar", "description": ""}, {"definition": "join(\nsplit(\njoin(\nsplit(studentAnswer[len(match_student_number[0])..len(studentAnswer)]\n,'-'),' ')\n,'\u00b0'),' deg')\n", "name": "student_units", "description": "<p>Does clumsy substitution to</p>\n<p>1.&nbsp;replace '-'&nbsp;with ' '&nbsp;</p>\n<p>2. replace '&deg;' with ' deg'&nbsp;</p>\n<p>to allow answers like 10 ft-lb and 30&deg;</p>"}, {"definition": "try(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n", "name": "good_units", "description": ""}, {"definition": "switch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n", "name": "student_quantity", "description": "<p>This fixes the student answer for two common errors. &nbsp;</p>\n<p>If student_units are wrong&nbsp;&nbsp;- replace with correct units</p>\n<p>If student_scalar has the wrong sign - replace with right sign</p>\n<p>If student makes both errors, only one gets fixed.</p>"}, {"definition": "try(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   ", "name": "percent_error", "description": ""}, {"definition": "percent_error <= settings['right']\n", "name": "right", "description": ""}, {"definition": "right_sign and percent_error <= settings['close']", "name": "close", "description": "<p>Only marked close if the student actually has the right sign.</p>"}, {"definition": "sign(student_scalar) = sign(correct_quantity) ", "name": "right_sign", "description": ""}], "name": "Engineering Accuracy with units", "source": {"edit_page": "/part_type/19/edit", "author": {"name": "William Haynes", "pk": 2530}, "pk": 19}, "marking_script": "mark:\nswitch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)\n\n\ninterpreted_answer:\nqty(student_scalar, student_units)\n\n\n\ncorrect_quantity:\nsettings[\"correctAnswer\"]\n\n\n\ncorrect_units:\nunits(correct_quantity)\n\n\nallowed_notation_styles:\n[\"plain\",\"en\"]\n\nmatch_student_number:\nmatchnumber(studentAnswer,allowed_notation_styles)\n\nstudent_scalar:\nmatch_student_number[1]\n\nstudent_units:\njoin(\nsplit(\njoin(\nsplit(studentAnswer[len(match_student_number[0])..len(studentAnswer)]\n,'-'),' ')\n,'\u00b0'),' deg')\n\n\ngood_units:\ntry(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n\n\nstudent_quantity:\nswitch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n\n\npercent_error:\ntry(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   \n\nright:\npercent_error <= settings['right']\n\n\nclose:\nright_sign and percent_error <= settings['close']\n\nright_sign:\nsign(student_scalar) = sign(correct_quantity)", "short_name": "engineering-answer", "input_widget": "string", "published": false, "help_url": "", "description": "<p>A value with units marked right if within an adjustable % error of the correct value. &nbsp;Marked close if within a wider margin of error.</p>", "extensions": ["quantities"], "public_availability": "restricted", "can_be_gap": true, "input_options": {"hint": {"value": "", "static": true}, "correctAnswer": "siground(settings['correctAnswer'],4)", "allowEmpty": {"value": true, "static": true}}}], "resources": [], "navigation": {"showfrontpage": false, "preventleave": false, "allowregen": true}, "question_groups": [{"questions": [{"metadata": {"description": "<p>Find forces required to hold a particle in equilibrium when subjected to a downward load. &nbsp;Directions of the reactions are given. &nbsp;</p>", "licence": "Creative Commons Attribution-NonCommercial 4.0 International"}, "ungrouped_variables": ["unit", "W", "alpha", "theta"], "functions": {}, "parts": [{"variableReplacements": [], "showCorrectAnswer": true, "type": "gapfill", "marks": 0, "variableReplacementStrategy": "originalfirst", "prompt": "<p>B =&nbsp;[[0]]&nbsp;&nbsp;</p>\n<p>C =&nbsp;[[1]]&nbsp;</p>", "sortAnswers": false, "customMarkingAlgorithm": "", "unitTests": [], "scripts": {}, "extendBaseMarkingAlgorithm": true, "gaps": [{"scripts": {}, "unitTests": [], "variableReplacements": [], "showCorrectAnswer": true, "type": "engineering-answer", "marks": "4", "variableReplacementStrategy": "originalfirst", "extendBaseMarkingAlgorithm": true, "customMarkingAlgorithm": "", "settings": {"correctAnswer": "qty(B,unit)", "close": "1.0", "C1": "75", "C3": "25", "right": "0.2", "C2": "50"}, "showFeedbackIcon": true}, {"scripts": {}, "unitTests": [], "variableReplacements": [], "showCorrectAnswer": true, "type": "engineering-answer", "marks": "4", "variableReplacementStrategy": "originalfirst", "extendBaseMarkingAlgorithm": true, "customMarkingAlgorithm": "", "settings": {"correctAnswer": "qty(C,unit)", "close": "1.0", "C1": "75", "C3": "25", "right": "0.2", "C2": "50"}, "showFeedbackIcon": true}], "showFeedbackIcon": true}], "type": "question", "name": "Equilibrium of a particle - equilibrium equation method", "contributors": [{"profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/", "name": "William Haynes"}], "advice": "<p><strong>Knowns:</strong></p>\n<p>$W = \\var{qty(W,unit)}$, $\\alpha = \\var{alpha}&deg;$ and $\\theta=\\var{theta}&deg;$.</p>\n<p><strong>Apply equations of equilibrium</strong></p>\n<p>$\\begin{align}\\Sigma F_x &amp;= 0 &amp; \\Sigma F_y &amp;= 0 \\\\-B_x + C_x &amp;= 0 &amp; B_y + C_y -W &amp;= 0\\\\ B \\sin\\alpha &amp; = C \\cos\\theta &amp; B \\cos\\alpha + C \\sin\\theta &amp; = W \\\\ B &amp;= C \\left(\\frac{\\cos\\theta}{\\sin\\alpha}\\right)\\\\B &amp;= \\var{K_1}\\, C\\\\&amp;&amp;(\\var{K_1}\\, C)&nbsp;\\cos\\alpha + C \\sin\\theta &amp; = W \\\\&amp;&amp; &nbsp;\\var{K_2}\\, C&amp;= \\var{qty(W,unit)} \\\\&amp;&amp;C = \\var{qty(siground(W/K_2,4),unit)} \\\\ B &amp;= \\var{K_1}\\,&nbsp;(\\var{qty(siground(W/K_2,4),unit)})\\\\B&amp;=\\var{qty(siground(B,4),unit)} \\end{align}$</p>", "statement": "<p>{geogebra_applet('uwmn9de6',[['alpha', alpha+'&deg;'],['theta', &nbsp;theta+'&deg;']])}</p>\n<p>Three forces act on particle <em>A </em>as shown. Use the equations of equilibrium to determine the magnitudes of&nbsp;<em>B</em>&nbsp;and <em>C &nbsp;</em>when load&nbsp;<em>W</em> = {qty(W,unit)}.</p>", "rulesets": {}, "variable_groups": [{"variables": ["B", "C", "K_1", "K_2"], "name": "Unnamed group"}], "variables": {"B": {"definition": "W cos(radians(theta))/sin(radians(90-alpha+theta))", "group": "Unnamed group", "templateType": "anything", "name": "B", "description": ""}, "K_1": {"definition": "{siground(cos(radians(theta))/sin(radians(alpha)),4)", "group": "Unnamed group", "templateType": "anything", "name": "K_1", "description": ""}, "alpha": {"definition": "random(15..80#5)", "group": "Ungrouped variables", "templateType": "anything", "name": "alpha", "description": ""}, "theta": {"definition": "random(15..80#5)", "group": "Ungrouped variables", "templateType": "anything", "name": "theta", "description": ""}, "unit": {"definition": "random(['lb', 'N', 'kN'])", "group": "Ungrouped variables", "templateType": "anything", "name": "unit", "description": ""}, "W": {"definition": "random(50..600#50)", "group": "Ungrouped variables", "templateType": "anything", "name": "W", "description": ""}, "K_2": {"definition": "precround(K_1 cos(radians(alpha)) + sin(radians(theta)),4)", "group": "Unnamed group", "templateType": "anything", "name": "K_2", "description": ""}, "C": {"definition": "W sin(radians(alpha))/sin(radians(90-alpha+theta))", "group": "Unnamed group", "templateType": "anything", "name": "C", "description": ""}}, "extensions": ["geogebra", "quantities", "weh"], "variablesTest": {"condition": "alpha <> (90 - theta)", "maxRuns": 100}, "tags": ["equilibrium", "Equilibrium", "mechanics", "Mechanics", "particle", "Statics", "statics", "vectors"], "preamble": {"js": "", "css": ""}}], "pickingStrategy": "all-ordered"}], "contributors": [{"profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/", "name": "William Haynes"}]}