// Numbas version: finer_feedback_settings
{"name": "Friction: Range of values 2", "extensions": ["geogebra", "quantities"], "custom_part_types": [{"source": {"pk": 19, "author": {"name": "William Haynes", "pk": 2530}, "edit_page": "/part_type/19/edit"}, "name": "Engineering Accuracy with units", "short_name": "engineering-answer", "description": "<p>A value with units marked right if within an adjustable % error of the correct value. &nbsp;Marked close if within a wider margin of error.</p>", "help_url": "", "input_widget": "string", "input_options": {"correctAnswer": "siground(settings['correctAnswer'],4)", "hint": {"static": true, "value": ""}, "allowEmpty": {"static": true, "value": true}}, "can_be_gap": true, "can_be_step": true, "marking_script": "mark:\nswitch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)\n\ninterpreted_answer:\nqty(student_scalar, student_units)\n\n\n\ncorrect_quantity:\nsettings[\"correctAnswer\"]\n\n\n\ncorrect_units:\nunits(correct_quantity)\n\n\nallowed_notation_styles:\n[\"plain\",\"en\"]\n\nmatch_student_number:\nmatchnumber(studentAnswer,allowed_notation_styles)\n\nstudent_scalar:\nmatch_student_number[1]\n\nstudent_units:\nreplace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")\n\ngood_units:\ntry(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n\n\nstudent_quantity:\nswitch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n\n\npercent_error:\ntry(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   \n\nright:\npercent_error <= settings['right']\n\n\nclose:\nright_sign and percent_error <= settings['close']\n\nright_sign:\nsign(student_scalar) = sign(correct_quantity)", "marking_notes": [{"name": "mark", "description": "This is the main marking note. It should award credit and provide feedback based on the student's answer.", "definition": "switch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)"}, {"name": "interpreted_answer", "description": "A value representing the student's answer to this part.", "definition": "qty(student_scalar, student_units)\n\n"}, {"name": "correct_quantity", "description": "", "definition": "settings[\"correctAnswer\"]\n\n"}, {"name": "correct_units", "description": "", "definition": "units(correct_quantity)\n"}, {"name": "allowed_notation_styles", "description": "", "definition": "[\"plain\",\"en\"]"}, {"name": "match_student_number", "description": "", "definition": "matchnumber(studentAnswer,allowed_notation_styles)"}, {"name": "student_scalar", "description": "", "definition": "match_student_number[1]"}, {"name": "student_units", "description": "<p>Modify&nbsp;the&nbsp;unit portion of&nbsp;the student's answer by</p>\n<p>1. replacing \"ohms\" with \"ohm\" &nbsp;case insensitive</p>\n<p>2.&nbsp;replacing '-'&nbsp;with ' '&nbsp;</p>\n<p>3. replacing '&deg;' with ' deg'&nbsp;</p>\n<p>to allow answers like 10 ft-lb and 30&deg;</p>", "definition": "replace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")"}, {"name": "good_units", "description": "", "definition": "try(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n"}, {"name": "student_quantity", "description": "<p>This fixes the student answer for two common errors. &nbsp;</p>\n<p>If student_units are wrong&nbsp;&nbsp;- replace with correct units</p>\n<p>If student_scalar has the wrong sign - replace with right sign</p>\n<p>If student makes both errors, only one gets fixed.</p>", "definition": "switch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n"}, {"name": "percent_error", "description": "", "definition": "try(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   "}, {"name": "right", "description": "", "definition": "percent_error <= settings['right']\n"}, {"name": "close", "description": "<p>Only marked close if the student actually has the right sign.</p>", "definition": "right_sign and percent_error <= settings['close']"}, {"name": "right_sign", "description": "", "definition": "sign(student_scalar) = sign(correct_quantity) "}], "settings": [{"name": "correctAnswer", "label": "Correct Quantity.", "help_url": "", "hint": "The correct answer given as a JME quantity.", "input_type": "code", "default_value": "", "evaluate": true}, {"name": "right", "label": "% Accuracy for right.", "help_url": "", "hint": "Question will be considered correct if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "0.2", "evaluate": true}, {"name": "close", "label": "% Accuracy for close.", "help_url": "", "hint": "Question will be considered close if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "1.0", "evaluate": true}, {"name": "C1", "label": "Close with units.", "help_url": "", "hint": "Partial Credit for close value with appropriate units. &nbsp;if correct answer is 100 N and close is &plusmn;1%,<br />99 &nbsp;N is accepted.", "input_type": "percent", "default_value": "75"}, {"name": "C2", "label": "No units or wrong sign", "help_url": "", "hint": "Partial credit for forgetting units or using wrong sign.<br />If the correct answer is 100 N, both 100 and -100 N are accepted.", "input_type": "percent", "default_value": "50"}, {"name": "C3", "label": "Close, no units.", "help_url": "", "hint": "Partial Credit for close value but forgotten units.<br />This value would be close&nbsp;if the expected units were provided. &nbsp;If the correct answer is 100 N, and close is &plusmn;1%,<br />99 is accepted.", "input_type": "percent", "default_value": "25"}], "public_availability": "always", "published": true, "extensions": ["quantities"]}], "resources": [["question-resources/friction-block_xtezuTb.ggb", "friction-block_xtezuTb.ggb"]], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Friction: Range of values 2", "tags": [], "metadata": {"description": "<p>Friction</p>", "licence": "Creative Commons Attribution-NonCommercial 4.0 International"}, "statement": "<p>A $\\var{qty(500,'N')}$ block is supported on an inclined rail as shown. The coefficient of friction between the block and the rail is $\\mu_s=\\var{mu_s}$.</p>\n<table style=\"margin-left: auto; margin-right: auto;\">\n<tbody>\n<tr>\n<td>&nbsp;{geogebra_applet('axvmx7au', params)}</td>\n</tr>\n</tbody>\n</table>\n<p></p>", "advice": "<p>{fbd}</p>\n<p>Given: $W = \\var{qty(500,\"N\")}, \\mu_s= \\var{mu_s}$</p>\n<p>$+\\!\\!\\nwarrow\\sum F_y=0,\\quad\\implies$</p>\n<p style=\"padding-left: 30px;\">$ N-W \\cos 35^\\circ+P\\sin(\\var{th-35}^\\circ)=0$</p>\n<p style=\"padding-left: 30px;\">$N=W \\cos 35^\\circ-P\\sin( \\var{th-35}^\\circ)$</p>\n<p>To start moving the block up the incline, point friction force down the incline:</p>\n<p></p>\n<p>$+\\!\\!\\swarrow\\sum F_x=0,\\quad\\implies$</p>\n<p>&nbsp; &nbsp; &nbsp; &nbsp;$\\mu_s N-P\\cos(\\theta-35^\\circ)+W\\sin 35^\\circ=0$</p>\n<p style=\"padding-left: 30px;\">$\\mu_s (W \\cos 35^\\circ-P\\sin \\var{th-35}^\\circ)-P\\cos \\var{th-35}^\\circ+W\\sin 35^\\circ=0$</p>\n<p style=\"padding-left: 30px;\">$P=\\dfrac{W\\,(\\sin 35^\\circ+\\mu_s\\cos 35^\\circ)}{(\\cos \\var{th-35}^\\circ+\\mu_s \\sin \\var{th-35}^\\circ)}=\\var{Pmax}\\,\\text{N}\\quad\\blacktriangleleft$</p>\n<p></p>\n<p>To prevent the block from moving down, point friction force up the incline:</p>\n<p>$+\\!\\!\\swarrow\\sum F_x=0,\\quad\\implies$</p>\n<p>&nbsp; &nbsp; &nbsp; &nbsp;$-\\mu_s N-P\\cos(\\theta-35^\\circ)+W\\sin 35^\\circ=0$</p>\n<p style=\"padding-left: 30px;\">$-\\mu_s (W \\cos 35^\\circ-P\\sin \\var{th-35}^\\circ)-P\\cos \\var{th-35}^\\circ+W\\sin 35^\\circ=0$</p>\n<p style=\"padding-left: 30px;\">$P=\\dfrac{W\\,(\\sin 35^\\circ-\\mu_s\\cos 35^\\circ)}{(\\cos\\var{th-35}^\\circ-\\mu_s \\sin \\var{th-35}^\\circ)}=\\var{Pmin}\\,\\text{N}\\quad\\blacktriangleleft$</p>\n<p style=\"padding-left: 30px;\">\\[\\var{Pmin}\\,\\text{N} \\le P \\le \\var{Pmax}\\,\\text{N}\\quad\\blacktriangleleft\\]</p>", "rulesets": {}, "extensions": ["geogebra", "quantities"], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"th": {"name": "th", "group": "Ungrouped variables", "definition": "random(40 .. 80#5)", "description": "", "templateType": "randrange", "can_override": false}, "mu_s": {"name": "mu_s", "group": "Ungrouped variables", "definition": "random(0.15 .. 0.4#0.05)", "description": "", "templateType": "randrange", "can_override": false}, "Pmax": {"name": "Pmax", "group": "Ungrouped variables", "definition": "precround(500*(sin(radians(alpha))+mu_S*cos(radians(35)))/(cos(radians(th-35))+mu_s*sin(radians(th-35))),1)", "description": "", "templateType": "anything", "can_override": false}, "Pmin": {"name": "Pmin", "group": "Ungrouped variables", "definition": "precround(500*(sin(radians(alpha))-mu_S*cos(radians(35)))/(cos(radians(th-35))-mu_s*sin(radians(th-35))),1)", "description": "", "templateType": "anything", "can_override": false}, "params": {"name": "params", "group": "Ungrouped variables", "definition": "[ theta: radians(th), W: 500, alpha: radians(alpha)]", "description": "<p>Currently advice is for W = 500 N, alpha = 35&deg;. &nbsp;Diagram will accept other values.</p>", "templateType": "anything", "can_override": false}, "fbd": {"name": "fbd", "group": "Ungrouped variables", "definition": "view2(geogebra_file(\"friction-block_xtezuTb.ggb\",params))", "description": "", "templateType": "anything", "can_override": false}, "alpha": {"name": "alpha", "group": "Ungrouped variables", "definition": "random(25..35#5)", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["mu_s", "th", "Pmin", "Pmax", "params", "fbd", "alpha"], "variable_groups": [], "functions": {"view2": {"parameters": [["app", "ggbapplet"]], "type": "anything", "language": "javascript", "definition": "// see https://wiki.geogebra.org/en/Reference:GeoGebra_Apps_API for other commands\n\napp.promise.then(function(d) {\n  d.app.setPerspective('D');\n  d.app.setGridVisible(2,false);\n});\nreturn new Numbas.jme.types.ggbapplet(app);\n"}}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Determine:</p>\n<ul>\n<li>The maximum value of $P$ before the block starts to move up the incline, i.e. impending upward motion.</li>\n</ul>\n<p style=\"padding-left: 40px;\">$P_\\text{max}$ = [[0]]</p>\n<ul>\n<li>The minimum value of $P$ required to prevent the block from moving down the incline, i.e. impending downward motion.</li>\n</ul>\n<p style=\"padding-left: 40px;\">$P_\\text{min}$ = [[1]]</p>", "gaps": [{"type": "engineering-answer", "useCustomName": true, "customName": "Pmax", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(Pmax,'N')", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "Pmin", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(Pmin,'N')", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}], "sortAnswers": false}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "Vladimir Vingoradov", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/1862/"}, {"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}, {"name": "Haoyu Huang", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/17657/"}], "resources": ["question-resources/friction-block_xtezuTb.ggb"]}]}], "contributors": [{"name": "Vladimir Vingoradov", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/1862/"}, {"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}, {"name": "Haoyu Huang", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/17657/"}]}