// Numbas version: finer_feedback_settings
{"name": "VM0: Section Cut Method", "extensions": ["jsxgraph", "shear-and-bending-moment-diagrams", "quantities"], "custom_part_types": [], "resources": [["question-resources/FBDS.svg", "FBDS_H6jeTl8.svg"]], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "VM0: Section Cut Method", "tags": ["bending moment", "Mechanics", "mechanics", "shear", "Statics", "statics"], "metadata": {"description": "<p>Determine expressions for internal shear and bending moment as a function of $x$ for a beam with two vertical loads.</p>", "licence": "Creative Commons Attribution-NonCommercial 4.0 International"}, "statement": "<p style=\"text-align: left;\">The $\\var{qty(12, 'm')}$ long beam is loaded with $\\var{qty(B, 'kN')}$ force $B$ and $\\var{qty(C, 'kN')}$ force $C$, as shown.</p>\n<p style=\"text-align: left;\">{diagram}</p>\n<p style=\"text-align: left;\">Derive the equations for the shear and bending moment functions for each of the three segments of the beam.</p>", "advice": "<p>1. &nbsp;Draw a FBD of the whole beam and take moments at $D$ to find $A = \\var{A'}$.</p>\n<p>2. &nbsp;Take cuts in each of the three segments, and draw three free-body diagrams. &nbsp;Show internal shear and moment at the cut with the standard sign convention.</p>\n<p>3. &nbsp;For each FBD, apply $\\Sigma F_y = 0$ and $\\Sigma M_\\text{cut}=0$ to determine expressions for $V(x)$ and $M(x)$ as &nbsp;functions of the forces and $x$.</p>\n<p><img src=\"resources/question-resources/FBDS.svg\" width=\"450\" height=\"517\"/></p>\n<p style=\"text-align: left;\"></p>", "rulesets": {}, "extensions": ["jsxgraph", "quantities", "shear-and-bending-moment-diagrams"], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"B": {"name": "B", "group": "Unnamed group", "definition": "random(100..1200#50)", "description": "", "templateType": "anything", "can_override": false}, "C": {"name": "C", "group": "Unnamed group", "definition": "random(100..1000#50)", "description": "", "templateType": "anything", "can_override": false}, "d1": {"name": "d1", "group": "Unnamed group", "definition": "random(2..8)", "description": "", "templateType": "anything", "can_override": false}, "d2": {"name": "d2", "group": "Unnamed group", "definition": "random(d1+2..10)\n", "description": "", "templateType": "anything", "can_override": false}, "D": {"name": "D", "group": "Unnamed group", "definition": "siground((b  d1 + c  *d2)/12,4)", "description": "", "templateType": "anything", "can_override": false}, "A": {"name": "A", "group": "Unnamed group", "definition": "B+C-D", "description": "", "templateType": "anything", "can_override": false}, "V1": {"name": "V1", "group": "Unnamed group", "definition": "A", "description": "", "templateType": "anything", "can_override": false}, "V2": {"name": "V2", "group": "Unnamed group", "definition": "A-B", "description": "", "templateType": "anything", "can_override": false}, "V3": {"name": "V3", "group": "Unnamed group", "definition": "A-B-C", "description": "", "templateType": "anything", "can_override": false}, "M1": {"name": "M1", "group": "Unnamed group", "definition": "substitute([\"A\": A], expression(\"A x\"))", "description": "", "templateType": "anything", "can_override": false}, "M2": {"name": "M2", "group": "Unnamed group", "definition": "substitute([\"A\": A, \"B\": b, \"d1\": d1], expression(\"A x - B*(x-d1)\"))", "description": "", "templateType": "anything", "can_override": false}, "M3": {"name": "M3", "group": "Unnamed group", "definition": "substitute([\"A\":A, \"B\": B, \"C\":C, \"d1\": d1, \"d2\":d2],\n  expression(\"A x - B*(x-d1) - c*(x- d2)\" )) ", "description": "", "templateType": "anything", "can_override": false}, "A'": {"name": "A'", "group": "Unnamed group", "definition": "qty(trunc(A),'kN')", "description": "", "templateType": "anything", "can_override": false}, "L": {"name": "L", "group": "Unnamed group", "definition": "12", "description": "", "templateType": "anything", "can_override": false}, "diagram": {"name": "diagram", "group": "Ungrouped variables", "definition": "vmloaddiagram(properties)", "description": "", "templateType": "anything", "can_override": false}, "properties": {"name": "properties", "group": "Ungrouped variables", "definition": "[beamLength: L, reactions: [], forces: forces,  moments: [], loads: [], symbols: symbols]", "description": "", "templateType": "anything", "can_override": false}, "forces": {"name": "forces", "group": "Ungrouped variables", "definition": "[[x: d1, value: 1, label: \"$B$\" , visible: true], [x: d2, value: 1, label: \"$C$\", visible: true]]", "description": "", "templateType": "anything", "can_override": false}, "symbols": {"name": "symbols", "group": "Ungrouped variables", "definition": "[[x: 0, type: \"pin\", label: \"$A$\", visible: true],[x: L, type: \"roller\", label: \"$D$\", visible: true], ]", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "c-b > 2", "maxRuns": 100}, "ungrouped_variables": ["diagram", "properties", "forces", "symbols"], "variable_groups": [{"name": "Unnamed group", "variables": ["D", "A", "V1", "V2", "V3", "M1", "M2", "M3", "A'", "L", "B", "C", "d1", "d2"]}], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": true, "customName": "Segment 1", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Shear and bending moment as functions of $x$ on the range: &nbsp;$0 \\text{ m}&lt; x &lt; \\var{d1}$ m</p>\n<p style=\"padding-left: 40px;\">$V_1(x)$ = [[0]] kN<br/><br/>$M_1(x)$ = [[1]]&nbsp; kN-m&nbsp;</p>\n<p></p>", "gaps": [{"type": "jme", "useCustomName": true, "customName": "$V_1$", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "answer": "{V1}", "showPreview": true, "checkingType": "sigfig", "checkingAccuracy": 3, "failureRate": 1, "vsetRangePoints": 5, "vsetRange": [0, 1], "checkVariableNames": false, "singleLetterVariables": false, "allowUnknownFunctions": true, "implicitFunctionComposition": false, "caseSensitive": false, "valuegenerators": []}, {"type": "jme", "useCustomName": true, "customName": "$M_1$", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "answer": "{M1}", "showPreview": true, "checkingType": "sigfig", "checkingAccuracy": 3, "failureRate": 1, "vsetRangePoints": 5, "vsetRange": [0, 1], "checkVariableNames": false, "singleLetterVariables": false, "allowUnknownFunctions": true, "implicitFunctionComposition": false, "caseSensitive": false, "valuegenerators": []}], "sortAnswers": false}, {"type": "gapfill", "useCustomName": true, "customName": "Segment 2", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Shear and bending moment as functions of $x$ on the range: $\\var{d1}\\text{ m} &lt; x &lt; \\var{d2}$ m</p>\n<p style=\"padding-left: 40px;\">$V_2(x)$ = [[1]] kN</p>\n<p style=\"padding-left: 40px;\">$M_2(x)$ = [[0]] kN-m</p>\n<p style=\"padding-left: 40px;\"></p>", "gaps": [{"type": "jme", "useCustomName": true, "customName": "$M_2$", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "answer": "{M2}", "showPreview": true, "checkingType": "sigfig", "checkingAccuracy": 3, "failureRate": 1, "vsetRangePoints": 5, "vsetRange": ["0", "1"], "checkVariableNames": false, "singleLetterVariables": false, "allowUnknownFunctions": true, "implicitFunctionComposition": false, "caseSensitive": false, "valuegenerators": []}, {"type": "jme", "useCustomName": true, "customName": "$V_2$", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "answer": "{V2}", "showPreview": true, "checkingType": "sigfig", "checkingAccuracy": 3, "failureRate": 1, "vsetRangePoints": 5, "vsetRange": [0, 1], "checkVariableNames": false, "singleLetterVariables": false, "allowUnknownFunctions": true, "implicitFunctionComposition": false, "caseSensitive": false, "valuegenerators": []}], "sortAnswers": false}, {"type": "gapfill", "useCustomName": true, "customName": "Segment 3", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Shear and bending moment as functions of $x$ on the range: &nbsp;$\\var{d2}\\text{ m} &lt; x &lt; 12$ m</p>\n<p style=\"padding-left: 40px;\">$V_3(x)$ = [[0]] kN <br/><br/>$ M_3(x)$ =[[1]] kN-m</p>", "gaps": [{"type": "jme", "useCustomName": true, "customName": "$V_3$", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "answer": "{V3}", "showPreview": true, "checkingType": "sigfig", "checkingAccuracy": 3, "failureRate": 1, "vsetRangePoints": 5, "vsetRange": [0, 1], "checkVariableNames": false, "singleLetterVariables": false, "allowUnknownFunctions": true, "implicitFunctionComposition": false, "caseSensitive": false, "valuegenerators": []}, {"type": "jme", "useCustomName": true, "customName": "$M_3$", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "answer": "{M3}", "showPreview": true, "checkingType": "sigfig", "checkingAccuracy": 3, "failureRate": 1, "vsetRangePoints": 5, "vsetRange": [0, 1], "checkVariableNames": false, "singleLetterVariables": false, "allowUnknownFunctions": true, "implicitFunctionComposition": false, "caseSensitive": false, "valuegenerators": []}], "sortAnswers": false}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}], "resources": ["question-resources/FBDS.svg"]}]}], "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}]}