// Numbas version: finer_feedback_settings
{"name": "Machine:  Pruning Shears", "extensions": ["quantities"], "custom_part_types": [{"source": {"pk": 19, "author": {"name": "William Haynes", "pk": 2530}, "edit_page": "/part_type/19/edit"}, "name": "Engineering Accuracy with units", "short_name": "engineering-answer", "description": "<p>A value with units marked right if within an adjustable % error of the correct value. &nbsp;Marked close if within a wider margin of error.</p>", "help_url": "", "input_widget": "string", "input_options": {"correctAnswer": "siground(settings['correctAnswer'],4)", "hint": {"static": true, "value": ""}, "allowEmpty": {"static": true, "value": true}}, "can_be_gap": true, "can_be_step": true, "marking_script": "mark:\nswitch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)\n\ninterpreted_answer:\nqty(student_scalar, student_units)\n\n\n\ncorrect_quantity:\nsettings[\"correctAnswer\"]\n\n\n\ncorrect_units:\nunits(correct_quantity)\n\n\nallowed_notation_styles:\n[\"plain\",\"en\"]\n\nmatch_student_number:\nmatchnumber(studentAnswer,allowed_notation_styles)\n\nstudent_scalar:\nmatch_student_number[1]\n\nstudent_units:\nreplace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")\n\ngood_units:\ntry(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n\n\nstudent_quantity:\nswitch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n\n\npercent_error:\ntry(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   \n\nright:\npercent_error <= settings['right']\n\n\nclose:\nright_sign and percent_error <= settings['close']\n\nright_sign:\nsign(student_scalar) = sign(correct_quantity)", "marking_notes": [{"name": "mark", "description": "This is the main marking note. It should award credit and provide feedback based on the student's answer.", "definition": "switch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)"}, {"name": "interpreted_answer", "description": "A value representing the student's answer to this part.", "definition": "qty(student_scalar, student_units)\n\n"}, {"name": "correct_quantity", "description": "", "definition": "settings[\"correctAnswer\"]\n\n"}, {"name": "correct_units", "description": "", "definition": "units(correct_quantity)\n"}, {"name": "allowed_notation_styles", "description": "", "definition": "[\"plain\",\"en\"]"}, {"name": "match_student_number", "description": "", "definition": "matchnumber(studentAnswer,allowed_notation_styles)"}, {"name": "student_scalar", "description": "", "definition": "match_student_number[1]"}, {"name": "student_units", "description": "<p>Modify&nbsp;the&nbsp;unit portion of&nbsp;the student's answer by</p>\n<p>1. replacing \"ohms\" with \"ohm\" &nbsp;case insensitive</p>\n<p>2.&nbsp;replacing '-'&nbsp;with ' '&nbsp;</p>\n<p>3. replacing '&deg;' with ' deg'&nbsp;</p>\n<p>to allow answers like 10 ft-lb and 30&deg;</p>", "definition": "replace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")"}, {"name": "good_units", "description": "", "definition": "try(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n"}, {"name": "student_quantity", "description": "<p>This fixes the student answer for two common errors. &nbsp;</p>\n<p>If student_units are wrong&nbsp;&nbsp;- replace with correct units</p>\n<p>If student_scalar has the wrong sign - replace with right sign</p>\n<p>If student makes both errors, only one gets fixed.</p>", "definition": "switch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n"}, {"name": "percent_error", "description": "", "definition": "try(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   "}, {"name": "right", "description": "", "definition": "percent_error <= settings['right']\n"}, {"name": "close", "description": "<p>Only marked close if the student actually has the right sign.</p>", "definition": "right_sign and percent_error <= settings['close']"}, {"name": "right_sign", "description": "", "definition": "sign(student_scalar) = sign(correct_quantity) "}], "settings": [{"name": "correctAnswer", "label": "Correct Quantity.", "help_url": "", "hint": "The correct answer given as a JME quantity.", "input_type": "code", "default_value": "", "evaluate": true}, {"name": "right", "label": "% Accuracy for right.", "help_url": "", "hint": "Question will be considered correct if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "0.2", "evaluate": true}, {"name": "close", "label": "% Accuracy for close.", "help_url": "", "hint": "Question will be considered close if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "1.0", "evaluate": true}, {"name": "C1", "label": "Close with units.", "help_url": "", "hint": "Partial Credit for close value with appropriate units. &nbsp;if correct answer is 100 N and close is &plusmn;1%,<br />99 &nbsp;N is accepted.", "input_type": "percent", "default_value": "75"}, {"name": "C2", "label": "No units or wrong sign", "help_url": "", "hint": "Partial credit for forgetting units or using wrong sign.<br />If the correct answer is 100 N, both 100 and -100 N are accepted.", "input_type": "percent", "default_value": "50"}, {"name": "C3", "label": "Close, no units.", "help_url": "", "hint": "Partial Credit for close value but forgotten units.<br />This value would be close&nbsp;if the expected units were provided. &nbsp;If the correct answer is 100 N, and close is &plusmn;1%,<br />99 is accepted.", "input_type": "percent", "default_value": "25"}], "public_availability": "always", "published": true, "extensions": ["quantities"]}], "resources": [["question-resources/shears_2.svg", "shears_2_0VriFhD.svg"]], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Machine:  Pruning Shears", "tags": ["machine", "mechanical advantage", "Mechanics", "mechanics", "Statics", "statics", "SVG"], "metadata": {"description": "<p>Classic machine problem. &nbsp;</p><p>Find cutting force and mechanical advantage of a set of compound pruning shears. &nbsp;This problem has the problem variables directly inserted into an SVG diagram.&nbsp;<br></p>", "licence": "Creative Commons Attribution-ShareAlike 4.0 International"}, "statement": "<p>A {F} force is applied to the handles of the compound pruning shears at the point shown.</p>\n<p>Assume that the cutting forces on the twig act along line $AB$ and answer the following questions. &nbsp;a) What is the force in member $DE\\,?$ b) What is the cutting force at $A$ and $B\\,?$ c) What is the mechanical advantage of this tool$?$&nbsp;</p>\n<table style=\"margin-left: auto; margin-right: auto;\">\n<tbody>\n<tr>\n<td><img src=\"resources/question-resources/shears_2.svg\" width=\"487\" height=\"345\"/></td>\n</tr>\n</tbody>\n</table>\n<p>Note: The mechanical advantage the ratio of the output force to the input force, i.e. the ratio of the cutting force to the gripping force.</p>", "advice": "<p>Recognize that member $DE$ is a two-force member in compression, thus the direction of the forces at $D$ and $E$ are known, and they are&nbsp;equal and opposite. &nbsp;Given: Input force $F = \\var{F}$, geometry of the tool.</p>\n<p>1. Draw a FBD of the lower handle, and take moments at $G$ to find the force in member $DE$.</p>\n<p style=\"padding-left: 40px;\">$DE = \\var{DE}$</p>\n<p>2. Draw a FBD of the upper jaw, and take moments about $C$ to find force $A$. &nbsp;Since the twig is a two-force member, the cutting force at $B$ equals the force at $A$.</p>\n<p style=\"padding-left: 40px;\">$F_A =\\var{FA}$</p>\n<p>3. The mechanical advantage of the tool is the ratio of the output force $F_A$ to input force $F$.</p>\n<p style=\"padding-left: 40px;\">Mechanical Advantage = $\\dfrac{F_A}{F} = \\var{R}:1$</p>\n<p>Checkpoints: &nbsp;</p>\n<ul style=\"padding-left: 40px;\">\n<li>Distance&nbsp;from $C$ to $D$ = {CD},</li>\n<li>Angle of $DE$ = {theta}&deg;, &nbsp;</li>\n<li>Angle of $CD$ = {phi}&deg;,&nbsp;</li>\n<li>Angle $CDE$ = {theta + phi}&deg;,&nbsp;</li>\n<li>$d_\\perp$ =&nbsp;&nbsp;{dperp},&nbsp;</li>\n<li>Moment of $DE$ about $C$ = { siground(F CD sin(radians(theta)),4)}</li>\n</ul>", "rulesets": {}, "extensions": ["quantities"], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"d1": {"name": "d1", "group": "dimensions", "definition": "qty(random(55..70),'mm')", "description": "", "templateType": "anything", "can_override": false}, "d2": {"name": "d2", "group": "dimensions", "definition": "qty(random(120..150),'mm')", "description": "", "templateType": "anything", "can_override": false}, "x1": {"name": "x1", "group": "dimensions", "definition": "qty(random(28..33),'mm')", "description": "", "templateType": "anything", "can_override": false}, "x2": {"name": "x2", "group": "dimensions", "definition": "qty(random(14..18),'mm')", "description": "", "templateType": "anything", "can_override": false}, "y1": {"name": "y1", "group": "dimensions", "definition": "qty(random(14..18),'mm')", "description": "", "templateType": "anything", "can_override": false}, "y2": {"name": "y2", "group": "dimensions", "definition": "qty(random(14..18),'mm')", "description": "", "templateType": "anything", "can_override": false}, "alpha": {"name": "alpha", "group": "dimensions", "definition": "random(24..32)", "description": "", "templateType": "anything", "can_override": false}, "a": {"name": "a", "group": "dimensions", "definition": "alpha+'\u00b0'", "description": "", "templateType": "anything", "can_override": false}, "F": {"name": "F", "group": "dimensions", "definition": "qty(random(30..100#5),'N')", "description": "", "templateType": "anything", "can_override": false}, "FA": {"name": "FA", "group": "Solution", "definition": "precround(DE dperp/ d1,1)", "description": "<p>Cutting force</p>", "templateType": "anything", "can_override": false}, "DE": {"name": "DE", "group": "Solution", "definition": "precround(F d2 /x1 / sin(radians(theta)),1)", "description": "", "templateType": "anything", "can_override": false}, "theta": {"name": "theta", "group": "Solution", "definition": "precround(degrees(arctan(scalar(y1/x2))),1)", "description": "<p>t</p>", "templateType": "anything", "can_override": false}, "dperp": {"name": "dperp", "group": "Solution", "definition": "precround(y2 cos(radians(theta)) + (x1 + x2) sin(radians(theta)),1)", "description": "<p>perpendicular distance from DE to C</p>", "templateType": "anything", "can_override": false}, "CD": {"name": "CD", "group": "Solution", "definition": "qty(precround(sqrt(scalar(y2)^2+ scalar((x1+x2))^2),1),'mm')", "description": "<p>distance from C to D</p>", "templateType": "anything", "can_override": false}, "dperp'": {"name": "dperp'", "group": "Solution", "definition": "precround(CD sin(radians(theta+phi)),1)", "description": "<p>Alternate calculation of dperp&nbsp;</p>", "templateType": "anything", "can_override": false}, "phi": {"name": "phi", "group": "Solution", "definition": "precround(degrees(arctan(scalar(y2/(x1+x2)))),1)", "description": "", "templateType": "anything", "can_override": false}, "R": {"name": "R", "group": "Solution", "definition": "precround(scalar(FA/F),2)", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": [], "variable_groups": [{"name": "dimensions", "variables": ["x1", "x2", "y1", "y2", "d1", "d2", "alpha", "a", "F"]}, {"name": "Solution", "variables": ["FA", "theta", "DE", "CD", "dperp'", "dperp", "phi", "R"]}], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p style=\"padding-left: 40px;\">$DE =$ [[0]]</p>\n<p style=\"padding-left: 40px;\">$A = B = $ &nbsp;[[1]]</p>\n<p style=\"padding-left: 40px;\">Mechanical Advantage = [[2]] : 1.&nbsp;</p>\n<p style=\"padding-left: 40px;\"></p>", "gaps": [{"type": "engineering-answer", "useCustomName": true, "customName": "DE", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "DE", "right": "0.2", "close": "2.0", "C1": "80", "C2": "80", "C3": "60"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "FA", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>s</p>", "settings": {"correctAnswer": "FA", "right": "0.2", "close": "2.0", "C1": "80", "C2": "80", "C3": "60"}}, {"type": "numberentry", "useCustomName": true, "customName": "R", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minValue": "0.99 R", "maxValue": "1.01 R", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}], "sortAnswers": false}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}], "resources": ["question-resources/shears_2.svg"]}]}], "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}]}