// Numbas version: finer_feedback_settings
{"name": "VM7: Symmetric parabolic loading", "extensions": ["jsxgraph", "polynomials", "quantities"], "custom_part_types": [{"source": {"pk": 19, "author": {"name": "William Haynes", "pk": 2530}, "edit_page": "/part_type/19/edit"}, "name": "Engineering Accuracy with units", "short_name": "engineering-answer", "description": "<p>A value with units marked right if within an adjustable % error of the correct value. &nbsp;Marked close if within a wider margin of error.</p>", "help_url": "", "input_widget": "string", "input_options": {"correctAnswer": "siground(settings['correctAnswer'],4)", "hint": {"static": true, "value": ""}, "allowEmpty": {"static": true, "value": true}}, "can_be_gap": true, "can_be_step": true, "marking_script": "mark:\nswitch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)\n\ninterpreted_answer:\nqty(student_scalar, student_units)\n\n\n\ncorrect_quantity:\nsettings[\"correctAnswer\"]\n\n\n\ncorrect_units:\nunits(correct_quantity)\n\n\nallowed_notation_styles:\n[\"plain\",\"en\"]\n\nmatch_student_number:\nmatchnumber(studentAnswer,allowed_notation_styles)\n\nstudent_scalar:\nmatch_student_number[1]\n\nstudent_units:\nreplace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")\n\ngood_units:\ntry(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n\n\nstudent_quantity:\nswitch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n\n\npercent_error:\ntry(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   \n\nright:\npercent_error <= settings['right']\n\n\nclose:\nright_sign and percent_error <= settings['close']\n\nright_sign:\nsign(student_scalar) = sign(correct_quantity)", "marking_notes": [{"name": "mark", "description": "This is the main marking note. It should award credit and provide feedback based on the student's answer.", "definition": "switch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)"}, {"name": "interpreted_answer", "description": "A value representing the student's answer to this part.", "definition": "qty(student_scalar, student_units)\n\n"}, {"name": "correct_quantity", "description": "", "definition": "settings[\"correctAnswer\"]\n\n"}, {"name": "correct_units", "description": "", "definition": "units(correct_quantity)\n"}, {"name": "allowed_notation_styles", "description": "", "definition": "[\"plain\",\"en\"]"}, {"name": "match_student_number", "description": "", "definition": "matchnumber(studentAnswer,allowed_notation_styles)"}, {"name": "student_scalar", "description": "", "definition": "match_student_number[1]"}, {"name": "student_units", "description": "<p>Modify&nbsp;the&nbsp;unit portion of&nbsp;the student's answer by</p>\n<p>1. replacing \"ohms\" with \"ohm\" &nbsp;case insensitive</p>\n<p>2.&nbsp;replacing '-'&nbsp;with ' '&nbsp;</p>\n<p>3. replacing '&deg;' with ' deg'&nbsp;</p>\n<p>to allow answers like 10 ft-lb and 30&deg;</p>", "definition": "replace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")"}, {"name": "good_units", "description": "", "definition": "try(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n"}, {"name": "student_quantity", "description": "<p>This fixes the student answer for two common errors. &nbsp;</p>\n<p>If student_units are wrong&nbsp;&nbsp;- replace with correct units</p>\n<p>If student_scalar has the wrong sign - replace with right sign</p>\n<p>If student makes both errors, only one gets fixed.</p>", "definition": "switch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n"}, {"name": "percent_error", "description": "", "definition": "try(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   "}, {"name": "right", "description": "", "definition": "percent_error <= settings['right']\n"}, {"name": "close", "description": "<p>Only marked close if the student actually has the right sign.</p>", "definition": "right_sign and percent_error <= settings['close']"}, {"name": "right_sign", "description": "", "definition": "sign(student_scalar) = sign(correct_quantity) "}], "settings": [{"name": "correctAnswer", "label": "Correct Quantity.", "help_url": "", "hint": "The correct answer given as a JME quantity.", "input_type": "code", "default_value": "", "evaluate": true}, {"name": "right", "label": "% Accuracy for right.", "help_url": "", "hint": "Question will be considered correct if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "0.2", "evaluate": true}, {"name": "close", "label": "% Accuracy for close.", "help_url": "", "hint": "Question will be considered close if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "1.0", "evaluate": true}, {"name": "C1", "label": "Close with units.", "help_url": "", "hint": "Partial Credit for close value with appropriate units. &nbsp;if correct answer is 100 N and close is &plusmn;1%,<br />99 &nbsp;N is accepted.", "input_type": "percent", "default_value": "75"}, {"name": "C2", "label": "No units or wrong sign", "help_url": "", "hint": "Partial credit for forgetting units or using wrong sign.<br />If the correct answer is 100 N, both 100 and -100 N are accepted.", "input_type": "percent", "default_value": "50"}, {"name": "C3", "label": "Close, no units.", "help_url": "", "hint": "Partial Credit for close value but forgotten units.<br />This value would be close&nbsp;if the expected units were provided. &nbsp;If the correct answer is 100 N, and close is &plusmn;1%,<br />99 is accepted.", "input_type": "percent", "default_value": "25"}], "public_availability": "always", "published": true, "extensions": ["quantities"]}], "resources": [["question-resources/newVM_cA1Sig0.ggb", "newVM_cA1Sig0.ggb"]], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "VM7: Symmetric parabolic loading", "tags": ["gets answer from ggb", "Mechanics", "mechanics", "shear and bending moment", "Statics", "statics"], "metadata": {"description": "<p>&nbsp;Students must derive the shear and bending moment functions for beam loaded with a parabolic distributed load described with a loading function. &nbsp;Diagrams are given, but the students must derive the equations by integration.</p>", "licence": "Creative Commons Attribution-NonCommercial 4.0 International"}, "statement": "<p>A {qty(2b, units[0])} long, simply-supported beam is subjected to a symmetrical parabolic loading, where</p>\n<p>\\[w(x) =\\simplify[fractionNumbers]{{w}} \\quad [\\var{units_string(qty(1,units[3]))}]\\].</p>\n<p><span data-jme-visible=\"debug\">$V(x) =\\simplify[fractionNumbers]{{v}}$ &nbsp;$[\\var{units_string(qty(1,units[1]))}]$.</span></p>\n<p><span data-jme-visible=\"debug\">$M(x) =\\simplify[fractionNumbers]{{m}}$ &nbsp; $[\\var{units_string(qty(1,units[2]))}]$.</span></p>\n<p>{loadBoard}</p>\n<p>{shearBoard}</p>\n<p>{momentBoard}</p>\n<p></p>", "advice": "<p><strong>Given</strong>: &nbsp;The beam is $\\var{qty(L, units[0])}$ long. The loading function is \\[w(x) = \\simplify[fractionNumbers]{{w}}.\\]</p>\n<p>The area under the loading function is the downward force on the beam. Here, because positive values of $w$ point down, we replace $w(x)$ with $-w(x)$ &nbsp;to agree with the standard convention that <em>negative</em> forces point down.</p>\n<p>1. &nbsp;The total load acting on the beam is the integral loading function from $0$ to $L$.</p>\n<p style=\"padding-left: 40px;\">\\[\\begin{align}W &amp;= -\\int_0^L w(x)\\; dx \\\\ <br/>&amp;= \\int_0^L \\simplify[fractionNumbers]{{-w}}\\; dx\\\\<br/>&amp;= &nbsp;\\left. \\simplify[fractionNumbers]{{polynomial(x,[ 0,0,-h/b, h/(b^2)/3])}}\\;\\right |_0^L \\\\<br/>&amp;= \\var{force(load_integral)}<br/>\\end{align}\\]</p>\n<p>By symmetry, the reactions at $x=0$ and $x= L$ are each half this load, pointing up, so:</p>\n<p style=\"padding-left: 40px;\">\\[\\begin{align}A &amp;= B=W/2 \\\\<br/>&nbsp;&amp;=\\var{-force(load_integral/2)}\\end{align}\\]</p>\n<p>2. &nbsp; The shear function $V(x)$ is the integral of the loading function $w(x)$. &nbsp;</p>\n<p style=\"padding-left: 40px;\">\\[\\begin{align} V(x) &amp;= -\\int w(x)\\; dx \\\\<br/>&amp;= \\int \\simplify[fractionNumbers]{{-w}}\\;dx\\\\<br/>&amp;= \\simplify[fractionNumbers]{{polynomial(x,[ 0,0, -h/b, h/(b^2)/3])}} + C_0 \\end{align}\\]</p>\n<p>The constant of integration can be found by evaluating the shear function at $x=0$, where $V(0) = A$ to get</p>\n<p style=\"padding-left: 40px;\">\\[V(0) = C_0 = \\var{force(-load_integral/2)} \\]</p>\n<p>5. &nbsp;The moment function is the integral of the shear function</p>\n<p style=\"padding-left: 40px;\">\\[\\begin{align} M(x) &amp;= \\int V(x)\\; dx \\\\<br/>&amp;= \\int \\simplify[fractionNumbers]{{polynomial(x,[ 0,0, -h/b, h/(b^2)/3])}} + C_0 \\;dx\\\\<br/>&amp; = \\simplify[fractionNumbers]{{m}} + C_1\\\\<br/>\\end{align}\\]</p>\n<p>Since the moment $M(0) = 0$, $C_1 = 0$</p>\n<p>6. The Maximum shear occurs at the ends, where $dV /dx = -w(x) = 0$</p>\n<p style=\"padding-left: 30px;\">\\[\\begin{align}V_{max} &amp;= V(0) = V(L)\\\\<br/>&amp;= \\var{force(maximums[1])}\\end{align} <br/>\\]</p>\n<p>7. The &nbsp;Maximum moment occurs at midpoint, $x = \\var{qty(L/2,units[0])}$, where $dM/dx = 0$</p>\n<p style=\"padding-left: 30px;\">\\[\\begin{align}M_{max} &amp;= M(\\var{b})\\\\ <br/>&amp;= \\var{moment(maximums[2])}\\end{align}\\]</p>\n<p></p>", "rulesets": {}, "extensions": ["jsxgraph", "polynomials", "quantities"], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"units": {"name": "units", "group": "Unnamed group", "definition": "random(['ft','lb', 'ft*lb', 'lb/ft'],['m','kN', 'kN*m', 'kN/m'])", "description": "", "templateType": "anything", "can_override": false}, "b": {"name": "b", "group": "Ungrouped variables", "definition": "0.5 L", "description": "<p>half width of beam \"base\"</p>", "templateType": "anything", "can_override": false}, "h": {"name": "h", "group": "Ungrouped variables", "definition": "random([1,2,2.5,5,7.5,10])random([10,100])", "description": "<p>distributed load at the center of the beam. &nbsp;The \"height\" of the parabola</p>", "templateType": "anything", "can_override": false}, "w": {"name": "w", "group": "Ungrouped variables", "definition": "polynomial(x,[0,2 h/b,-h/(b^2)])", "description": "<p style=\"text-align: left;\">loading function w(x)</p>", "templateType": "anything", "can_override": false}, "F_A": {"name": "F_A", "group": "Unnamed group", "definition": "rational(2/3 b h)", "description": "<p>Reaction at A = half the downward force ( 2/3 bh)</p>\n<p>Reaction at B is the same.</p>", "templateType": "anything", "can_override": false}, "v": {"name": "v", "group": "Ungrouped variables", "definition": "polynomial(x,[ F_A,0, -h/b, h/(b^2)/3])", "description": "<p>the shear function. result of integrating -w(x)</p>", "templateType": "anything", "can_override": false}, "m": {"name": "m", "group": "Ungrouped variables", "definition": "polynomial(x,[0, F_A,0, -h/b/3, h/(b^2)/3/4])", "description": "<p>The moment function m(x), result of integrating v(x)</p>", "templateType": "anything", "can_override": false}, "w_t": {"name": "w_t", "group": "Unnamed group", "definition": "2 F_A", "description": "<p>Total Weight supported by the beam</p>", "templateType": "anything", "can_override": false}, "debug": {"name": "debug", "group": "Ungrouped variables", "definition": "false", "description": "", "templateType": "anything", "can_override": false}, "loadBoard": {"name": "loadBoard", "group": "Ungrouped variables", "definition": "let(max, maximums[0],\n  jsxgraph(\n    480,200,[-2, 1.2 max, 2 b+2, - max/2],\n    merge(\n    [   \"$w(x)$\": ['functiongraph', [load_function, 0, L], [name: \"a\", withLabel: true, strokeWidth: 3, strokeColor: \"#CC0000\", fillColor: \"#CC0000\", \n            fillOpacity: 0.1,label:[fixed: false, fontSize: 14]]]],\narrows,\n    [\"beam\":   ['segment', [[0,0],[L,0]], [strokeWidth: 5, color: \"black\"]] ],\n    [ \"$A$\": ['segment', [[0,-max 0.4], [0,0]], merge(reactionoptions, [label: [fixed: false, color: \"#CC0000\", fontSize: 16]])]],\n    [ \"$B$\": ['segment', [[L,-max 0.4],[L,0],], merge(reactionoptions, [label: [fixed: false, color: \"#CC0000\", fontSize: 16]])]]\n  )\n)\n)", "description": "", "templateType": "anything", "can_override": false}, "load_function": {"name": "load_function", "group": "Ungrouped variables", "definition": "expr(polynomials[0])", "description": "", "templateType": "anything", "can_override": false}, "maximums": {"name": "maximums", "group": "Ungrouped variables", "definition": "let(\n  w, polynomials[0],\n  v, polynomials[1],\n  m, polynomials[2],\n  [eval(w,b),eval(v,0),eval(m,b)])", "description": "<p>load function at peak</p>", "templateType": "anything", "can_override": false}, "L": {"name": "L", "group": "Ungrouped variables", "definition": "20", "description": "<p>beamLength</p>", "templateType": "anything", "can_override": false}, "shear_function": {"name": "shear_function", "group": "Ungrouped variables", "definition": "expr(polynomials[1])", "description": "", "templateType": "anything", "can_override": false}, "moment_function": {"name": "moment_function", "group": "Ungrouped variables", "definition": "expr(polynomials[2])", "description": "", "templateType": "anything", "can_override": false}, "polynomials": {"name": "polynomials", "group": "Ungrouped variables", "definition": "[\npolynomial(x,[0,2 h/b,-h/(b^2)]),\npolynomial(x,[ F_A,0, -h/b, h/(b^2)/3]),\npolynomial(x,[0, F_A,0, -h/b/3, h/(b^2)/3/4])\n]", "description": "", "templateType": "anything", "can_override": false}, "shearBoard": {"name": "shearBoard", "group": "Ungrouped variables", "definition": "let(max, 1.2*maximums[1],\n  jsxgraph(\n    480,200,[-2, 1.1 max, 2 b+2, -max],\n    [   \n        \"$V(x)$\": ['functiongraph', [shear_function, 0, L], [withLabel: true, strokeWidth: 3, strokeColor: \"#CC0000\" ,label:[fixed: false, fontSize: 14]]],\n\n    ]\n))", "description": "", "templateType": "anything", "can_override": false}, "momentBoard": {"name": "momentBoard", "group": "Ungrouped variables", "definition": " let(max, maximums[2], color,  \"#3070AD\", \n  jsxgraph(\n    480,200,[-2, 1.2 max, 2 b+2, -max/3],\n    [\"$M(x)$\": ['functiongraph', [moment_function, 0, L], [withLabel: true, strokeWidth: 3, strokeColor: color, label:[fixed: false, fontSize: 14]]],\n ],\n    ))", "description": "", "templateType": "anything", "can_override": false}, "shearload": {"name": "shearload", "group": "Ungrouped variables", "definition": "eval(-w,L)", "description": "", "templateType": "anything", "can_override": false}, "force": {"name": "force", "group": "Ungrouped variables", "definition": "f -> siground(qty(f,units[1]),4)", "description": "", "templateType": "anything", "can_override": false}, "moment": {"name": "moment", "group": "Ungrouped variables", "definition": "m -> siground(qty(m,units[2]),4)", "description": "", "templateType": "anything", "can_override": false}, "load_integral": {"name": "load_integral", "group": "Ungrouped variables", "definition": "eval(polynomial(x,[0,0, -h/b, h/(b^2)/3]),L)", "description": "", "templateType": "anything", "can_override": false}, "Arrow": {"name": "Arrow", "group": "Ungrouped variables", "definition": "let( w , polynomials[0],\n  (x) ->  ['segment', [[x, eval(w,x)],[x,0]],[ fixed: true, color: \"black\",strokewidth: 1.5,highlight: false, lastArrow: [ type: 4 ]]])", "description": "", "templateType": "anything", "can_override": false}, "reactionOptions": {"name": "reactionOptions", "group": "Ungrouped variables", "definition": "[withLabel: true, \"fixed\": true, \"color\": \"#CC0000\", \"strokewidth\": 4, \"highlight\": false, \"lastArrow\": [ \"type\": 4 ] ]", "description": "", "templateType": "anything", "can_override": false}, "Arrows": {"name": "Arrows", "group": "Ungrouped variables", "definition": "dict(map((i)->([(\"A\"+i),Arrow(i)]), 1..L))", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["debug", "L", "b", "h", "w", "v", "m", "maximums", "load_function", "shear_function", "moment_function", "polynomials", "loadBoard", "shearBoard", "momentBoard", "shearload", "force", "moment", "load_integral", "Arrow", "reactionOptions", "Arrows"], "variable_groups": [{"name": "Unnamed group", "variables": ["w_t", "F_A", "units"]}], "functions": {"show_curvesxxx": {"parameters": [["app", "?"]], "type": "anything", "language": "javascript", "definition": "app.promise.then(function(d) {\n  d.app.setVisible(\"showV\", true,false);\n  d.app.setVisible(\"showM\", true,false);\n  d.app.setValue(\"showV\",true);\n  d.app.setValue(\"showM\",true);\n});\nreturn \"\"//new Numbas.jme.types.ggbapplet(app);\n"}}, "preamble": {"js": "\n", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": true, "customName": "Load and Reactions", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Determine the total weight supported by the beam and the reactions at $A$ and $B$.</p>\n<p>$W = $ [[0]]</p>\n<p>$R_A =$&nbsp;[[1]]</p>\n<p>$R_B = $ [[2]]</p>\n<p><span data-jme-visible=\"debug\">$R_A = R_B =\\simplify[fractionNumbers]{{F_A}}$ = &nbsp;{force(F_A)}</span></p>", "stepsPenalty": "10", "steps": [{"type": "information", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>The total weight is found by integrating the loading function $w(x)$. &nbsp;Positive values indicate a downward load.</p>\n<p style=\"padding-left: 40px;\">$\\begin{align}W &amp;= \\simplify{defint(w,x,0,L)}\\\\ &amp;= \\simplify[fractionNumbers,canonicalOrder]{defint({w},x,0,{2b})}\\end{align}$</p>\n<p>By symmetry, the reactions are half the weight:</p>\n<p style=\"padding-left: 40px;\">$F_A = F_B = \\dfrac{W}{2}$. &nbsp;These point up.</p>"}], "gaps": [{"type": "engineering-answer", "useCustomName": true, "customName": "$W$", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(2 F_A,units[1])", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "$R_A$", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "force(F_A)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "$R_B$", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "force(F_A)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}], "sortAnswers": false}, {"type": "gapfill", "useCustomName": true, "customName": "Shear and Bending Moment Functions", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Determine expressions for:</p>\n<p>$V\\!(x) =$&nbsp;[[0]]</p>\n<p>$M\\!(x) =$&nbsp;[[1]]</p>\n<p>and then plot the shear and bending moment diagrams.</p>\n<p><span data-jme-visible=\"debug\">$V(x) = \\simplify[canonicalOrder,fractionNumbers]{{v}}$</span></p>\n<p><span data-jme-visible=\"debug\">$M(x) = \\simplify[canonicalOrder,fractionNumbers]{{m}}$</span></p>\n<p></p>", "stepsPenalty": "10", "steps": [{"type": "information", "useCustomName": false, "customName": "", "marks": 0, "scripts": {"validate": {"script": "{show_curves(applet)}", "order": "after"}}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>The shear function is found by integrating the loading function. The negative sign is because positive values of w(x) represent a load which points down.</p>\n<p style=\"padding-left: 40px;\">$\\begin{align}V(x) &amp;= &nbsp;\\simplify{-defint(w,x,0,x)} \\\\&amp;= \\simplify[fractionNumbers,canonicalOrder]{defint({-w},x,0,L)} \\\\&amp;= \\simplify[fractionNumbers,canonicalOrder]{{v}}\\end{align}$</p>\n<p>The bending moment function is found by integrating the shear function.</p>\n<p style=\"padding-left: 40px;\">$\\begin{align}M(x) &amp;= \\simplify{defint(V,x,0,x)}\\\\&amp;= \\simplify[fractionNumbers,canonicalOrder]{defint({v},x,0,L)} \\\\&amp;= \\simplify[fractionNumbers,canonicalOrder]{{M}}\\end{align}$</p>\n<p>In both cases, the resulting constant of integration is determined by considering the values at $x=0$.</p>\n<p style=\"padding-left: 40px;\">$V(0) = F_A, \\text{ and } M(0) = 0$</p>\n<p></p>"}], "gaps": [{"type": "jme", "useCustomName": true, "customName": "V(x)", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "answer": "{expr(v)}", "answerSimplification": "fractionNumbers", "showPreview": true, "checkingType": "sigfig", "checkingAccuracy": "2", "failureRate": 1, "vsetRangePoints": 5, "vsetRange": [0, 1], "checkVariableNames": false, "singleLetterVariables": false, "allowUnknownFunctions": true, "implicitFunctionComposition": false, "caseSensitive": false, "valuegenerators": []}, {"type": "jme", "useCustomName": true, "customName": "M(x)", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "answer": "{expr(m)}", "answerSimplification": "fractionNumbers", "showPreview": true, "checkingType": "sigfig", "checkingAccuracy": "2", "failureRate": 1, "vsetRangePoints": 5, "vsetRange": [0, 1], "checkVariableNames": false, "singleLetterVariables": false, "allowUnknownFunctions": true, "implicitFunctionComposition": false, "caseSensitive": false, "valuegenerators": []}], "sortAnswers": false}, {"type": "gapfill", "useCustomName": true, "customName": "Maximum internal loads", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Give the (signed) values of the largest&nbsp;shear and bending moment occuring anywhere within the beam. &nbsp;Use the standard sign convention for shear and bending moments.</p>\n<p>$V_{\\text{max}} = $ &nbsp;[[0]]&nbsp;$\\qquad M_{\\text{max}} = $ &nbsp;[[1]]</p>\n<p><span data-jme-visible=\"debug\">$V_{max} = \\var{siground(qty(eval(v,0),units[1]),4)}$&nbsp;</span></p>\n<p><span data-jme-visible=\"debug\">$M_{max} = \\var{siground(qty(eval(m,b),units[0]+\" \" + units[1]),4)}$</span></p>\n<p></p>", "stepsPenalty": "10", "steps": [{"type": "information", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>The maximum shear occurs at the ends, where $\\simplify{diff(V,x,1)} = -w(x) = 0$ i.e. at $x=0$, or $x=\\var{qty(2b,units[0])}$.</p>\n<p style=\"padding-left: 40px;\">$\\begin{align}V_{max} &amp;= V(0) \\\\ &amp;= \\var{siground(qty(F_A,units[1]),4)}\\end{align}$</p>\n<p>Maximum moment occurs at midpoint, where $\\simplify{diff(M,x,1)} = V = 0$</p>\n<p style=\"padding-left: 40px;\">$\\begin{align}M_{max} &amp;= M(\\var{b}) \\\\&amp;= \\var{siground(qty(eval(m,b),units[0] + ' ' + units[1]),4)}\\end{align}$</p>"}], "gaps": [{"type": "engineering-answer", "useCustomName": true, "customName": "$V_{max}$", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "force(F_A)", "right": "0.5", "close": "1", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "$M_{max}$", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "moment(eval(m,b))", "right": "0.5", "close": "1", "C1": "75", "C2": "50", "C3": "25"}}], "sortAnswers": false}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}], "resources": ["question-resources/newVM_cA1Sig0.ggb"]}]}], "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}]}