// Numbas version: finer_feedback_settings
{"name": "Parallel Axis Theorem: rectangles and triangles", "extensions": ["geogebra", "quantities"], "custom_part_types": [{"source": {"pk": 19, "author": {"name": "William Haynes", "pk": 2530}, "edit_page": "/part_type/19/edit"}, "name": "Engineering Accuracy with units", "short_name": "engineering-answer", "description": "<p>A value with units marked right if within an adjustable % error of the correct value. &nbsp;Marked close if within a wider margin of error.</p>", "help_url": "", "input_widget": "string", "input_options": {"correctAnswer": "siground(settings['correctAnswer'],4)", "hint": {"static": true, "value": ""}, "allowEmpty": {"static": true, "value": true}}, "can_be_gap": true, "can_be_step": true, "marking_script": "mark:\nswitch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)\n\ninterpreted_answer:\nqty(student_scalar, student_units)\n\n\n\ncorrect_quantity:\nsettings[\"correctAnswer\"]\n\n\n\ncorrect_units:\nunits(correct_quantity)\n\n\nallowed_notation_styles:\n[\"plain\",\"en\"]\n\nmatch_student_number:\nmatchnumber(studentAnswer,allowed_notation_styles)\n\nstudent_scalar:\nmatch_student_number[1]\n\nstudent_units:\nreplace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")\n\ngood_units:\ntry(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n\n\nstudent_quantity:\nswitch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n\n\npercent_error:\ntry(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   \n\nright:\npercent_error <= settings['right']\n\n\nclose:\nright_sign and percent_error <= settings['close']\n\nright_sign:\nsign(student_scalar) = sign(correct_quantity)", "marking_notes": [{"name": "mark", "description": "This is the main marking note. It should award credit and provide feedback based on the student's answer.", "definition": "switch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)"}, {"name": "interpreted_answer", "description": "A value representing the student's answer to this part.", "definition": "qty(student_scalar, student_units)\n\n"}, {"name": "correct_quantity", "description": "", "definition": "settings[\"correctAnswer\"]\n\n"}, {"name": "correct_units", "description": "", "definition": "units(correct_quantity)\n"}, {"name": "allowed_notation_styles", "description": "", "definition": "[\"plain\",\"en\"]"}, {"name": "match_student_number", "description": "", "definition": "matchnumber(studentAnswer,allowed_notation_styles)"}, {"name": "student_scalar", "description": "", "definition": "match_student_number[1]"}, {"name": "student_units", "description": "<p>Modify&nbsp;the&nbsp;unit portion of&nbsp;the student's answer by</p>\n<p>1. replacing \"ohms\" with \"ohm\" &nbsp;case insensitive</p>\n<p>2.&nbsp;replacing '-'&nbsp;with ' '&nbsp;</p>\n<p>3. replacing '&deg;' with ' deg'&nbsp;</p>\n<p>to allow answers like 10 ft-lb and 30&deg;</p>", "definition": "replace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")"}, {"name": "good_units", "description": "", "definition": "try(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n"}, {"name": "student_quantity", "description": "<p>This fixes the student answer for two common errors. &nbsp;</p>\n<p>If student_units are wrong&nbsp;&nbsp;- replace with correct units</p>\n<p>If student_scalar has the wrong sign - replace with right sign</p>\n<p>If student makes both errors, only one gets fixed.</p>", "definition": "switch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n"}, {"name": "percent_error", "description": "", "definition": "try(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   "}, {"name": "right", "description": "", "definition": "percent_error <= settings['right']\n"}, {"name": "close", "description": "<p>Only marked close if the student actually has the right sign.</p>", "definition": "right_sign and percent_error <= settings['close']"}, {"name": "right_sign", "description": "", "definition": "sign(student_scalar) = sign(correct_quantity) "}], "settings": [{"name": "correctAnswer", "label": "Correct Quantity.", "help_url": "", "hint": "The correct answer given as a JME quantity.", "input_type": "code", "default_value": "", "evaluate": true}, {"name": "right", "label": "% Accuracy for right.", "help_url": "", "hint": "Question will be considered correct if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "0.2", "evaluate": true}, {"name": "close", "label": "% Accuracy for close.", "help_url": "", "hint": "Question will be considered close if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "1.0", "evaluate": true}, {"name": "C1", "label": "Close with units.", "help_url": "", "hint": "Partial Credit for close value with appropriate units. &nbsp;if correct answer is 100 N and close is &plusmn;1%,<br />99 &nbsp;N is accepted.", "input_type": "percent", "default_value": "75"}, {"name": "C2", "label": "No units or wrong sign", "help_url": "", "hint": "Partial credit for forgetting units or using wrong sign.<br />If the correct answer is 100 N, both 100 and -100 N are accepted.", "input_type": "percent", "default_value": "50"}, {"name": "C3", "label": "Close, no units.", "help_url": "", "hint": "Partial Credit for close value but forgotten units.<br />This value would be close&nbsp;if the expected units were provided. &nbsp;If the correct answer is 100 N, and close is &plusmn;1%,<br />99 is accepted.", "input_type": "percent", "default_value": "25"}], "public_availability": "always", "published": true, "extensions": ["quantities"]}], "resources": [["question-resources/triangle-rectangle.ggb", "triangle-rectangle.ggb"]], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Parallel Axis Theorem: rectangles and triangles", "tags": ["Area moment of inertia", "area moment of inertia", "mechanics", "Mechanics", "parallel axis theorem", "second moment of area", "statics", "Statics"], "metadata": {"description": "<p>Use the parallel axis theorem to find the area moment of inertia of a triangle and a rectangle about various axes.</p>", "licence": "Creative Commons Attribution-NonCommercial 4.0 International"}, "statement": "<p style=\"text-align: left;\">Determine the moment of inertia of the rectangle and the triangle about the $x$ and $y$ axes using the parallel axis theorem when necessary. &nbsp;1 grid square = 1 {units}.</p>\n<table style=\"margin-left: auto; margin-right: auto;\">\n<tbody>\n<tr>\n<td style=\"text-align: center;\">{geogebra_applet('qqws9ffu', [['R1',R1],['R2',R2],['T1',T1],['T2',T2]])}</td>\n</tr>\n</tbody>\n</table>", "advice": "<p><strong>For the x-axis</strong></p>\n<p><strong>Rectangle:</strong></p>\n<p style=\"padding-left: 30px;\">$b = \\var{R_b} \\qquad h = \\var{R_h} \\qquad A = b \\cdot h = \\var{R_area}$</p>\n<p style=\"padding-left: 30px;\">$\\bar{I}_x = \\frac{1}{12}b h^3 = \\var{display(R_Ibarx')}$</p>\n<p style=\"padding-left: 30px;\">$I_x = \\bar{I}_x + A d^2$ where&nbsp;$d$ is the distance from the x-axis to the centroid of the rectangle:&nbsp;$d =\\var{qty(abs(R_c[1]),units)}$</p>\n<p style=\"padding-left: 30px;\">$I_x = \\var{display(R_ibarx')}&nbsp;+ (\\var{R_area})&nbsp;(\\var{qty(abs(R_c[1]),units)})^2 = \\var{display(R_Ix)}$</p>\n<p><strong>Triangle:</strong></p>\n<p style=\"padding-left: 30px;\">$b = \\var{T_b} \\qquad h = \\var{T_h} \\qquad A = b \\cdot h/2 = \\var{T_area}$</p>\n<p style=\"padding-left: 30px;\">$\\bar{I}_x = \\frac{1}{36}b h^3 = \\var{display(T_Ibarx')}$</p>\n<p style=\"padding-left: 30px;\">$I_x = \\bar{I}_x + A d^2$ where $d$ is the distance from the x-axis to the centroid of the triangle: $d =\\var{qty(abs(T_c[1]),units)}$</p>\n<p style=\"padding-left: 30px;\">$I_x = \\var{display(T_ibarx')} + (\\var{T_area}) (\\var{qty(abs(T_c[1]),units)})^2 = \\var{display(T_Ix)}$</p>\n<p><strong>For the y-axis</strong></p>\n<p style=\"padding-left: 30px;\">Follow a similar procedure to find the moments of inertia with respect to the <em>y</em>-axis.</p>\n<p style=\"padding-left: 30px;\"><strong>Rectangle:</strong> $I_y&nbsp;&nbsp;= \\bar{I}_y + A d^2 = \\var{display(R_ibary')}&nbsp;+ (\\var{R_area})&nbsp;(\\var{qty(abs(R_c[0]),units)})^2 = \\var{display(R_Iy)}$</p>\n<p style=\"padding-left: 30px;\"><strong>Triangle:</strong>&nbsp; &nbsp; $I_y &nbsp;= \\bar{I}_y + A d^2 = \\var{display(T_ibary')} + (\\var{T_area}) (\\var{qty(abs(T_c[0]),units)})^2 = \\var{display(T_Iy)}$</p>", "rulesets": {}, "extensions": ["geogebra", "quantities"], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"T_b": {"name": "T_b", "group": "derived ", "definition": "qty(abs(T1[0]-T2[0]),units)", "description": "", "templateType": "anything", "can_override": false}, "T_h": {"name": "T_h", "group": "derived ", "definition": "qty(abs(T1[1]-T2[1]) ,units)", "description": "", "templateType": "anything", "can_override": false}, "R2": {"name": "R2", "group": "inputs", "definition": "R1 + random(-1,1) vector(random(3..8),random(3..8))\n", "description": "", "templateType": "anything", "can_override": false}, "R_c": {"name": "R_c", "group": "derived ", "definition": "R1 + (R2-R1)/2", "description": "<p>centroid</p>", "templateType": "anything", "can_override": false}, "R_xbar": {"name": "R_xbar", "group": "answers", "definition": "qty(R_c[0], units)\n", "description": "", "templateType": "anything", "can_override": false}, "R_Ibarx'": {"name": "R_Ibarx'", "group": "answers", "definition": "r_b r_h r_h r_h/12", "description": "", "templateType": "anything", "can_override": false}, "T1": {"name": "T1", "group": "inputs", "definition": "vector(random(-10..10),random(-10..10))", "description": "", "templateType": "anything", "can_override": false}, "R_area": {"name": "R_area", "group": "answers", "definition": "R_b R_h", "description": "", "templateType": "anything", "can_override": false}, "debug": {"name": "debug", "group": "inputs", "definition": "false", "description": "", "templateType": "anything", "can_override": false}, "units": {"name": "units", "group": "inputs", "definition": "'in'", "description": "", "templateType": "anything", "can_override": false}, "T_Ix": {"name": "T_Ix", "group": "answers", "definition": "T_ibarx' + T_area T_ybar T_ybar", "description": "", "templateType": "anything", "can_override": false}, "R_b": {"name": "R_b", "group": "derived ", "definition": "qty(abs(R1[0]-R2[0]) ,units)", "description": "", "templateType": "anything", "can_override": false}, "T_Ibarx'": {"name": "T_Ibarx'", "group": "answers", "definition": "t_b t_h t_h t_h/36", "description": "", "templateType": "anything", "can_override": false}, "T2": {"name": "T2", "group": "inputs", "definition": "T1 + random(-1,1) vector(random(3,6,9),random(3,6,9))\n\n", "description": "", "templateType": "anything", "can_override": false}, "R_Ibary'": {"name": "R_Ibary'", "group": "answers", "definition": "r_h r_b r_b r_b/12", "description": "", "templateType": "anything", "can_override": false}, "T_Ibary'": {"name": "T_Ibary'", "group": "answers", "definition": "t_h t_b t_b t_b/36", "description": "<p>s</p>", "templateType": "anything", "can_override": false}, "R1": {"name": "R1", "group": "inputs", "definition": "vector(random(-10..10),random(-10..10))", "description": "", "templateType": "anything", "can_override": false}, "R_Iy": {"name": "R_Iy", "group": "answers", "definition": "R_ibary' + R_area R_xbar R_xbar", "description": "", "templateType": "anything", "can_override": false}, "R_ybar": {"name": "R_ybar", "group": "answers", "definition": "qty(R_c[1], units) ", "description": "", "templateType": "anything", "can_override": false}, "T_c": {"name": "T_c", "group": "derived ", "definition": "T1 + (T2-T1)/3", "description": "", "templateType": "anything", "can_override": false}, "T_xbar": {"name": "T_xbar", "group": "answers", "definition": "qty(T_c[0], units) ", "description": "", "templateType": "anything", "can_override": false}, "T_area": {"name": "T_area", "group": "answers", "definition": "T_b T_h / 2\n", "description": "", "templateType": "anything", "can_override": false}, "T_Iy": {"name": "T_Iy", "group": "answers", "definition": "T_ibary' + T_area T_xbar T_xbar", "description": "", "templateType": "anything", "can_override": false}, "R_Ix": {"name": "R_Ix", "group": "answers", "definition": "R_ibarx' + R_area R_ybar R_ybar", "description": "", "templateType": "anything", "can_override": false}, "T_ybar": {"name": "T_ybar", "group": "answers", "definition": "qty(T_c[1], units) ", "description": "", "templateType": "anything", "can_override": false}, "version": {"name": "version", "group": "inputs", "definition": "random('x','y')", "description": "<p>v</p>", "templateType": "anything", "can_override": false}, "R_h": {"name": "R_h", "group": "derived ", "definition": "qty(abs(R1[1]-R2[1]) ,units)", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "not some(map(d=0,d,[R_c[0],R_c[1],T_c[0],T_c[1]]))\n// this prevents a centroid from landing on either axis", "maxRuns": 100}, "ungrouped_variables": [], "variable_groups": [{"name": "inputs", "variables": ["R1", "R2", "T1", "T2", "units", "version", "debug"]}, {"name": "derived ", "variables": ["R_b", "R_h", "T_b", "T_h", "R_c", "T_c"]}, {"name": "answers", "variables": ["R_area", "T_area", "R_xbar", "R_ybar", "T_xbar", "T_ybar", "R_Ibarx'", "R_Ibary'", "T_Ibary'", "T_Ibarx'", "R_Ix", "R_Iy", "T_Ix", "T_Iy"]}], "functions": {"display": {"parameters": [["q", "quantity"]], "type": "quantity", "language": "jme", "definition": "string(siground(q,4))"}}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": true, "customName": "Area and Centroid of Rectangle", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p style=\"text-align: left;\">Find the area and centroidal coordinates of the rectangle.</p>\n<p style=\"text-align: left; padding-left: 40px;\">$A$ = &nbsp;[[0]]&nbsp; &nbsp;</p>\n<p style=\"text-align: left; padding-left: 40px;\">$\\bar{x}$ = [[1]]&nbsp; &nbsp;</p>\n<p style=\"text-align: left; padding-left: 40px;\">$\\bar{y}$ = [[2]]</p>\n<p style=\"text-align: left; padding-left: 40px;\"><span data-jme-visible=\"debug\">$A = \\var{R_area} \\qquad &nbsp; \\bar{x} = \\var{R_xbar} \\qquad \\bar{y} = \\var{R_ybar}$</span></p>\n<p style=\"text-align: left;\"></p>", "gaps": [{"type": "engineering-answer", "useCustomName": false, "customName": "", "marks": "2", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "R_area", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": false, "customName": "", "marks": "2", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "R_xbar", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": false, "customName": "", "marks": "2", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "R_ybar", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}], "sortAnswers": false}, {"type": "gapfill", "useCustomName": true, "customName": "Centroidal Moment of Inertia of Rectangle", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p style=\"text-align: left;\">Find the centroidal moments of inertia&nbsp;for the rectangle:</p>\n<p style=\"text-align: left; padding-left: 40px;\">$\\bar{I}_{x'} = $ [[0]]&nbsp;</p>\n<p style=\"text-align: left; padding-left: 40px;\">$\\bar{I}_{y'} =$ [[1]]&nbsp;</p>\n<p style=\"text-align: left;\"><span data-jme-visible=\"debug\">{display(R_ibarx')}&nbsp;{display(R_ibary')}</span></p>\n<p style=\"text-align: left;\"></p>", "gaps": [{"type": "engineering-answer", "useCustomName": false, "customName": "", "marks": "3", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "R_ibarx'", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": false, "customName": "", "marks": "3", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "R_ibary'", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}], "sortAnswers": false}, {"type": "gapfill", "useCustomName": true, "customName": "Moment of Inertia about $x$ and $y$ Axis", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p style=\"text-align: left;\">Apply the parallel axis theorem $I = \\bar{I} + A d^2 $ to&nbsp;find the moments of inertia with respect to the $x$- and $y$- axes for the rectangle&nbsp;.</p>\n<p style=\"padding-left: 40px;\">$I_x =$ [[0]] &nbsp;</p>\n<p style=\"padding-left: 40px;\">$I_y =$ [[1]]&nbsp;</p>\n<p style=\"padding-left: 40px;\"><span data-jme-visible=\"debug\">{display(R_iy)}&nbsp;{display(R_ix)}</span></p>", "gaps": [{"type": "engineering-answer", "useCustomName": false, "customName": "", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "R_Ix", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": false, "customName": "", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "R_Iy", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}], "sortAnswers": false}, {"type": "gapfill", "useCustomName": true, "customName": "Triangle", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Use a similar process to find the moment of inertia of the triangle with respect to the $\\var{latex(version)}$-axis.</p>\n<p style=\"padding-left: 40px;\">$I_\\var{latex(version)}$ = [[0]] &nbsp;<span data-jme-visible=\"debug\">{display(if(version='x',T_Ix,T_Iy))}</span></p>\n<p></p>", "gaps": [{"type": "engineering-answer", "useCustomName": false, "customName": "", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "if(version='x',T_Ix,T_Iy)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}], "sortAnswers": false}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}], "resources": ["question-resources/triangle-rectangle.ggb"]}]}], "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}]}