// Numbas version: finer_feedback_settings
{"name": "Moment of inertia: built-up beam with angles", "extensions": ["geogebra", "quantities"], "custom_part_types": [{"source": {"pk": 19, "author": {"name": "William Haynes", "pk": 2530}, "edit_page": "/part_type/19/edit"}, "name": "Engineering Accuracy with units", "short_name": "engineering-answer", "description": "<p>A value with units marked right if within an adjustable % error of the correct value. &nbsp;Marked close if within a wider margin of error.</p>", "help_url": "", "input_widget": "string", "input_options": {"correctAnswer": "siground(settings['correctAnswer'],4)", "hint": {"static": true, "value": ""}, "allowEmpty": {"static": true, "value": true}}, "can_be_gap": true, "can_be_step": true, "marking_script": "mark:\nswitch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)\n\ninterpreted_answer:\nqty(student_scalar, student_units)\n\n\n\ncorrect_quantity:\nsettings[\"correctAnswer\"]\n\n\n\ncorrect_units:\nunits(correct_quantity)\n\n\nallowed_notation_styles:\n[\"plain\",\"en\"]\n\nmatch_student_number:\nmatchnumber(studentAnswer,allowed_notation_styles)\n\nstudent_scalar:\nmatch_student_number[1]\n\nstudent_units:\nreplace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")\n\ngood_units:\ntry(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n\n\nstudent_quantity:\nswitch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n\n\npercent_error:\ntry(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   \n\nright:\npercent_error <= settings['right']\n\n\nclose:\nright_sign and percent_error <= settings['close']\n\nright_sign:\nsign(student_scalar) = sign(correct_quantity)", "marking_notes": [{"name": "mark", "description": "This is the main marking note. It should award credit and provide feedback based on the student's answer.", "definition": "switch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)"}, {"name": "interpreted_answer", "description": "A value representing the student's answer to this part.", "definition": "qty(student_scalar, student_units)\n\n"}, {"name": "correct_quantity", "description": "", "definition": "settings[\"correctAnswer\"]\n\n"}, {"name": "correct_units", "description": "", "definition": "units(correct_quantity)\n"}, {"name": "allowed_notation_styles", "description": "", "definition": "[\"plain\",\"en\"]"}, {"name": "match_student_number", "description": "", "definition": "matchnumber(studentAnswer,allowed_notation_styles)"}, {"name": "student_scalar", "description": "", "definition": "match_student_number[1]"}, {"name": "student_units", "description": "<p>Modify&nbsp;the&nbsp;unit portion of&nbsp;the student's answer by</p>\n<p>1. replacing \"ohms\" with \"ohm\" &nbsp;case insensitive</p>\n<p>2.&nbsp;replacing '-'&nbsp;with ' '&nbsp;</p>\n<p>3. replacing '&deg;' with ' deg'&nbsp;</p>\n<p>to allow answers like 10 ft-lb and 30&deg;</p>", "definition": "replace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")"}, {"name": "good_units", "description": "", "definition": "try(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n"}, {"name": "student_quantity", "description": "<p>This fixes the student answer for two common errors. &nbsp;</p>\n<p>If student_units are wrong&nbsp;&nbsp;- replace with correct units</p>\n<p>If student_scalar has the wrong sign - replace with right sign</p>\n<p>If student makes both errors, only one gets fixed.</p>", "definition": "switch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n"}, {"name": "percent_error", "description": "", "definition": "try(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   "}, {"name": "right", "description": "", "definition": "percent_error <= settings['right']\n"}, {"name": "close", "description": "<p>Only marked close if the student actually has the right sign.</p>", "definition": "right_sign and percent_error <= settings['close']"}, {"name": "right_sign", "description": "", "definition": "sign(student_scalar) = sign(correct_quantity) "}], "settings": [{"name": "correctAnswer", "label": "Correct Quantity.", "help_url": "", "hint": "The correct answer given as a JME quantity.", "input_type": "code", "default_value": "", "evaluate": true}, {"name": "right", "label": "% Accuracy for right.", "help_url": "", "hint": "Question will be considered correct if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "0.2", "evaluate": true}, {"name": "close", "label": "% Accuracy for close.", "help_url": "", "hint": "Question will be considered close if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "1.0", "evaluate": true}, {"name": "C1", "label": "Close with units.", "help_url": "", "hint": "Partial Credit for close value with appropriate units. &nbsp;if correct answer is 100 N and close is &plusmn;1%,<br />99 &nbsp;N is accepted.", "input_type": "percent", "default_value": "75"}, {"name": "C2", "label": "No units or wrong sign", "help_url": "", "hint": "Partial credit for forgetting units or using wrong sign.<br />If the correct answer is 100 N, both 100 and -100 N are accepted.", "input_type": "percent", "default_value": "50"}, {"name": "C3", "label": "Close, no units.", "help_url": "", "hint": "Partial Credit for close value but forgotten units.<br />This value would be close&nbsp;if the expected units were provided. &nbsp;If the correct answer is 100 N, and close is &plusmn;1%,<br />99 is accepted.", "input_type": "percent", "default_value": "25"}], "public_availability": "always", "published": true, "extensions": ["quantities"]}], "resources": [["question-resources/plate-angles.ggb", "plate-angles.ggb"]], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Moment of inertia: built-up beam with angles", "tags": ["beam", "Mechanics", "mechanics", "moment of inertia", "neutral axis", "parallel axis theorem", "standard sections", "Statics", "statics"], "metadata": {"description": "<p>Find the centroid and the centroidal moments of inertia for a beam composed&nbsp;of a flat plate and&nbsp;two angle sections.</p>", "licence": "Creative Commons Attribution-NonCommercial 4.0 International"}, "statement": "<p>A built-up beam is made up of two {name} angles&nbsp;welded to a {b} $\\times$ {h} plate, as shown in cross section below. &nbsp;</p>\n<table style=\"margin-left: auto; margin-right: auto;\">\n<tbody>\n<tr>\n<td>{geogebra_applet('vvcydmsb',['A': A,'d': scalar(depth),'C':[ definition: C, fixed: true]])}</td>\n</tr>\n</tbody>\n</table>\n<p>Use this&nbsp;<a href=\"https://weh.maritime.edu/mechanics/HW/reference/unequalangle.html\">table of properties</a>&nbsp;for the angle sections.</p>", "advice": "<p>Let the rectangle be designated by subscript $R$, the rectangle by $L$, and the total composite shape by $T$.</p>\n<p><strong>Find rectangle properties</strong></p>\n<p style=\"padding-left: 30px;\">$b = \\var{b}, h = \\var{h}$</p>\n<p style=\"padding-left: 30px;\">$A_R = b&nbsp;\\cdot h = \\var{A_R}$</p>\n<p style=\"padding-left: 30px;\">$\\bar{y}_R = h/2 = \\var{ybar_R}$</p>\n<p><strong>Find the angle properties</strong></p>\n<p style=\"padding-left: 30px;\">From the table:</p>\n<p style=\"padding-left: 60px;\">$A_L= \\var{A_L}$,</p>\n<p style=\"padding-left: 60px;\">$\\bar{x}=\\var{xbar}, \\bar{y} = \\var{ybar}$, &nbsp;</p>\n<p style=\"padding-left: 60px;\">$I_{xx} = \\var{Ixx}, I_{yy} = \\var{Iyy}$</p>\n<p><strong>Find the centroid of the angle shapes</strong></p>\n<p style=\"padding-left: 30px;\">Note that&nbsp;the angle shape is rotated 90&deg; from the table diagram.</p>\n<p style=\"padding-left: 30px;\">Measuring from the&nbsp;bottom center of the rectangle:</p>\n<p style=\"padding-left: 30px;\">$\\bar{y}_L = h + \\bar{x} = \\var{ybar_L}$&nbsp;</p>\n<p style=\"padding-left: 30px;\">$\\bar{x}_L = \\pm (b/2 - \\bar{y}) = \\pm \\var{xbar_L}$</p>\n<p><strong>Find the centroid of composite shape</strong></p>\n<p style=\"padding-left: 30px;\">Measuring from the bottom center of the rectangle:</p>\n<p style=\"padding-left: 30px;\">$\\begin{align}A_T &amp;= &nbsp;A_R + 2\\, A_L \\\\&amp;= \\var{A_R} + 2 \\,(\\var{A_L})\\\\ &amp;=\\var{A_T}\\end{align}\\\\$</p>\n<p><strong>&nbsp; &nbsp; &nbsp; &nbsp;</strong>$\\bar{x}_T = 0$, by symmetry.</p>\n<p style=\"padding-left: 30px;\">$\\begin{align} &nbsp;\\bar{y}_T &amp;= \\dfrac{\\Sigma A_i \\bar{y}_i}{\\Sigma A_i}\\\\ &amp;=\\dfrac{A_R \\bar{y}_R + 2\\, A_L \\bar{y}_L}{A_T}\\\\&amp;= \\dfrac{(\\var{A_R}) (\\var{ybar_R}) + 2\\, (\\var{A_L}) (\\var{ybar_L})}{\\var{A_T} }\\\\&amp;=\\var{disp(ybar_T)} = \\bar{y}\\end{align}$&nbsp;</p>\n<p><strong>Find the moment of inertia about the centroidal <em>y-</em> axis</strong></p>\n<p style=\"padding-left: 30px;\">$\\begin{align} \\bar{I}_y &amp;= (I_y)_R + 2\\, (I_y)_L\\\\&amp;= \\frac{h b^3}{12} + 2\\, [\\bar{I} + A d^2]_L\\\\ &amp;=\\frac{h b^3}{12} + 2\\, [I_{xx} + A_L (\\bar{x}_L)^2]\\\\ &amp;=\\frac{(\\var{h})(\\var{b})^3}{12} + 2\\, [\\var{Ixx} + (\\var{A_L}) (\\var{xbar_L})^2]\\\\&amp;=\\var{disp(Iy_R)} + 2\\,[\\var{disp(Iy_L)}] \\\\&amp;=\\var{disp(Iy_T)}\\end{align}$</p>\n<p><strong>Find the moment of inertia about the x-axis</strong></p>\n<p style=\"padding-left: 30px;\">$\\begin{align} I_x &amp;= (I_x)_R + 2\\, (I_x)_L\\\\&amp;= \\frac{b&nbsp;h^3}{3} + 2\\, [\\bar{I} + A d^2]_L\\\\ &amp;=\\frac{b&nbsp;h^3}{3} + 2\\, [I_{yy} + A_L (\\bar{y}_L)^2]\\\\ &amp;=\\frac{(\\var{b})(\\var{h})^3}{3} + 2\\, [\\var{Iyy} + (\\var{A_L}) (\\var{ybar_L})^2]\\\\&amp;= \\var{disp(Ix_r)} + 2\\, [\\var{disp(Ix_L)}]\\\\&amp;=\\var{disp(Ix_T)}\\end{align}$</p>\n<p><strong>Use the parallel axis theorem to find the moment of inertia about the centroidal x'-axis</strong></p>\n<p style=\"padding-left: 30px;\">$\\begin{align} I &amp;= \\bar{I} + A d^2 \\\\ \\bar{I}_{x'} &amp;= I_x - A d^2\\\\&nbsp;&amp;=I_x - A_T (\\bar{y}_T)^2\\\\ &amp;=\\var{disp(Ix_T)}&nbsp;- (\\var{A_T}) (\\var{disp(ybar_T)})^2\\\\&amp;=\\var{disp(Ibarx')}\\end{align}$</p>\n<p><strong>Find the radius of gyration with respect to the $x'$- axis.</strong></p>\n<p style=\"padding-left: 30px;\">$\\begin{align} k_{x'} &amp;= \\sqrt{\\dfrac{\\bar{I}_{x'}}{A_T}} \\\\&amp;=\\sqrt{\\dfrac{\\var{disp(Ibarx')}}{\\var{A_T}}} \\\\&amp;=\\var{disp(kx')}\\end{align}$&nbsp;</p>", "rulesets": {}, "extensions": ["geogebra", "quantities"], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"debug": {"name": "debug", "group": "Setup values", "definition": "false", "description": "", "templateType": "anything", "can_override": false}, "A": {"name": "A", "group": "Setup values", "definition": "vector(scalar(b)/2,scalar(h))", "description": "<p>coordinates of top right corner of rectangle</p>", "templateType": "anything", "can_override": false}, "Ibarx'": {"name": "Ibarx'", "group": "MOI", "definition": "Ix_T - A_T ybar_T ybar_T", "description": "", "templateType": "anything", "can_override": false}, "ybar_T": {"name": "ybar_T", "group": "centroid", "definition": "Qx/A_T", "description": "", "templateType": "anything", "can_override": false}, "index": {"name": "index", "group": "Setup values", "definition": "random(0..14)", "description": "<p>i</p>", "templateType": "anything", "can_override": false}, "h": {"name": "h", "group": "Setup values", "definition": "qty(random(0.5..2#0.5),'in')", "description": "<p>height of rectangle</p>", "templateType": "anything", "can_override": false}, "xbar_L": {"name": "xbar_L", "group": "centroid", "definition": "b/2 - ybar", "description": "", "templateType": "anything", "can_override": false}, "A_T": {"name": "A_T", "group": "centroid", "definition": "A_R + 2 A_L", "description": "", "templateType": "anything", "can_override": false}, "name": {"name": "name", "group": "angle properties", "definition": "[\"L9  $\\\\times$  4  $\\\\times$  1/2\",\"L8  $\\\\times$  6  $\\\\times$  5/8\",\"L8  $\\\\times$  6  $\\\\times$  1/2\",\"L8  $\\\\times$  4  $\\\\times$  3/4\",\"L8  $\\\\times$  4  $\\\\times$  1/2\",\"L7  $\\\\times$  4  $\\\\times$  3/4\",\"L7  $\\\\times$  4  $\\\\times$  5/8\",\"L7  $\\\\times$  4  $\\\\times$  1/2\",\"L7  $\\\\times$  4  $\\\\times$  3/8\",\"L6  $\\\\times$  4  $\\\\times$  7/8\",\"L6  $\\\\times$  4  $\\\\times$  3/4\",\"L6  $\\\\times$  4  $\\\\times$  1/2\",\"L6  $\\\\times$  4  $\\\\times$  3/8\",\"L5  $\\\\times$  3  $\\\\times$  5/8\",\"L5  $\\\\times$  3  $\\\\times$  1/2\"][index]", "description": "", "templateType": "anything", "can_override": false}, "Ix_L": {"name": "Ix_L", "group": "MOI", "definition": "Iyy + A_L ybar_L ybar_L", "description": "", "templateType": "anything", "can_override": false}, "C": {"name": "C", "group": "Setup values", "definition": "vector(0,scalar(ybar_T))", "description": "<p>centroid of composite shape as a point for geogebra</p>", "templateType": "anything", "can_override": false}, "Iy_L": {"name": "Iy_L", "group": "MOI", "definition": "Ixx + A_L xbar_L xbar_L", "description": "", "templateType": "anything", "can_override": false}, "kx'": {"name": "kx'", "group": "MOI", "definition": "qty(sqrt(scalar(Ibarx'/A_T)),'in')", "description": "", "templateType": "anything", "can_override": false}, "ybar": {"name": "ybar", "group": "angle properties", "definition": "qty([3.31,2.52,2.47,2.95,2.86,2.51,2.46,2.42,2.37,2.12,2.08,1.99,1.94,1.8,1.75][index],'in')", "description": "", "templateType": "anything", "can_override": false}, "Ixx": {"name": "Ixx", "group": "angle properties", "definition": "qty([53.2,54.1,44.3,54.9,38.5,37.8,32.4,26.7,20.6,27.7,24.5,17.4,13.5,11.4,9.45][index],'in^4')", "description": "", "templateType": "anything", "can_override": false}, "b": {"name": "b", "group": "Setup values", "definition": "qty(random(2..4#0.5) scalar(depth),'in')", "description": "<p>base width of rectangle</p>", "templateType": "anything", "can_override": false}, "Iy_T": {"name": "Iy_T", "group": "MOI", "definition": "Iy_R + 2 Iy_L", "description": "", "templateType": "anything", "can_override": false}, "ybar_R": {"name": "ybar_R", "group": "centroid", "definition": "h/2", "description": "", "templateType": "anything", "can_override": false}, "A_L": {"name": "A_L", "group": "angle properties", "definition": "qty([6.25,8.36,6.75,8.44,5.75,7.69,6.48,5.25,3.98,7.98,6.94,4.75,3.61,4.61,3.75][index],'in^2')", "description": "", "templateType": "anything", "can_override": false}, "Qx": {"name": "Qx", "group": "centroid", "definition": "A_R ybar_R + 2 A_L ybar_L", "description": "", "templateType": "anything", "can_override": false}, "Iy_R": {"name": "Iy_R", "group": "MOI", "definition": "h b b b/12", "description": "", "templateType": "anything", "can_override": false}, "depth": {"name": "depth", "group": "angle properties", "definition": "qty([9,8,8,8,8,7,7,7,7,6,6,6,6,5,5][index],'in')", "description": "", "templateType": "anything", "can_override": false}, "Ix_R": {"name": "Ix_R", "group": "MOI", "definition": "b h h h/3", "description": "", "templateType": "anything", "can_override": false}, "ybar_L": {"name": "ybar_L", "group": "centroid", "definition": "h + xbar", "description": "", "templateType": "anything", "can_override": false}, "Iyy": {"name": "Iyy", "group": "angle properties", "definition": "qty([6.92,26.3,21.7,9.36,6.74,9.05,7.84,6.53,5.1,9.75,8.68,6.27,4.9,3.06,2.58][index],'in^4')", "description": "", "templateType": "anything", "can_override": false}, "Ix_T": {"name": "Ix_T", "group": "MOI", "definition": "Ix_R + 2 Ix_L", "description": "", "templateType": "anything", "can_override": false}, "xbar": {"name": "xbar", "group": "angle properties", "definition": "qty([0.81,1.52,1.47,0.953,0.859,1.01,0.963,0.917,0.87,1.12,1.08,0.987,0.941,0.796,0.75][index],'in')", "description": "", "templateType": "anything", "can_override": false}, "A_R": {"name": "A_R", "group": "centroid", "definition": "b h", "description": "", "templateType": "anything", "can_override": false}, "applet": {"name": "applet", "group": "Setup values", "definition": "geogebra_file('resources/question-resources/plate-angles.ggb',[\n    'A': A,\n    'd': scalar(depth),\n    'C':[definition:  C, fixed: true]\n  ]\n)", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": [], "variable_groups": [{"name": "angle properties", "variables": ["name", "depth", "A_L", "Ixx", "ybar", "Iyy", "xbar"]}, {"name": "centroid", "variables": ["A_R", "A_T", "ybar_R", "ybar_L", "xbar_L", "Qx", "ybar_T"]}, {"name": "MOI", "variables": ["Ix_R", "Ix_L", "Ix_T", "Iy_R", "Iy_L", "Iy_T", "Ibarx'", "kx'"]}, {"name": "Setup values", "variables": ["index", "b", "h", "A", "C", "debug", "applet"]}], "functions": {"disp": {"parameters": [["q", "quantity"]], "type": "quantity", "language": "jme", "definition": "siground(q,4)"}}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Find the area of the composite shape.</p>\n<p style=\"padding-left: 40px;\">$A_T = $&nbsp;[[0]] <span data-jme-visible=\"debug\">{A_T}</span></p>\n<p>Find the distance between the $x$ axis at&nbsp;the bottom of the plate&nbsp;and the parallel $x'$ axis passing&nbsp;through&nbsp;the centroid of the composite shape.</p>\n<p style=\"padding-left: 40px;\">$\\bar{y} = $&nbsp;[[1]] <span data-jme-visible=\"debug\">{ybar_T}</span></p>\n<p>Find the moment of inertia of the composite shape about the centroidal $y$ axis.</p>\n<p style=\"padding-left: 40px;\">$\\bar{I}_y = $&nbsp;[[2]] <span data-jme-visible=\"debug\">{Iy_T}</span></p>\n<p>Find the moment of inertia of the composite shape about the $x$-axis&nbsp;</p>\n<p style=\"padding-left: 40px;\">$I_x=$&nbsp;[[3]] <span data-jme-visible=\"debug\">{Ix_T}</span></p>\n<p>Use the parallel axis theorem to find the moment of inertia of the composite shape about the centroidal $x'$-axis.</p>\n<p style=\"padding-left: 40px;\">$\\bar{I}_{x'}= $&nbsp;[[4]] <span data-jme-visible=\"debug\">{Ibarx'}</span></p>\n<p>Find the corresponding radius of gyration</p>\n<p style=\"padding-left: 40px;\">$k_{x'} = $&nbsp;[[5]] <span data-jme-visible=\"debug\">{kx'}</span></p>", "gaps": [{"type": "engineering-answer", "useCustomName": true, "customName": "$A$", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "A_T", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "$\\bar{y}$", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "ybar_T", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "$\\bar{I}_y$", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "Iy_T", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "$\\bar{I}_x$", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "Ix_T", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "$\\bar{I}_{x'}$", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "Ibarx'", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "$k_{x'}$", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "kx'", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}], "sortAnswers": false}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}], "resources": ["question-resources/plate-angles.ggb"]}]}], "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}]}