// Numbas version: finer_feedback_settings
{"name": "Truss: maximum load", "extensions": ["quantities"], "custom_part_types": [{"source": {"pk": 19, "author": {"name": "William Haynes", "pk": 2530}, "edit_page": "/part_type/19/edit"}, "name": "Engineering Accuracy with units", "short_name": "engineering-answer", "description": "<p>A value with units marked right if within an adjustable % error of the correct value. &nbsp;Marked close if within a wider margin of error.</p>", "help_url": "", "input_widget": "string", "input_options": {"correctAnswer": "siground(settings['correctAnswer'],4)", "hint": {"static": true, "value": ""}, "allowEmpty": {"static": true, "value": true}}, "can_be_gap": true, "can_be_step": true, "marking_script": "mark:\nswitch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)\n\ninterpreted_answer:\nqty(student_scalar, student_units)\n\n\n\ncorrect_quantity:\nsettings[\"correctAnswer\"]\n\n\n\ncorrect_units:\nunits(correct_quantity)\n\n\nallowed_notation_styles:\n[\"plain\",\"en\"]\n\nmatch_student_number:\nmatchnumber(studentAnswer,allowed_notation_styles)\n\nstudent_scalar:\nmatch_student_number[1]\n\nstudent_units:\nreplace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")\n\ngood_units:\ntry(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n\n\nstudent_quantity:\nswitch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n\n\npercent_error:\ntry(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   \n\nright:\npercent_error <= settings['right']\n\n\nclose:\nright_sign and percent_error <= settings['close']\n\nright_sign:\nsign(student_scalar) = sign(correct_quantity)", "marking_notes": [{"name": "mark", "description": "This is the main marking note. It should award credit and provide feedback based on the student's answer.", "definition": "switch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)"}, {"name": "interpreted_answer", "description": "A value representing the student's answer to this part.", "definition": "qty(student_scalar, student_units)\n\n"}, {"name": "correct_quantity", "description": "", "definition": "settings[\"correctAnswer\"]\n\n"}, {"name": "correct_units", "description": "", "definition": "units(correct_quantity)\n"}, {"name": "allowed_notation_styles", "description": "", "definition": "[\"plain\",\"en\"]"}, {"name": "match_student_number", "description": "", "definition": "matchnumber(studentAnswer,allowed_notation_styles)"}, {"name": "student_scalar", "description": "", "definition": "match_student_number[1]"}, {"name": "student_units", "description": "<p>Modify&nbsp;the&nbsp;unit portion of&nbsp;the student's answer by</p>\n<p>1. replacing \"ohms\" with \"ohm\" &nbsp;case insensitive</p>\n<p>2.&nbsp;replacing '-'&nbsp;with ' '&nbsp;</p>\n<p>3. replacing '&deg;' with ' deg'&nbsp;</p>\n<p>to allow answers like 10 ft-lb and 30&deg;</p>", "definition": "replace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")"}, {"name": "good_units", "description": "", "definition": "try(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n"}, {"name": "student_quantity", "description": "<p>This fixes the student answer for two common errors. &nbsp;</p>\n<p>If student_units are wrong&nbsp;&nbsp;- replace with correct units</p>\n<p>If student_scalar has the wrong sign - replace with right sign</p>\n<p>If student makes both errors, only one gets fixed.</p>", "definition": "switch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n"}, {"name": "percent_error", "description": "", "definition": "try(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   "}, {"name": "right", "description": "", "definition": "percent_error <= settings['right']\n"}, {"name": "close", "description": "<p>Only marked close if the student actually has the right sign.</p>", "definition": "right_sign and percent_error <= settings['close']"}, {"name": "right_sign", "description": "", "definition": "sign(student_scalar) = sign(correct_quantity) "}], "settings": [{"name": "correctAnswer", "label": "Correct Quantity.", "help_url": "", "hint": "The correct answer given as a JME quantity.", "input_type": "code", "default_value": "", "evaluate": true}, {"name": "right", "label": "% Accuracy for right.", "help_url": "", "hint": "Question will be considered correct if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "0.2", "evaluate": true}, {"name": "close", "label": "% Accuracy for close.", "help_url": "", "hint": "Question will be considered close if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "1.0", "evaluate": true}, {"name": "C1", "label": "Close with units.", "help_url": "", "hint": "Partial Credit for close value with appropriate units. &nbsp;if correct answer is 100 N and close is &plusmn;1%,<br />99 &nbsp;N is accepted.", "input_type": "percent", "default_value": "75"}, {"name": "C2", "label": "No units or wrong sign", "help_url": "", "hint": "Partial credit for forgetting units or using wrong sign.<br />If the correct answer is 100 N, both 100 and -100 N are accepted.", "input_type": "percent", "default_value": "50"}, {"name": "C3", "label": "Close, no units.", "help_url": "", "hint": "Partial Credit for close value but forgotten units.<br />This value would be close&nbsp;if the expected units were provided. &nbsp;If the correct answer is 100 N, and close is &plusmn;1%,<br />99 is accepted.", "input_type": "percent", "default_value": "25"}], "public_availability": "always", "published": true, "extensions": ["quantities"]}], "resources": [["question-resources/truss150_ST2n1NE.png", "truss150_ST2n1NE.png"], ["question-resources/FBD1_0gNiD65.png", "FBD1_0gNiD65.png"], ["question-resources/FBD2_brSZxtl.png", "FBD2_brSZxtl.png"], ["question-resources/FBD3_IU8eiX9.png", "FBD3_IU8eiX9.png"], ["question-resources/FBD2_Vmzm7ke.png", "FBD2_Vmzm7ke.png"]], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Truss: maximum load", "tags": ["Mechanics", "mechanics", "Statics", "statics", "truss", "zero-force-member"], "metadata": {"description": "Determine the maximum load on a truss given two constraints. &nbsp;Truss &nbsp;includes zero-force members.", "licence": "Creative Commons Attribution-NonCommercial 4.0 International"}, "statement": "<p>Determine the maximum mass&nbsp;$m$ which may be applied to the truss without exceeding {tmax} kN in tension or {cmax} kN in compression in any member.</p>\n<table style=\"margin-left: auto; margin-right: auto;\">\n<tbody>\n<tr>\n<td><img src=\"resources/question-resources/truss150_ST2n1NE.png\" width=\"354\" height=\"354\"/></td>\n</tr>\n</tbody>\n</table>", "advice": "<p>Note that members $AE$, $DF$ and $AF$ are zero&nbsp;force members, that $AD = DC$ and $BE = EF = FC$, and that angle $\\beta = 45&deg;$</p>\n<p>Let $T$ be the weight of the load which has a mass $m$. &nbsp; &nbsp;$T = m g$.</p>\n<p><strong></strong></p>\n<p><strong>Joint C</strong></p>\n<p style=\"padding-left: 30px;\">Draw a FBD of joint $C$. &nbsp;The tensile forces on the pulley are equivalent to forces acting on the pin $C$.</p>\n<p><img src=\"resources/question-resources/FBD1_0gNiD65.png\" style=\"width: 263px; height: 216px;\"/></p>\n<p></p>\n<p style=\"padding-left: 30px;\">The two tension forces acting on joint $C$ resolve to a single force $T' = \\sqrt{2} T$ acting down and to the left&nbsp;at a 45&deg; angle. Redraw the FBD to this.&nbsp;</p>\n<p><img src=\"resources/question-resources/FBD2_Vmzm7ke.png\" style=\"width: 249px; height: 203px;\"/></p>\n<p style=\"padding-left: 30px;\">By inspection you can conclude:</p>\n<p style=\"padding-left: 60px;\">$DC = 0$ &nbsp;</p>\n<p style=\"padding-left: 60px;\">$FC = T'$ (compression).</p>\n<p><strong>Joint B</strong></p>\n<p style=\"padding-left: 30px;\">Draw a&nbsp;free body diagram of joint $B$. &nbsp;</p>\n<p style=\"padding-left: 30px;\"><img src=\"resources/question-resources/FBD3_IU8eiX9.png\" style=\"width: 279px; height: 228px;\"/></p>\n<p style=\"padding-left: 30px;\">From $\\Sigma F_y = 0 $&nbsp;&nbsp;you conclude that</p>\n<p style=\"padding-left: 60px;\">&nbsp;$AB = T' \\sin{45&deg;} = T$ (tension).</p>\n<p style=\"padding-left: 60px;\"></p>\n<p><strong>Apply Design Conditions</strong></p>\n<p style=\"padding-left: 30px;\">Since we have two design constraints, you will need to check them both.</p>\n<p></p>\n<p style=\"padding-left: 30px;\"><strong></strong><em>For Compression</em></p>\n<p style=\"padding-left: 60px;\">$FC = T' = \\sqrt{2} T=\\sqrt{2} m g $</p>\n<p style=\"padding-left: 60px;\">$FC&nbsp;&lt; &nbsp;\\var{cmax}$ &nbsp;kN</p>\n<p style=\"padding-left: 60px;\">$\\sqrt{2} m g &nbsp;&lt; \\var{cmax * 1000} $ N</p>\n<p style=\"padding-left: 60px;\">$m &lt; \\frac{\\var{1000 cmax}}{\\sqrt{2} g} &nbsp;= \\var{max1}$ kg</p>\n<p style=\"padding-left: 60px;\"></p>\n<p style=\"padding-left: 30px;\"><em>For&nbsp;Tension</em></p>\n<p style=\"padding-left: 60px;\">$AB&nbsp;= T = &nbsp;m g $</p>\n<p style=\"padding-left: 60px;\">$AB &lt; &nbsp;\\var{tmax}$ &nbsp;kN</p>\n<p style=\"padding-left: 60px;\">$ m g &nbsp;&lt; \\var{1000 tmax} $ N</p>\n<p style=\"padding-left: 60px;\">$m &lt; \\frac{\\var{1000 tmax}}{g} &nbsp;= \\var{max2}$ kg</p>\n<p style=\"padding-left: 60px;\"></p>\n<p style=\"padding-left: 30px;\"><em>Solution</em></p>\n<p style=\"padding-left: 60px;\">The maximum mass will be the smaller of these two values.</p>\n<p style=\"padding-left: 60px;\">$m_{max} = \\var{answer}$&nbsp;</p>", "rulesets": {}, "extensions": ["quantities"], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"tmax": {"name": "tmax", "group": "Ungrouped variables", "definition": "cmax + random(-40..40#5 except 0)\n", "description": "", "templateType": "anything", "can_override": false}, "cmax": {"name": "cmax", "group": "Ungrouped variables", "definition": "random(50..100#5)\n", "description": "", "templateType": "anything", "can_override": false}, "max1": {"name": "max1", "group": "Ungrouped variables", "definition": "siground(cmax/sqrt(2)/9.81,3) 1000", "description": "", "templateType": "anything", "can_override": false}, "max2": {"name": "max2", "group": "Ungrouped variables", "definition": "siground(tmax/9.81,3) 1000", "description": "", "templateType": "anything", "can_override": false}, "answer": {"name": "answer", "group": "Ungrouped variables", "definition": "qty(min(max1,max2),'kg') \n", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["cmax", "tmax", "max1", "max2", "answer"], "variable_groups": [], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p style=\"padding-left: 40px;\">$m_\\text{max} = $ [[0]]</p>\n<p></p>", "gaps": [{"type": "engineering-answer", "useCustomName": true, "customName": "Max Load", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "{answer}", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}], "sortAnswers": false}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}], "resources": ["question-resources/truss150_ST2n1NE.png", "question-resources/FBD1_0gNiD65.png", "question-resources/FBD2_brSZxtl.png", "question-resources/FBD3_IU8eiX9.png", "question-resources/FBD2_Vmzm7ke.png"]}]}], "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}]}