// Numbas version: finer_feedback_settings
{"name": "Method of sections: Howe Truss", "extensions": ["geogebra", "quantities"], "custom_part_types": [{"source": {"pk": 19, "author": {"name": "William Haynes", "pk": 2530}, "edit_page": "/part_type/19/edit"}, "name": "Engineering Accuracy with units", "short_name": "engineering-answer", "description": "<p>A value with units marked right if within an adjustable % error of the correct value. &nbsp;Marked close if within a wider margin of error.</p>", "help_url": "", "input_widget": "string", "input_options": {"correctAnswer": "siground(settings['correctAnswer'],4)", "hint": {"static": true, "value": ""}, "allowEmpty": {"static": true, "value": true}}, "can_be_gap": true, "can_be_step": true, "marking_script": "mark:\nswitch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)\n\ninterpreted_answer:\nqty(student_scalar, student_units)\n\n\n\ncorrect_quantity:\nsettings[\"correctAnswer\"]\n\n\n\ncorrect_units:\nunits(correct_quantity)\n\n\nallowed_notation_styles:\n[\"plain\",\"en\"]\n\nmatch_student_number:\nmatchnumber(studentAnswer,allowed_notation_styles)\n\nstudent_scalar:\nmatch_student_number[1]\n\nstudent_units:\nreplace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")\n\ngood_units:\ntry(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n\n\nstudent_quantity:\nswitch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n\n\npercent_error:\ntry(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   \n\nright:\npercent_error <= settings['right']\n\n\nclose:\nright_sign and percent_error <= settings['close']\n\nright_sign:\nsign(student_scalar) = sign(correct_quantity)", "marking_notes": [{"name": "mark", "description": "This is the main marking note. It should award credit and provide feedback based on the student's answer.", "definition": "switch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)"}, {"name": "interpreted_answer", "description": "A value representing the student's answer to this part.", "definition": "qty(student_scalar, student_units)\n\n"}, {"name": "correct_quantity", "description": "", "definition": "settings[\"correctAnswer\"]\n\n"}, {"name": "correct_units", "description": "", "definition": "units(correct_quantity)\n"}, {"name": "allowed_notation_styles", "description": "", "definition": "[\"plain\",\"en\"]"}, {"name": "match_student_number", "description": "", "definition": "matchnumber(studentAnswer,allowed_notation_styles)"}, {"name": "student_scalar", "description": "", "definition": "match_student_number[1]"}, {"name": "student_units", "description": "<p>Modify&nbsp;the&nbsp;unit portion of&nbsp;the student's answer by</p>\n<p>1. replacing \"ohms\" with \"ohm\" &nbsp;case insensitive</p>\n<p>2.&nbsp;replacing '-'&nbsp;with ' '&nbsp;</p>\n<p>3. replacing '&deg;' with ' deg'&nbsp;</p>\n<p>to allow answers like 10 ft-lb and 30&deg;</p>", "definition": "replace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")"}, {"name": "good_units", "description": "", "definition": "try(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n"}, {"name": "student_quantity", "description": "<p>This fixes the student answer for two common errors. &nbsp;</p>\n<p>If student_units are wrong&nbsp;&nbsp;- replace with correct units</p>\n<p>If student_scalar has the wrong sign - replace with right sign</p>\n<p>If student makes both errors, only one gets fixed.</p>", "definition": "switch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n"}, {"name": "percent_error", "description": "", "definition": "try(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   "}, {"name": "right", "description": "", "definition": "percent_error <= settings['right']\n"}, {"name": "close", "description": "<p>Only marked close if the student actually has the right sign.</p>", "definition": "right_sign and percent_error <= settings['close']"}, {"name": "right_sign", "description": "", "definition": "sign(student_scalar) = sign(correct_quantity) "}], "settings": [{"name": "correctAnswer", "label": "Correct Quantity.", "help_url": "", "hint": "The correct answer given as a JME quantity.", "input_type": "code", "default_value": "", "evaluate": true}, {"name": "right", "label": "% Accuracy for right.", "help_url": "", "hint": "Question will be considered correct if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "0.2", "evaluate": true}, {"name": "close", "label": "% Accuracy for close.", "help_url": "", "hint": "Question will be considered close if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "1.0", "evaluate": true}, {"name": "C1", "label": "Close with units.", "help_url": "", "hint": "Partial Credit for close value with appropriate units. &nbsp;if correct answer is 100 N and close is &plusmn;1%,<br />99 &nbsp;N is accepted.", "input_type": "percent", "default_value": "75"}, {"name": "C2", "label": "No units or wrong sign", "help_url": "", "hint": "Partial credit for forgetting units or using wrong sign.<br />If the correct answer is 100 N, both 100 and -100 N are accepted.", "input_type": "percent", "default_value": "50"}, {"name": "C3", "label": "Close, no units.", "help_url": "", "hint": "Partial Credit for close value but forgotten units.<br />This value would be close&nbsp;if the expected units were provided. &nbsp;If the correct answer is 100 N, and close is &plusmn;1%,<br />99 is accepted.", "input_type": "percent", "default_value": "25"}], "public_availability": "always", "published": true, "extensions": ["quantities"]}], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Method of sections: Howe Truss", "tags": ["equilibrium", "Equilibrium", "Mechanics", "mechanics", "method of sections", "Statics", "statics", "truss"], "metadata": {"description": "<p>Solve for the internal force in three members of a truss.</p>", "licence": "Creative Commons Attribution-NonCommercial 4.0 International"}, "statement": "<p>The Howe truss shown supports three loads: {load_html(0,units[0])}, {load_html(1,units[0])}, and {load_html(2, units[0])}. Determine the reactions and the forces in members <span data-jme-visible=\"version=0\">$DE$, $KE$ and $KL$.</span> <span data-jme-visible=\"version=1\">$BC$, $CJ$, and $JK$.</span>&nbsp;Indicate tension or compression.</p>\n<table style=\"margin-left: auto; margin-right: auto; height: 18px;\">\n<tbody>\n<tr style=\"height: 18px;\">\n<td style=\"height: 18px; width: 734.828125px;\">{geogebra_applet('mzwuspm4',[['width',width],['height',height],['units','\"'+units[1]+'\"']]+ggb_forces)}</td>\n</tr>\n</tbody>\n</table>", "advice": "<ol>\n<li>Draw a free body diagram of the whole truss and take moments at <em>A</em> to find force <em>G</em>, and&nbsp;take moments at <em>G</em> to find <em>A</em>. &nbsp;Check your answer by ensuring that $\\Sigma F_y = 0$.</li>\n<li>Take an imaginary cut through the three members of interest to divide the truss into two sections.</li>\n<li>Draw a free body diagram of one side of the truss. &nbsp;Indicate forces at the cut members as tension (pulling away from the joint). &nbsp;Show only the loads which act on the isolated portion of the truss</li>\n<li>Apply three equations of equilibrium to find the forces in the cut members. &nbsp;Negative values will indicate compression. &nbsp;It is good idea to take moments at a point where two unknowns intersect when possible.</li>\n</ol>", "rulesets": {}, "extensions": ["geogebra", "quantities"], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"units": {"name": "units", "group": "Inputs", "definition": "random(['N','m'],['lb','ft'])", "description": "", "templateType": "anything", "can_override": false}, "BC": {"name": "BC", "group": "magnitudes", "definition": "BC_y/sin(radians(alpha))", "description": "", "templateType": "anything", "can_override": false}, "KE_y": {"name": "KE_y", "group": "magnitudes", "definition": "siground(-(f[\"F\"] + f[\"M\"]  + 2 (f[\"E\"]+ f[\"L\"]))/ 3,4)", "description": "<p>sum of moment about G/dperp</p>", "templateType": "anything", "can_override": false}, "BC_y": {"name": "BC_y", "group": "magnitudes", "definition": "(-2A + f[\"H\"] + f[\"B\"])/2", "description": "<p>sum of moments about J / dperp</p>", "templateType": "anything", "can_override": false}, "DE": {"name": "DE", "group": "magnitudes", "definition": "DE_x/cos(radians(alpha))", "description": "", "templateType": "anything", "can_override": false}, "test_ans": {"name": "test_ans", "group": "Ungrouped variables", "definition": "ans(KL,units[1])", "description": "", "templateType": "anything", "can_override": false}, "JK": {"name": "JK", "group": "magnitudes", "definition": "((2 A - f[\"B\"] - f[\"H\"]) width)/(2 height /3)", "description": "<p>sum of moments about C/dperp</p>", "templateType": "anything", "can_override": false}, "KL": {"name": "KL", "group": "magnitudes", "definition": "-((f[\"F\"]+f[\"M\"]) width - g 2 width)/(2 height /3)", "description": "", "templateType": "anything", "can_override": false}, "alpha": {"name": "alpha", "group": "Ungrouped variables", "definition": "degrees(arctan(height/(3 width)))", "description": "", "templateType": "anything", "can_override": false}, "debug": {"name": "debug", "group": "Inputs", "definition": "false", "description": "", "templateType": "anything", "can_override": false}, "version": {"name": "version", "group": "Inputs", "definition": "random(0,1)", "description": "", "templateType": "anything", "can_override": false}, "G": {"name": "G", "group": "magnitudes", "definition": "sum(map(F[k]*X[k],k,loads))/6", "description": "<p>sum of moments about A / dperp&nbsp;</p>", "templateType": "anything", "can_override": false}, "F": {"name": "F", "group": "Inputs", "definition": "dict(map([n,if(n in loads,random(100..500#25),0)],n,names))", "description": "", "templateType": "anything", "can_override": false}, "beta": {"name": "beta", "group": "Ungrouped variables", "definition": "degrees(arctan(2 height/3/width))", "description": "", "templateType": "anything", "can_override": false}, "DE_x": {"name": "DE_x", "group": "magnitudes", "definition": "siground((2(f[\"F\"]+f[\"M\"]) +  (f[\"E\"]+ f[\"L\"]) - 3 G ) \n  width/ height,4)", "description": "<p>sum of moments about K</p>", "templateType": "anything", "can_override": false}, "names": {"name": "names", "group": "Inputs", "definition": "[\"A\",\"B\",\"H\",\"C\",\"J\",\"D\",\"K\",\"E\",\"L\",\"F\",\"M\",\"G\"]", "description": "", "templateType": "anything", "can_override": false}, "loads": {"name": "loads", "group": "Inputs", "definition": "sort(shuffle(names[1..10])[0..3])", "description": "<p>only apply loads at middle 10 (out of 12) points.</p>", "templateType": "anything", "can_override": false}, "height": {"name": "height", "group": "Inputs", "definition": "3 random(1/2,2/3,3/4,2/5,3/5) width\n", "description": "", "templateType": "anything", "can_override": false}, "A": {"name": "A", "group": "magnitudes", "definition": "sum(map(F[k],k,loads))-G", "description": "", "templateType": "anything", "can_override": false}, "ggb_forces": {"name": "ggb_forces", "group": "Inputs", "definition": "map([lower(n),F[n]],n,loads)", "description": "", "templateType": "anything", "can_override": false}, "X": {"name": "X", "group": "Inputs", "definition": "dict(map([names[n],[0,1,1,2,2,3,3,4,4,5,5,6][n]],n,0..11))", "description": "<p>horizontal location of point. &nbsp;multiply by width for actual position.</p>", "templateType": "anything", "can_override": false}, "KE": {"name": "KE", "group": "magnitudes", "definition": "KE_y/sin(radians(beta))", "description": "", "templateType": "anything", "can_override": false}, "CJ": {"name": "CJ", "group": "magnitudes", "definition": "(2 f[\"J\"]+ F[\"B\"] + F[\"H\"])/2", "description": "<p>sum of moments about A/dperp</p>", "templateType": "anything", "can_override": false}, "width": {"name": "width", "group": "Inputs", "definition": "random(3,4,5,6)", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["alpha", "beta", "test_ans"], "variable_groups": [{"name": "Inputs", "variables": ["width", "units", "height", "ggb_forces", "names", "loads", "F", "X", "version", "debug"]}, {"name": "tests", "variables": []}, {"name": "magnitudes", "variables": ["A", "G", "KL", "DE_x", "DE", "KE_y", "CJ", "BC_y", "BC", "JK", "KE"]}, {"name": "quantities", "variables": []}], "functions": {"display": {"parameters": [["q", "quantity"]], "type": "string", "language": "jme", "definition": "string(abs(q))"}, "ans": {"parameters": [["f", "number"], ["u", "string"]], "type": "string", "language": "jme", "definition": "string(siground(qty(abs(f),u),4)) + if(f<0,' (C)', ' (T)')"}, "load_html": {"parameters": [["k", "number"], ["u", "string"]], "type": "html", "language": "jme", "definition": "\"$\" + loads[k] + \"= $ \" + qty(f[loads[k]],u)"}}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p style=\"padding-left: 40px;\">$\\phantom{A}A$ = [[0]]</p>\n<p style=\"padding-left: 40px;\">$\\phantom{A}G$ = [[1]] <span data-jme-visible=\"debug\"><em>A</em>&nbsp;= {siground(qty(A,units[0]),4)} &nbsp;<em>G</em> = {siground(qty(G,units[0]),4)} &nbsp; &nbsp;</span></p>\n<p style=\"padding-left: 40px;\"><span data-jme-visible=\"version=0\">$DE$</span>&nbsp;<span data-jme-visible=\"version=1\">$BC$</span> =&nbsp;[[2]]&nbsp;[[3]]&nbsp; &nbsp;<span data-jme-visible=\"debug\"><em>DE</em> = {ans(DE, units[0])}&nbsp;$BC$ = {ans(BC, units[0])}</span><span data-jme-visible=\"debug\"></span></p>\n<p style=\"padding-left: 40px;\"><em><span data-jme-visible=\"version=0\">$KE$</span>&nbsp;<span data-jme-visible=\"version=1\">$CJ$</span></em> =&nbsp;[[4]]&nbsp;[[5]]&nbsp; &nbsp;<span data-jme-visible=\"debug\"><em>KE</em> = &nbsp;{ans(KE,units[0])} &nbsp;<em>CJ</em> &nbsp;= {ans(CJ,units[0])}</span></p>\n<p style=\"padding-left: 40px;\"><i><span data-jme-visible=\"version=0\">$KL$</span>&nbsp;<span data-jme-visible=\"version=1\">$JK$</span></i> =&nbsp;[[6]]&nbsp;[[7]]&nbsp;<span data-jme-visible=\"debug\"><em>KL</em> = {ans(KL,units[0])}&nbsp;<em>JK</em> = {ans(JK,units[0])}</span></p>", "gaps": [{"type": "engineering-answer", "useCustomName": false, "customName": "", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(A,units[0])", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": false, "customName": "", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(G,units[0])", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": false, "customName": "", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "abs(qty(if(version=0,DE,BC),units[0]))", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "choices": ["Tension", "Compression", "Neither"], "matrix": ["if(if(version=0,DE,BC)>0,1,0)", "if(if(version=0,DE,BC)<0,1,0)", "if(if(version=0,DE,BC)=0,1,0)"], "distractors": ["", "", ""]}, {"type": "engineering-answer", "useCustomName": false, "customName": "", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "abs(qty(if(version=0,KE,CJ),units[0]))", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "choices": ["Tension", "Compression", "Neither"], "matrix": ["if(if(version=0,KE,CJ)>0,1,0)", "if(if(version=0,KE,CJ)<0,1,0)", "if(if(version=0,KE,CJ)=0,1,0)"], "distractors": ["", "", ""]}, {"type": "engineering-answer", "useCustomName": false, "customName": "", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "abs(qty(if(version=0,KL,JK),units[0]))", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "choices": ["Tension", "Compression", "Neither"], "matrix": ["if(if(version=0,KL,JK)>0,1,0)", "if(if(version=0,KL,JK)<0,1,0)", "if(if(version=0,KL,JK)=0,1,0)"], "distractors": ["", "", ""]}], "sortAnswers": false}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}], "resources": []}]}], "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}]}