// Numbas version: finer_feedback_settings
{"name": "Vector addition by summing scalar components", "extensions": ["geogebra", "quantities"], "custom_part_types": [{"source": {"pk": 19, "author": {"name": "William Haynes", "pk": 2530}, "edit_page": "/part_type/19/edit"}, "name": "Engineering Accuracy with units", "short_name": "engineering-answer", "description": "<p>A value with units marked right if within an adjustable % error of the correct value. &nbsp;Marked close if within a wider margin of error.</p>", "help_url": "", "input_widget": "string", "input_options": {"correctAnswer": "siground(settings['correctAnswer'],4)", "hint": {"static": true, "value": ""}, "allowEmpty": {"static": true, "value": true}}, "can_be_gap": true, "can_be_step": true, "marking_script": "mark:\nswitch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)\n\ninterpreted_answer:\nqty(student_scalar, student_units)\n\n\n\ncorrect_quantity:\nsettings[\"correctAnswer\"]\n\n\n\ncorrect_units:\nunits(correct_quantity)\n\n\nallowed_notation_styles:\n[\"plain\",\"en\"]\n\nmatch_student_number:\nmatchnumber(studentAnswer,allowed_notation_styles)\n\nstudent_scalar:\nmatch_student_number[1]\n\nstudent_units:\nreplace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")\n\ngood_units:\ntry(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n\n\nstudent_quantity:\nswitch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n\n\npercent_error:\ntry(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   \n\nright:\npercent_error <= settings['right']\n\n\nclose:\nright_sign and percent_error <= settings['close']\n\nright_sign:\nsign(student_scalar) = sign(correct_quantity)", "marking_notes": [{"name": "mark", "description": "This is the main marking note. It should award credit and provide feedback based on the student's answer.", "definition": "switch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)"}, {"name": "interpreted_answer", "description": "A value representing the student's answer to this part.", "definition": "qty(student_scalar, student_units)\n\n"}, {"name": "correct_quantity", "description": "", "definition": "settings[\"correctAnswer\"]\n\n"}, {"name": "correct_units", "description": "", "definition": "units(correct_quantity)\n"}, {"name": "allowed_notation_styles", "description": "", "definition": "[\"plain\",\"en\"]"}, {"name": "match_student_number", "description": "", "definition": "matchnumber(studentAnswer,allowed_notation_styles)"}, {"name": "student_scalar", "description": "", "definition": "match_student_number[1]"}, {"name": "student_units", "description": "<p>Modify&nbsp;the&nbsp;unit portion of&nbsp;the student's answer by</p>\n<p>1. replacing \"ohms\" with \"ohm\" &nbsp;case insensitive</p>\n<p>2.&nbsp;replacing '-'&nbsp;with ' '&nbsp;</p>\n<p>3. replacing '&deg;' with ' deg'&nbsp;</p>\n<p>to allow answers like 10 ft-lb and 30&deg;</p>", "definition": "replace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")"}, {"name": "good_units", "description": "", "definition": "try(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n"}, {"name": "student_quantity", "description": "<p>This fixes the student answer for two common errors. &nbsp;</p>\n<p>If student_units are wrong&nbsp;&nbsp;- replace with correct units</p>\n<p>If student_scalar has the wrong sign - replace with right sign</p>\n<p>If student makes both errors, only one gets fixed.</p>", "definition": "switch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n"}, {"name": "percent_error", "description": "", "definition": "try(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   "}, {"name": "right", "description": "", "definition": "percent_error <= settings['right']\n"}, {"name": "close", "description": "<p>Only marked close if the student actually has the right sign.</p>", "definition": "right_sign and percent_error <= settings['close']"}, {"name": "right_sign", "description": "", "definition": "sign(student_scalar) = sign(correct_quantity) "}], "settings": [{"name": "correctAnswer", "label": "Correct Quantity.", "help_url": "", "hint": "The correct answer given as a JME quantity.", "input_type": "code", "default_value": "", "evaluate": true}, {"name": "right", "label": "% Accuracy for right.", "help_url": "", "hint": "Question will be considered correct if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "0.2", "evaluate": true}, {"name": "close", "label": "% Accuracy for close.", "help_url": "", "hint": "Question will be considered close if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "1.0", "evaluate": true}, {"name": "C1", "label": "Close with units.", "help_url": "", "hint": "Partial Credit for close value with appropriate units. &nbsp;if correct answer is 100 N and close is &plusmn;1%,<br />99 &nbsp;N is accepted.", "input_type": "percent", "default_value": "75"}, {"name": "C2", "label": "No units or wrong sign", "help_url": "", "hint": "Partial credit for forgetting units or using wrong sign.<br />If the correct answer is 100 N, both 100 and -100 N are accepted.", "input_type": "percent", "default_value": "50"}, {"name": "C3", "label": "Close, no units.", "help_url": "", "hint": "Partial Credit for close value but forgotten units.<br />This value would be close&nbsp;if the expected units were provided. &nbsp;If the correct answer is 100 N, and close is &plusmn;1%,<br />99 is accepted.", "input_type": "percent", "default_value": "25"}], "public_availability": "always", "published": true, "extensions": ["quantities"]}], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Vector addition by summing scalar components", "tags": ["Mechanics", "mechanics", "Statics", "statics", "Vector addition", "vector addition", "Vector Addition"], "metadata": {"description": "<p>Add three vectors by determining their scalar components, summing them and then resolving the rectangular components to find the magnitude and direction of the resultant</p>", "licence": "Creative Commons Attribution-NonCommercial 4.0 International"}, "statement": "<p><strong></strong>Three forces act on the bracket at point $A$ as shown: &nbsp;$\\color{#006400}{ \\mathbf{A} = \\var{FA}\\ &nbsp;\\var{units[1]}}$, $\\color{#0033cc}{\\mathbf{B} = \\var{FB}\\ \\var{units[1]}}$ and, $\\color{#cc0000}{\\mathbf{C} = \\var{FC}\\ \\var{units[1]}}$. Find the magnitude and direction of the resultant force&nbsp;$\\mathbf{R}$ by summing scalar components.</p>\n<table style=\"margin-left: auto; margin-right: auto;\">\n<tbody>\n<tr>\n<td>{geogebra_applet('btgqdpqx', [['A',A],['B',B],['C',C],['&alpha;', radians(mod(alpha,360))]])}</td>\n</tr>\n</tbody>\n</table>\n<p></p>\n<p><span data-jme-visible=\"debug\">Force A: {fa} at {alpha} = {forceA}&nbsp;</span></p>\n<p><span data-jme-visible=\"debug\">Force B: {fb} at {beta} = {forceB}</span></p>\n<p><span data-jme-visible=\"debug\">Force C: {fc} at {gamma} = {forceC}</span></p>\n<p><span data-jme-visible=\"debug\">Resultant: {R} at {rho} &nbsp;= {forceR}</span></p>", "advice": "<p>Vector Addition: &nbsp;</p>\n<p>$R_x = \\Sigma F_x = &nbsp;(\\var{siground(ForceA[0],4)}) &nbsp;+ (\\var{siground(ForceB[0],4)}) +(&nbsp;\\var{siground(ForceC[0],4)}) = \\var{siground(ForceR[0],4)}$&nbsp;{units[1]}</p>\n<p>$R_y = \\Sigma F_y = &nbsp;(\\var{siground(ForceA[1],4)}) &nbsp;+ (\\var{siground(ForceB[1],4)}) +(&nbsp;\\var{siground(ForceC[1],4)}) = \\var{siground(ForceR[1],4)}$&nbsp;{units[1]}</p>\n<p></p>\n<p>$ R=\\sqrt{R_x^2 + R_y^2} = \\sqrt{(\\var{siground(ForceR[0],4)})^2 +(\\var{siground(ForceR[0],4)})^2} = \\var{siground(abs(ForceR),4)}$&nbsp;{units[1]}</p>\n<p></p>\n<p>$\\theta = \\tan^{-1}\\left(\\left|\\frac{R_y}{R_x}\\right| \\right) = \\tan^{-1}\\left(\\frac{\\var{abs(siground(ForceR[1],4))}}{\\var{abs(siground(ForceR[0],4))}} \\right) &nbsp;= \\var{siground(degrees(arctan(abs(forceR[1]/forceR[0]))),4)}&deg; \\text{ from the } x \\text{-axis in quadrant } \\var{quadrant(forceR)} \\text{ i.e. } &nbsp;\\var{siground(rho,4)}&deg; $ from the positive x-axis.</p>", "rulesets": {}, "extensions": ["geogebra", "quantities"], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"FC": {"name": "FC", "group": "Inputs", "definition": "random(10..100#5)", "description": "<p>Magnitude of force C</p>", "templateType": "anything", "can_override": false}, "A": {"name": "A", "group": "Inputs", "definition": "vector(0,0)", "description": "<p>Position of point A</p>", "templateType": "anything", "can_override": false}, "C": {"name": "C", "group": "Inputs", "definition": "vector( random(1,-1) random(6..13),random(1,-1) random(6..9))", "description": "<p>Position of point C</p>", "templateType": "anything", "can_override": false}, "B": {"name": "B", "group": "Inputs", "definition": "vector( random(1,-1) random(5..13),random(1,-1) random(5..9))", "description": "<p>Position of point B</p>", "templateType": "anything", "can_override": false}, "FA": {"name": "FA", "group": "Inputs", "definition": "random(10..100#5)", "description": "<p>Magnitude of force A</p>", "templateType": "anything", "can_override": false}, "debug": {"name": "debug", "group": "Inputs", "definition": "false", "description": "", "templateType": "anything", "can_override": false}, "ForceR": {"name": "ForceR", "group": "Outputs", "definition": "ForceA+ForceB+ForceC", "description": "<p>Resultant as a vector</p>", "templateType": "anything", "can_override": false}, "R": {"name": "R", "group": "Outputs", "definition": "sqrt(ForceR[0]^2+forceR[1]^2)", "description": "<p>Magnitude of resultant</p>", "templateType": "anything", "can_override": false}, "ForceC": {"name": "ForceC", "group": "Outputs", "definition": "precround(FC *( vector(cos(radians(gamma)),sin(radians(gamma)))),4)", "description": "<p>Force C as a vector</p>", "templateType": "anything", "can_override": false}, "FB": {"name": "FB", "group": "Inputs", "definition": "random(10..100#10)", "description": "<p>Magnitude of Force B</p>", "templateType": "anything", "can_override": false}, "rho": {"name": "rho", "group": "Outputs", "definition": "degrees(direction(ForceR))", "description": "<p style=\"text-align: left;\">direction of resultant</p>", "templateType": "anything", "can_override": false}, "beta": {"name": "beta", "group": "Outputs", "definition": "degrees(direction(vector(B[0]-A[0],B[1]-A[1])))", "description": "<p>Direction of force B</p>", "templateType": "anything", "can_override": false}, "ForceA": {"name": "ForceA", "group": "Outputs", "definition": "precround(FA *( vector(cos(radians(alpha)),sin(radians(alpha)))),4)", "description": "<p>force A&nbsp;as a vector</p>", "templateType": "anything", "can_override": false}, "angle_from_ref": {"name": "angle_from_ref", "group": "Outputs", "definition": "let(ang,rho,\n[if(ang>180,ang-360,ang),\nif(ang>270,ang-450,if(ang < -90,ang+270,ang-90)),\nif(ang>0,ang-180,ang+180),\nif(ang>90,ang-270,90+ang)])", "description": "", "templateType": "anything", "can_override": false}, "gamma": {"name": "gamma", "group": "Outputs", "definition": "degrees(direction(vector(C[0]-A[0],C[1]-A[1])))", "description": "<p>Direction of force C in degrees</p>", "templateType": "anything", "can_override": false}, "ForceB": {"name": "ForceB", "group": "Outputs", "definition": "precround(FB *( vector(cos(radians(beta)),sin(radians(beta)))),4)", "description": "<p>Force B as a vector</p>", "templateType": "anything", "can_override": false}, "units": {"name": "units", "group": "Inputs", "definition": "random(['ft','lb'],['in','lb'],['cm','N'])", "description": "", "templateType": "anything", "can_override": false}, "alpha": {"name": "alpha", "group": "Inputs", "definition": "random(0..360#5)", "description": "<p>direction of force A</p>", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "abs(mod(beta,180) - mod(gamma,180)) > 30 and\nabs(mod(beta,180) - mod(alpha,180)) > 30 and\nabs(mod(alpha,180) - mod(gamma,180)) > 30 and\nquadrant(B) <> quadrant(C)\n", "maxRuns": 100}, "ungrouped_variables": [], "variable_groups": [{"name": "Inputs", "variables": ["A", "B", "C", "alpha", "FA", "FB", "FC", "units", "debug"]}, {"name": "Outputs", "variables": ["ForceA", "beta", "gamma", "ForceB", "ForceC", "ForceR", "rho", "R", "angle_from_ref"]}], "functions": {"direction": {"parameters": [["v", "vector"]], "type": "number", "language": "javascript", "definition": "return Math.atan2(v[1],v[0])"}, "quadrant": {"parameters": [["pt", "vector"]], "type": "anything", "language": "jme", "definition": "let(a,mod(degrees(atan2(pt[1],pt[0])),360),if(a< 90, 1, if(a < 180, 2, (if(a < 270, 3,4)))))"}}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": true, "customName": "Components of Components", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Find the scalar components of the component forces <strong>A</strong>, <strong>B</strong>, and <strong>C</strong>.&nbsp;</p>\n<p style=\"padding-left: 40px;\">$A_x =$&nbsp;[[0]]&nbsp;</p>\n<p style=\"padding-left: 40px;\">$A_y =$&nbsp;[[1]]&nbsp;</p>\n<p style=\"padding-left: 40px;\">$B_x =$&nbsp;[[2]]&nbsp;</p>\n<p style=\"padding-left: 40px;\">$B_y =$ [[3]]&nbsp;</p>\n<p style=\"padding-left: 40px;\">$C_x =$&nbsp;[[4]]</p>\n<p style=\"padding-left: 40px;\">&nbsp;$C_y =$&nbsp;[[5]]</p>", "gaps": [{"type": "engineering-answer", "useCustomName": true, "customName": "Ax", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(forceA[0],units[1])", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "Ay", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(forceA[1],units[1])", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "Bx", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(forceB[0],units[1])", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "By", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(forceB[1],units[1])", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "Cx", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(forceC[0],units[1])", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "Cy", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(forceC[1],units[1])", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}], "sortAnswers": false}, {"type": "gapfill", "useCustomName": true, "customName": "Components of Resultant", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Sum the scalar components to get&nbsp;the scalar components of the resultant <strong>R</strong>.</p>\n<p style=\"padding-left: 40px;\">$R_x = \\Sigma F_x &nbsp;= A_x + B_x + C_x =$&nbsp;[[0]]&nbsp;</p>\n<p style=\"padding-left: 40px;\">$R_y = \\Sigma F_y &nbsp;= A_y + B_y + C_y =$ [[1]]&nbsp;&nbsp;</p>", "gaps": [{"type": "engineering-answer", "useCustomName": true, "customName": "Rx", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(ForceR[0],units[1])", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "Ry", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(ForceR[1],units[1])", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}], "sortAnswers": false}, {"type": "gapfill", "useCustomName": true, "customName": "Magnitude", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Resolve&nbsp;$R_x$ and $R_y$ to find the magnitude and direction of <strong>R</strong>.</p>\n<p style=\"padding-left: 40px;\">$R = \\sqrt{{R_x}^2 +{R_y}^2}$ =&nbsp;[[0]]&nbsp;</p>\n<p></p>", "gaps": [{"type": "engineering-answer", "useCustomName": true, "customName": "R", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(R,units[1])", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}], "sortAnswers": false}, {"type": "gapfill", "useCustomName": true, "customName": "Direction", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "mark:\n  if(iscorrect, set_credit(1,\"Correct. \"), set_credit(0, correct_angle))\n\ntarget_angle (the correct angle):\n scalar(gaps[0][\"settings\"][\"correctAnswer\"])\n\nreference (student choice of reference axis: 0,90,180,-90 ):\n   answers[2]\n\nanswers0:\n    answers[2]\n\nsign (student choice of directions: 1, -1):\n   answers[1]\n\nstudent_angle (student input as as standard angle):\n  reference + sign * scalar(answers[0])\n\ncorrect_angle:\n  \"Correct answer: \" + precround((target_angle-reference)/sign,2) + \"\u00b0 for your selections.\"\n\nisCorrect:\n  resultsequal(cos(radians(student_angle)),cos(radians(target_angle)),\"absdiff\", radians(0.1)) and\n  resultsequal(sin(radians(student_angle)),sin(radians(target_angle)),\"absdiff\", radians(0.1))\n", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p style=\"padding-left: 40px;\">&nbsp;$\\theta = \\tan^{-1}\\left(\\left|\\dfrac{R_y}{R_x}\\right| \\right) =$ [[0]]&nbsp;</p>\n<p style=\"padding-left: 40px;\">measured&nbsp;[[1]] from the&nbsp;[[2]].</p>", "gaps": [{"type": "engineering-answer", "useCustomName": true, "customName": "angle", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "\n", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(rho,\"deg\")", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": true, "customName": "dir", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "interpreted_answer:\n  [1,-1][indices( studentAnswer,[true])[0]]", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": false, "choices": ["CCW", "CW"], "matrix": [0, 0], "distractors": ["", ""]}, {"type": "1_n_2", "useCustomName": true, "customName": "ref", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "interpreted_answer:\n  [0,90,180,-90][indices(studentAnswer,[true])[0]]", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": false, "choices": ["Positive x axis", "Positive y axis", "Negative x axis", "Negative y axis"], "matrix": [0, 0, 0, 0], "distractors": ["", "", "", ""]}], "sortAnswers": false}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}], "resources": []}]}], "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}]}