// Numbas version: finer_feedback_settings
{"name": "Moment of inertia of composite shape - triangle and rectangles", "extensions": ["geogebra", "quantities"], "custom_part_types": [{"source": {"pk": 19, "author": {"name": "William Haynes", "pk": 2530}, "edit_page": "/part_type/19/edit"}, "name": "Engineering Accuracy with units", "short_name": "engineering-answer", "description": "<p>A value with units marked right if within an adjustable % error of the correct value. &nbsp;Marked close if within a wider margin of error.</p>", "help_url": "", "input_widget": "string", "input_options": {"correctAnswer": "siground(settings['correctAnswer'],4)", "hint": {"static": true, "value": ""}, "allowEmpty": {"static": true, "value": true}}, "can_be_gap": true, "can_be_step": true, "marking_script": "mark:\nswitch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)\n\ninterpreted_answer:\nqty(student_scalar, student_units)\n\n\n\ncorrect_quantity:\nsettings[\"correctAnswer\"]\n\n\n\ncorrect_units:\nunits(correct_quantity)\n\n\nallowed_notation_styles:\n[\"plain\",\"en\"]\n\nmatch_student_number:\nmatchnumber(studentAnswer,allowed_notation_styles)\n\nstudent_scalar:\nmatch_student_number[1]\n\nstudent_units:\nreplace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")\n\ngood_units:\ntry(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n\n\nstudent_quantity:\nswitch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n\n\npercent_error:\ntry(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   \n\nright:\npercent_error <= settings['right']\n\n\nclose:\nright_sign and percent_error <= settings['close']\n\nright_sign:\nsign(student_scalar) = sign(correct_quantity)", "marking_notes": [{"name": "mark", "description": "This is the main marking note. It should award credit and provide feedback based on the student's answer.", "definition": "switch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)"}, {"name": "interpreted_answer", "description": "A value representing the student's answer to this part.", "definition": "qty(student_scalar, student_units)\n\n"}, {"name": "correct_quantity", "description": "", "definition": "settings[\"correctAnswer\"]\n\n"}, {"name": "correct_units", "description": "", "definition": "units(correct_quantity)\n"}, {"name": "allowed_notation_styles", "description": "", "definition": "[\"plain\",\"en\"]"}, {"name": "match_student_number", "description": "", "definition": "matchnumber(studentAnswer,allowed_notation_styles)"}, {"name": "student_scalar", "description": "", "definition": "match_student_number[1]"}, {"name": "student_units", "description": "<p>Modify&nbsp;the&nbsp;unit portion of&nbsp;the student's answer by</p>\n<p>1. replacing \"ohms\" with \"ohm\" &nbsp;case insensitive</p>\n<p>2.&nbsp;replacing '-'&nbsp;with ' '&nbsp;</p>\n<p>3. replacing '&deg;' with ' deg'&nbsp;</p>\n<p>to allow answers like 10 ft-lb and 30&deg;</p>", "definition": "replace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")"}, {"name": "good_units", "description": "", "definition": "try(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n"}, {"name": "student_quantity", "description": "<p>This fixes the student answer for two common errors. &nbsp;</p>\n<p>If student_units are wrong&nbsp;&nbsp;- replace with correct units</p>\n<p>If student_scalar has the wrong sign - replace with right sign</p>\n<p>If student makes both errors, only one gets fixed.</p>", "definition": "switch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n"}, {"name": "percent_error", "description": "", "definition": "try(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   "}, {"name": "right", "description": "", "definition": "percent_error <= settings['right']\n"}, {"name": "close", "description": "<p>Only marked close if the student actually has the right sign.</p>", "definition": "right_sign and percent_error <= settings['close']"}, {"name": "right_sign", "description": "", "definition": "sign(student_scalar) = sign(correct_quantity) "}], "settings": [{"name": "correctAnswer", "label": "Correct Quantity.", "help_url": "", "hint": "The correct answer given as a JME quantity.", "input_type": "code", "default_value": "", "evaluate": true}, {"name": "right", "label": "% Accuracy for right.", "help_url": "", "hint": "Question will be considered correct if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "0.2", "evaluate": true}, {"name": "close", "label": "% Accuracy for close.", "help_url": "", "hint": "Question will be considered close if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "1.0", "evaluate": true}, {"name": "C1", "label": "Close with units.", "help_url": "", "hint": "Partial Credit for close value with appropriate units. &nbsp;if correct answer is 100 N and close is &plusmn;1%,<br />99 &nbsp;N is accepted.", "input_type": "percent", "default_value": "75"}, {"name": "C2", "label": "No units or wrong sign", "help_url": "", "hint": "Partial credit for forgetting units or using wrong sign.<br />If the correct answer is 100 N, both 100 and -100 N are accepted.", "input_type": "percent", "default_value": "50"}, {"name": "C3", "label": "Close, no units.", "help_url": "", "hint": "Partial Credit for close value but forgotten units.<br />This value would be close&nbsp;if the expected units were provided. &nbsp;If the correct answer is 100 N, and close is &plusmn;1%,<br />99 is accepted.", "input_type": "percent", "default_value": "25"}], "public_availability": "always", "published": true, "extensions": ["quantities"]}], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Moment of inertia of composite shape - triangle and rectangles", "tags": ["Area moment of inertia", "area moment of inertia", "Mechanics", "mechanics", "moment of inertia", "radius of gyration", "Statics", "statics"], "metadata": {"description": "<p>Find moment of inertia of a composite shape consisting of a rectangle and two triangles with respect to the <em>x</em>-axis. Shapes rest on the x-axis so the parallel axis theorem is not required.</p>", "licence": "Creative Commons Attribution-NonCommercial 4.0 International"}, "statement": "<p>Determine the total area of the composite shape, and its moment of inertia and radius of gyration with respect to the $x$-axis. &nbsp;Grid units are [{units}].</p>\n<table style=\"margin-left: auto; margin-right: auto; width: 433px;\" height=\"44\">\n<tbody>\n<tr>\n<td style=\"width: 416px;\">{geogebra_applet('wsef7939',[['A',A],['B',B],['C',C]])}</td>\n</tr>\n</tbody>\n</table>", "advice": "<p><strong></strong>Consider the shape to be made of three parts.</p>\n<ol>\n<li>$\\qquad&nbsp;\\var{b1} \\times \\var{h1}$ rectangle.</li>\n<li>$\\qquad \\var{b2} \\times \\var{h2}$ triangle.</li>\n<li>$\\qquad \\var{b3} \\times \\var{h3}$&nbsp;<strong>{if(removed,'negative','')}&nbsp;</strong>triangle.</li>\n</ol>\n<p>Area:</p>\n<p style=\"padding-left: 40px;\">$A_1 = b_1 h_1 = \\var{A1}$ <strong>&nbsp;</strong></p>\n<p style=\"padding-left: 40px;\">$A_2 = \\frac{1}{2} b_2 h_2= \\var{A2}$<strong>&nbsp;</strong></p>\n<p style=\"padding-left: 40px;\">$A_3 = \\frac{1}{2}&nbsp;b_3 h_3= \\var{A3}$ &nbsp;<strong>{if(removed,'Removed','')}</strong></p>\n<p style=\"padding-left: 40px;\">$A = \\Sigma A_i &nbsp;= \\var{A_t}$</p>\n<p>Moment of inertia with respect to the x-axis<strong>:</strong></p>\n<p style=\"padding-left: 40px;\">$\\begin{align}I_x &amp;= \\Sigma I_x \\\\&nbsp;&amp;= \\frac{b_1 h_1^3}{3} + \\frac{b_2 h_2^3&nbsp;}{12} +&nbsp;\\var{if(removed,'&ndash;','')} \\frac{b_3 h_3^3}{12}\\\\&amp;= (\\var{siground(I1x,4)}) + &nbsp;(\\var{siground(I2x,4)}) + (\\var{if(removed,'&ndash;','')}\\var{ siground(I3x,4)})\\\\&amp;= \\var{siground(I_t,4)}\\end{align}$</p>\n<p>Radius of gyration with respect to the x-axis:</p>\n<p style=\"padding-left: 40px;\">$k_x = \\sqrt{\\dfrac{I_x}{A}} &nbsp;= \\var{siground(kx,4)}$</p>", "rulesets": {}, "extensions": ["geogebra", "quantities"], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"A_t": {"name": "A_t", "group": "derived geometry", "definition": "A1+ A2 + if(removed,-1,1) A3", "description": "", "templateType": "anything", "can_override": false}, "A2": {"name": "A2", "group": "derived geometry", "definition": "B2 H2 /2", "description": "", "templateType": "anything", "can_override": false}, "h2": {"name": "h2", "group": "derived geometry", "definition": "qty(B[1],units)", "description": "<p>part 2 left triangle</p>", "templateType": "anything", "can_override": false}, "expression": {"name": "expression", "group": "derived geometry", "definition": "random([\n  \"\\\\frac\\{1\\}\\{2\\}\",\n  \"\\\\sin(x)\",\n  \"2 + \\\\hat\\{x\\}\"\n])", "description": "", "templateType": "anything", "can_override": false}, "b2": {"name": "b2", "group": "derived geometry", "definition": "qty(abs(C[0]),units)\n", "description": "<p>base of left triangle &nbsp;x(c) is always &lt; 0</p>", "templateType": "anything", "can_override": false}, "A3": {"name": "A3", "group": "derived geometry", "definition": "B3 h3/2", "description": "", "templateType": "anything", "can_override": false}, "units": {"name": "units", "group": "inputs", "definition": "random('in','ft','mm','cm','m')", "description": "", "templateType": "anything", "can_override": false}, "I3x": {"name": "I3x", "group": "answers", "definition": "(b3 * h3 * h3 * h3)/12", "description": "", "templateType": "anything", "can_override": false}, "C": {"name": "C", "group": "inputs", "definition": "vector(random(-9..-1),0)", "description": "<p>this point determines the triangle to the left of the y axis</p>", "templateType": "anything", "can_override": false}, "debug": {"name": "debug", "group": "inputs", "definition": "false", "description": "", "templateType": "anything", "can_override": false}, "A": {"name": "A", "group": "inputs", "definition": "vector(random(C[0]..10 except B[0] ),random(1..B[1]))", "description": "<p>This points sets the coordinates of the second triangle, at (x,0) and (x(b),y)</p>", "templateType": "anything", "can_override": false}, "I1x": {"name": "I1x", "group": "answers", "definition": "b1 * h1 * h1 * h1/3", "description": "", "templateType": "anything", "can_override": false}, "b1": {"name": "b1", "group": "derived geometry", "definition": "qty(B[0],units)", "description": "<p>base of rectangle 1</p>", "templateType": "anything", "can_override": false}, "I_t": {"name": "I_t", "group": "answers", "definition": "I1x + I2x + if(removed,-1,1) I3x", "description": "<p>total moment of inertia wrt x axis</p>", "templateType": "anything", "can_override": false}, "B": {"name": "B", "group": "inputs", "definition": "vector(random(1..8),random(1..9))", "description": "<p>top right corner of the rectangle</p>", "templateType": "anything", "can_override": false}, "kx": {"name": "kx", "group": "answers", "definition": "qty(sqrt(scalar(I_t/A_t)),units)", "description": "", "templateType": "anything", "can_override": false}, "h3": {"name": "h3", "group": "derived geometry", "definition": "qty(A[1],units)", "description": "", "templateType": "anything", "can_override": false}, "h1": {"name": "h1", "group": "derived geometry", "definition": "qty(B[1],units)", "description": "<p>height of rectangle 1</p>\n<p></p>", "templateType": "anything", "can_override": false}, "axis": {"name": "axis", "group": "inputs", "definition": "random('x','y')", "description": "", "templateType": "anything", "can_override": false}, "I2x": {"name": "I2x", "group": "answers", "definition": "(b2 * h2 * h2 * h2)/12\n", "description": "", "templateType": "anything", "can_override": false}, "A1": {"name": "A1", "group": "derived geometry", "definition": "B1 H1", "description": "<p>area of rectangle 1</p>", "templateType": "anything", "can_override": false}, "removed": {"name": "removed", "group": "derived geometry", "definition": "if(A[0]<B[0],true,false)", "description": "", "templateType": "anything", "can_override": false}, "b3": {"name": "b3", "group": "derived geometry", "definition": "qty(abs(B[0]-A[0]),units)", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": [], "variable_groups": [{"name": "inputs", "variables": ["A", "B", "C", "units", "debug", "axis"]}, {"name": "derived geometry", "variables": ["h1", "b2", "h2", "b3", "h3", "b1", "removed", "A1", "A2", "A3", "A_t", "expression"]}, {"name": "answers", "variables": ["I1x", "I2x", "I3x", "I_t", "kx"]}], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p style=\"padding-left: 40px;\">$A$ =&nbsp;[[0]]&nbsp;&nbsp; <span data-jme-visible=\"debug\">{A_t}</span></p>\n<p style=\"padding-left: 40px;\">$I_x =$&nbsp;[[1]]&nbsp;<span data-jme-visible=\"debug\">{siground(I_t,4)}</span></p>\n<p style=\"padding-left: 40px;\">$k_x =$&nbsp;[[2]]&nbsp; <span data-jme-visible=\"debug\">{siground(kx,4)}</span></p>", "gaps": [{"type": "engineering-answer", "useCustomName": true, "customName": "$A$", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "abs(A_t)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "$I_x$", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "I_t", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "$k_x$", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "kx", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}], "sortAnswers": false}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}], "resources": []}]}], "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}]}