// Numbas version: finer_feedback_settings
{"name": "Parallel Axis Theorem: semi and quarter circles", "extensions": ["geogebra", "quantities"], "custom_part_types": [{"source": {"pk": 19, "author": {"name": "William Haynes", "pk": 2530}, "edit_page": "/part_type/19/edit"}, "name": "Engineering Accuracy with units", "short_name": "engineering-answer", "description": "<p>A value with units marked right if within an adjustable % error of the correct value. &nbsp;Marked close if within a wider margin of error.</p>", "help_url": "", "input_widget": "string", "input_options": {"correctAnswer": "siground(settings['correctAnswer'],4)", "hint": {"static": true, "value": ""}, "allowEmpty": {"static": true, "value": true}}, "can_be_gap": true, "can_be_step": true, "marking_script": "mark:\nswitch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)\n\ninterpreted_answer:\nqty(student_scalar, student_units)\n\n\n\ncorrect_quantity:\nsettings[\"correctAnswer\"]\n\n\n\ncorrect_units:\nunits(correct_quantity)\n\n\nallowed_notation_styles:\n[\"plain\",\"en\"]\n\nmatch_student_number:\nmatchnumber(studentAnswer,allowed_notation_styles)\n\nstudent_scalar:\nmatch_student_number[1]\n\nstudent_units:\nreplace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")\n\ngood_units:\ntry(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n\n\nstudent_quantity:\nswitch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n\n\npercent_error:\ntry(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   \n\nright:\npercent_error <= settings['right']\n\n\nclose:\nright_sign and percent_error <= settings['close']\n\nright_sign:\nsign(student_scalar) = sign(correct_quantity)", "marking_notes": [{"name": "mark", "description": "This is the main marking note. It should award credit and provide feedback based on the student's answer.", "definition": "switch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)"}, {"name": "interpreted_answer", "description": "A value representing the student's answer to this part.", "definition": "qty(student_scalar, student_units)\n\n"}, {"name": "correct_quantity", "description": "", "definition": "settings[\"correctAnswer\"]\n\n"}, {"name": "correct_units", "description": "", "definition": "units(correct_quantity)\n"}, {"name": "allowed_notation_styles", "description": "", "definition": "[\"plain\",\"en\"]"}, {"name": "match_student_number", "description": "", "definition": "matchnumber(studentAnswer,allowed_notation_styles)"}, {"name": "student_scalar", "description": "", "definition": "match_student_number[1]"}, {"name": "student_units", "description": "<p>Modify&nbsp;the&nbsp;unit portion of&nbsp;the student's answer by</p>\n<p>1. replacing \"ohms\" with \"ohm\" &nbsp;case insensitive</p>\n<p>2.&nbsp;replacing '-'&nbsp;with ' '&nbsp;</p>\n<p>3. replacing '&deg;' with ' deg'&nbsp;</p>\n<p>to allow answers like 10 ft-lb and 30&deg;</p>", "definition": "replace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")"}, {"name": "good_units", "description": "", "definition": "try(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n"}, {"name": "student_quantity", "description": "<p>This fixes the student answer for two common errors. &nbsp;</p>\n<p>If student_units are wrong&nbsp;&nbsp;- replace with correct units</p>\n<p>If student_scalar has the wrong sign - replace with right sign</p>\n<p>If student makes both errors, only one gets fixed.</p>", "definition": "switch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n"}, {"name": "percent_error", "description": "", "definition": "try(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   "}, {"name": "right", "description": "", "definition": "percent_error <= settings['right']\n"}, {"name": "close", "description": "<p>Only marked close if the student actually has the right sign.</p>", "definition": "right_sign and percent_error <= settings['close']"}, {"name": "right_sign", "description": "", "definition": "sign(student_scalar) = sign(correct_quantity) "}], "settings": [{"name": "correctAnswer", "label": "Correct Quantity.", "help_url": "", "hint": "The correct answer given as a JME quantity.", "input_type": "code", "default_value": "", "evaluate": true}, {"name": "right", "label": "% Accuracy for right.", "help_url": "", "hint": "Question will be considered correct if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "0.2", "evaluate": true}, {"name": "close", "label": "% Accuracy for close.", "help_url": "", "hint": "Question will be considered close if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "1.0", "evaluate": true}, {"name": "C1", "label": "Close with units.", "help_url": "", "hint": "Partial Credit for close value with appropriate units. &nbsp;if correct answer is 100 N and close is &plusmn;1%,<br />99 &nbsp;N is accepted.", "input_type": "percent", "default_value": "75"}, {"name": "C2", "label": "No units or wrong sign", "help_url": "", "hint": "Partial credit for forgetting units or using wrong sign.<br />If the correct answer is 100 N, both 100 and -100 N are accepted.", "input_type": "percent", "default_value": "50"}, {"name": "C3", "label": "Close, no units.", "help_url": "", "hint": "Partial Credit for close value but forgotten units.<br />This value would be close&nbsp;if the expected units were provided. &nbsp;If the correct answer is 100 N, and close is &plusmn;1%,<br />99 is accepted.", "input_type": "percent", "default_value": "25"}], "public_availability": "always", "published": true, "extensions": ["quantities"]}], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Parallel Axis Theorem: semi and quarter circles", "tags": ["Area moment of inertia", "area moment of inertia", "mechancs", "moment of inertia", "parallel axis theorem", "Statics", "statics"], "metadata": {"description": "<p>Find the moment of inertia of semi and quarter circles using the parallel axis theorem.</p>", "licence": "Creative Commons Attribution-NonCommercial 4.0 International"}, "statement": "<p>Determine the moment of inertia and radius of gyration about the $\\var{latex(axis)}$<em>-</em>axis for the {lower(names[pts[0]])} and the {lower(names[pts[1]])}. Grid units are [{units}].&nbsp;</p>\n<table style=\"margin-left: auto; margin-right: auto;\">\n<tbody>\n<tr>\n<td style=\"text-align: center;\">\n<p>{geogebra_applet('emjtbcxv',ggb_points)}</p>\n</td>\n</tr>\n</tbody>\n</table>\n<table data-jme-visible=\"debug\">\n<tbody>\n<tr>\n<td><span>axis</span></td>\n<td><span>radius</span></td>\n<td><span>area</span></td>\n<td><span>ibar</span></td>\n<td><span>d</span></td>\n<td><span>I</span></td>\n<td>Centroid</td>\n</tr>\n<tr>\n<td><span>x</span></td>\n<td><span>{r[\"A\"]}</span></td>\n<td><span>{area[\"A\"]}</span></td>\n<td><span>{ibarx[\"A\"]}</span></td>\n<td><span>{d[\"x\"][\"A\"]}</span></td>\n<td><span>{Ix[\"A\"]}</span></td>\n<td>{siground(Centroid[\"A\"],3)}</td>\n</tr>\n<tr>\n<td><span>y</span></td>\n<td><span>{r[\"A\"]}</span></td>\n<td><span>{area[\"A\"]}</span></td>\n<td><span>{ibary[\"A\"]}</span></td>\n<td><span>{d[\"y\"][\"A\"]}</span></td>\n<td><span>{Iy[\"A\"]}</span></td>\n<td></td>\n</tr>\n</tbody>\n</table>\n<p></p>\n<table data-jme-visible=\"debug\">\n<tbody>\n<tr>\n<td><span>axis</span></td>\n<td><span>radius</span></td>\n<td><span>area</span></td>\n<td><span>ibar</span></td>\n<td><span>d</span></td>\n<td><span>I</span></td>\n<td>Centroid</td>\n</tr>\n<tr>\n<td><span>x</span></td>\n<td><span>{r[\"B\"]}</span></td>\n<td><span>{area[\"B\"]}</span></td>\n<td><span>{ibarx[\"B\"]}</span></td>\n<td><span>{d[\"x\"][\"B\"]}</span></td>\n<td><span>{Ix[\"B\"]}</span></td>\n<td>{siground(Centroid[\"B\"],3)}</td>\n</tr>\n<tr>\n<td><span>y</span></td>\n<td><span>{r[\"B\"]}</span></td>\n<td><span>{area[\"B\"]}</span></td>\n<td><span>{ibary[\"B\"]}</span></td>\n<td><span>{d[\"y\"][\"B\"]}</span></td>\n<td><span>{Iy[\"B\"]}</span></td>\n<td></td>\n</tr>\n</tbody>\n</table>\n<p></p>\n<table data-jme-visible=\"debug\">\n<tbody>\n<tr>\n<td><span>axis</span></td>\n<td><span>radius</span></td>\n<td><span>area</span></td>\n<td><span>ibar</span></td>\n<td><span>d</span></td>\n<td><span>I</span></td>\n<td>Centroid</td>\n</tr>\n<tr>\n<td><span>x</span></td>\n<td><span>{r[\"C\"]}</span></td>\n<td><span>{area[\"C\"]}</span></td>\n<td><span>{ibarx[\"C\"]}</span></td>\n<td><span>{d[\"x\"][\"C\"]}</span></td>\n<td><span>{Ix[\"C\"]}</span></td>\n<td>{siground(Centroid[\"C\"],3)}</td>\n</tr>\n<tr>\n<td><span>y</span></td>\n<td><span>{r[\"C\"]}</span></td>\n<td><span>{area[\"C\"]}</span></td>\n<td><span>{ibary[\"C\"]}</span></td>\n<td><span>{d[\"y\"][\"C\"]}</span></td>\n<td><span>{Iy[\"C\"]}</span></td>\n<td></td>\n</tr>\n</tbody>\n</table>", "advice": "<div data-jme-visible=\"&quot;A&quot; in pts[0..2]\">\n<p><strong>For the circle:</strong></p>\n<p style=\"padding-left: 60px;\">$r = \\var{q(r[\"A\"],units)} \\qquad A = \\pi r^2 = \\var{q(area[\"A\"],units+'^2')}$</p>\n<p style=\"padding-left: 60px;\">$\\bar{x} =\\var{q(centroid[\"A\"][0],units)}&nbsp;\\qquad \\bar{y} = \\var{q(centroid[\"A\"][1],units)}$</p>\n<p style=\"padding-left: 30px;\">Place a {if(axis=\"x\", \"horizontal\", \"vertical\" )} axis through the centroid of the {names[\"A\"]}. &nbsp;The centroidal moment of inertia&nbsp;with respect to this axis is:</p>\n<p style=\"padding-left: 60px;\">$\\bar{I} &nbsp;= \\dfrac{\\pi r^4}{4} = \\var{q(Ibar[axis][\"A\"],units+'^4')}$</p>\n<p style=\"padding-left: 30px;\">and the&nbsp;distance between the parallel axes is:</p>\n<p style=\"padding-left: 60px;\">$d =&nbsp;\\var{q(abs(d[axis][\"A\"]),units)}$</p>\n<p style=\"padding-left: 30px;\">Applying the parallel axis theorem and the definition of radius of gyration:</p>\n<p style=\"padding-left: 60px;\">$I_{\\var{axis}} = \\bar{I} + A d^2 = \\var{q(I'[axis][\"A\"],units+'^4')}$</p>\n<p style=\"padding-left: 60px;\">$k_{\\var{axis}} = \\sqrt{\\dfrac{I_{\\var{axis}}}{A}} = \\var{q(k[axis][\"A\"],units)}$</p>\n<p style=\"padding-left: 60px;\"></p>\n</div>\n<div data-jme-visible=\"&quot;B&quot; in pts[0..2]\">\n<p><strong>For the semi-circle:</strong></p>\n<p style=\"padding-left: 60px;\">$r = \\var{q(r[\"B\"],units)} \\qquad A = \\dfrac{\\pi r^2}{2} = \\var{q(area[\"B\"],units+'^2')}$</p>\n<p style=\"padding-left: 60px;\">$\\bar{x} = \\var{latex( offset(0,centers[\"B\"], radii[\"B\"]))} \\var{q(centroid[\"B\"][0],units)}&nbsp;\\qquad&nbsp;\\bar{y} = \\var{latex( offset(1,centers[\"B\"], radii[\"B\"]))} \\var{q(centroid[\"B\"][1],units)}$</p>\n<p style=\"padding-left: 30px;\">Place a {if(axis=\"x\", \"horizontal\", \"vertical\" )} axis through the centroid of the {names[\"B\"]}. &nbsp;The centroidal moment of inertia&nbsp;with respect to this axis is:</p>\n<p style=\"padding-left: 60px;\">$\\bar{I} &nbsp;=\\var{ if(orientation=axis,latex(\"0.1098 r^4\"),latex(\"\\\\dfrac\\{\\\\pi r^4\\}\\{8\\}\")) }= \\var{q(Ibar[axis][\"B\"],units+'^4')}$</p>\n<p style=\"padding-left: 30px;\">and the&nbsp;distance between the parallel axes is:</p>\n<p style=\"padding-left: 60px;\">$d =&nbsp;\\var{q(abs(d[axis][\"B\"]),units)}$</p>\n<p style=\"padding-left: 30px;\">Applying the parallel axis theorem and the definition of radius of gyration:</p>\n<p style=\"padding-left: 60px;\">$I_{\\var{axis}} = \\bar{I} + A d^2 = \\var{q(I'[axis][\"B\"],units+'^4')}$</p>\n<p style=\"padding-left: 60px;\">$k_{\\var{axis}} = \\sqrt{\\dfrac{I_{\\var{axis}}}{A}} = \\var{q(k[axis][\"B\"],units)}$</p>\n<p style=\"padding-left: 60px;\"></p>\n</div>\n<div data-jme-visible=\"&quot;C&quot; in pts[0..2]\">\n<p><strong>For the quarter-circle:</strong></p>\n<p style=\"padding-left: 60px;\">$r = \\var{q(r[\"C\"],units)} \\qquad A = \\dfrac{\\pi r^2}{4} = \\var{q(area[\"C\"],units+'^2')}$</p>\n<p style=\"padding-left: 60px;\">$\\bar{x} =&nbsp;\\var{latex( offset(0,centers[\"C\"], radii[\"C\"]))} \\var{q(centroid[\"C\"][0],units)}&nbsp;\\qquad&nbsp;\\bar{y} =\\var{latex( offset(1,centers[\"C\"], radii[\"C\"]))} \\var{q(centroid[\"C\"][1],units)}$</p>\n<p style=\"padding-left: 30px;\">Place a {if(axis=\"x\", \"horizontal\", \"vertical\" )} axis through the centroid of the {names[\"C\"]}. &nbsp;The centroidal moment of inertia&nbsp;with respect to this axis is:</p>\n<p style=\"padding-left: 60px;\">$\\bar{I} &nbsp;= 0.0549 r^4= \\var{q(Ibar[axis][\"C\"],units+'^4')}$</p>\n<p style=\"padding-left: 30px;\">and the&nbsp;distance between the parallel axes is:</p>\n<p style=\"padding-left: 60px;\">$d =&nbsp;\\var{q(abs(d[axis][\"C\"]),units)}$</p>\n<p style=\"padding-left: 30px;\">Applying the parallel axis theorem and the definition of radius of gyration:</p>\n<p style=\"padding-left: 60px;\">$I_{\\var{axis}} = \\bar{I} + A d^2 = \\var{q(I'[axis][\"C\"],units+'^4')}$</p>\n<p style=\"padding-left: 60px;\">$k_{\\var{axis}} = \\sqrt{\\dfrac{I_{\\var{axis}}}{A}} = \\var{q(k[axis][\"C\"],units)}$</p>\n</div>", "rulesets": {}, "extensions": ["geogebra", "quantities"], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"area": {"name": "area", "group": "calculated", "definition": "[\"A\": pi r[\"A\"]^2, \"B\": pi r[\"B\"]^2/2, \"C\": pi r[\"C\"]^2 / 4]", "description": "", "templateType": "anything", "can_override": false}, "I0": {"name": "I0", "group": "Unnamed group", "definition": "qty(I'[axis][pts[0]],units+\"^4\")", "description": "", "templateType": "anything", "can_override": false}, "ky": {"name": "ky", "group": "calculated", "definition": "dict(map([p,sqrt(Iy[p]/area[p])], p, sort(pts)))", "description": "", "templateType": "anything", "can_override": false}, "Iy": {"name": "Iy", "group": "calculated", "definition": "dict(map([p,Ibary[p] + Area[p] * \n      d[\"y\"][p]^2],p,sort(pts)))", "description": "", "templateType": "anything", "can_override": false}, "k1": {"name": "k1", "group": "Unnamed group", "definition": "qty(k[axis][pts[1]],units)", "description": "", "templateType": "anything", "can_override": false}, "Ibar": {"name": "Ibar", "group": "calculated", "definition": "[\"x\":Ibarx,\"y\":Ibary]", "description": "", "templateType": "anything", "can_override": false}, "kx": {"name": "kx", "group": "calculated", "definition": "dict(map([p,sqrt(Ix[p]/area[p])], p, sort(pts)))", "description": "", "templateType": "anything", "can_override": false}, "Ibarx": {"name": "Ibarx", "group": "calculated", "definition": "[\"A\": pi r[\"A\"]^4 /4, \"B\": if(orientation=\"x\",0.1098 r[\"B\"]^4, pi r[\"B\"]^4/8), \"C\": 0.0549 r[\"C\"]^4]", "description": "", "templateType": "anything", "can_override": false}, "I1": {"name": "I1", "group": "Unnamed group", "definition": "qty(I'[axis][pts[1]],units+\"^4\")", "description": "", "templateType": "anything", "can_override": false}, "Ibary": {"name": "Ibary", "group": "calculated", "definition": "[\"A\": pi r[\"A\"]^4 /4, \"B\": if(orientation=\"x\",pi r[\"B\"]^4/8,0.1098 r[\"B\"]^4), \"C\": 0.0549 r[\"C\"]^4]", "description": "", "templateType": "anything", "can_override": false}, "centers": {"name": "centers", "group": "Ungrouped variables", "definition": "dict(map([p,\n      vector(random(-6..6),random(-5..5))],p,pts))", "description": "<p>randomly located center points for the shapes</p>", "templateType": "anything", "can_override": false}, "k0": {"name": "k0", "group": "Unnamed group", "definition": "qty(k[axis][pts[0]],units)", "description": "", "templateType": "anything", "can_override": false}, "radii": {"name": "radii", "group": "Ungrouped variables", "definition": "[\"A\":centers[\"A\"] + random(1..5) * vector(0,1), \"B\":centers[\"B\"] + random(1..5)  * vector(random([1,0],[-1,0],[0,-1],[0,1])), \"C\":centers[\"C\"] + random(1..5) / sqrt(2) * vector(random([1,1],[-1,1],[-1,-1],[-1,1]))]", "description": "<p>This is a point which defines the radius. &nbsp;The centroid is on a line passing through center and this point.</p>\n<p>The actual radius is the distance between center and radii.</p>\n<p></p>", "templateType": "anything", "can_override": false}, "names": {"name": "names", "group": "Ungrouped variables", "definition": "[\"A\":\"Circle\",\"B\":\"Semi-circle\",\"C\":\"Quarter-circle\"]\n", "description": "<p>the names of the shapes</p>", "templateType": "anything", "can_override": false}, "centroid": {"name": "centroid", "group": "calculated", "definition": "[\"A\": centers[\"A\"], \"B\": centers[\"B\"] + 4 (radii[\"B\"] - centers[\"B\"]) /3 /pi, \"C\": centers[\"C\"] +   4 sqrt(2) (radii[\"C\"] - centers[\"C\"]) /3 /pi]", "description": "<p>need sqrt(2) for quarter circle because of the way radii is defined</p>", "templateType": "anything", "can_override": false}, "axis": {"name": "axis", "group": "Ungrouped variables", "definition": "random(latex('x'),latex('y'))", "description": "<p>random choice for question&nbsp;</p>", "templateType": "anything", "can_override": false}, "I'": {"name": "I'", "group": "calculated", "definition": "[\"x\":Ix,\"y\":Iy]", "description": "<p>Cant call it I because I is a reserved variable</p>", "templateType": "anything", "can_override": false}, "k": {"name": "k", "group": "calculated", "definition": "[\"x\":kx,\"y\":ky]", "description": "", "templateType": "anything", "can_override": false}, "units": {"name": "units", "group": "Ungrouped variables", "definition": "random('in','cm','ft','mm')", "description": "", "templateType": "anything", "can_override": false}, "d": {"name": "d", "group": "calculated", "definition": "[\"x\": dict(map([p,centroid[p][1]],p,sort(pts))), \"y\": dict(map([p,centroid[p][0]],p,sort(pts)))]", "description": "<p>d[\"x\"] gives distance from the&nbsp;y-axis</p>\n<p>This is the value needed for parallel axis theorem when calculating I_x</p>", "templateType": "anything", "can_override": false}, "Ix": {"name": "Ix", "group": "calculated", "definition": "dict(map([p,Ibarx[p] + Area[p] * \n      d[\"x\"][p]^2],p,sort(pts)))", "description": "", "templateType": "anything", "can_override": false}, "orientation": {"name": "orientation", "group": "calculated", "definition": "if(centers[\"B\"][0] = radii[\"B\"][0],\"x\",\"y\")", "description": "<p>This determines which axis the semicircle's diameter is parallel to.</p>", "templateType": "anything", "can_override": false}, "debug": {"name": "debug", "group": "Ungrouped variables", "definition": "false", "description": "", "templateType": "anything", "can_override": false}, "pts": {"name": "pts", "group": "Ungrouped variables", "definition": "shuffle([\"A\",\"B\",\"C\"])", "description": "<p>this is where the randomizing of the shapes occurs.</p>", "templateType": "anything", "can_override": false}, "ggb_points": {"name": "ggb_points", "group": "Ungrouped variables", "definition": "map([p,centers[p]],p,pts) +\nmap(['R'+p,radii[p]],p,pts)[0..2]+\nmap(['R'+p,centers[p]],p,pts)[2..3]", "description": "<p>The problem is about two shapes randomly selected, which are the first two values in the pts array.</p>\n<p>The first map is the center coordinates for all three shapes.</p>\n<p>the second map is the random radii for the first two shapes,</p>\n<p>The last map&nbsp;sets the radius of the third&nbsp;shapes to its center, which hides it in GGB.</p>\n<p></p>\n<p>map([p,centers[p]],p,pts) +<br />map(['R'+p,radii[p]],p,pts)[0..2]+<br />map(['R'+p,centers[p]],p,pts)[2..3]</p>", "templateType": "anything", "can_override": false}, "r": {"name": "r", "group": "calculated", "definition": "dict(map([p,round(abs(centers[p]-radii[p]))],p,pts))", "description": "<p>the radii of the three shapes</p>", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["pts", "axis", "ggb_points", "names", "centers", "radii", "units", "debug"], "variable_groups": [{"name": "calculated", "variables": ["r", "area", "centroid", "Ibarx", "Iy", "Ibary", "Ix", "ky", "kx", "Ibar", "k", "I'", "d", "orientation"]}, {"name": "Unnamed group", "variables": ["I0", "I1", "k0", "k1"]}], "functions": {"": {"parameters": [], "type": "number", "language": "jme", "definition": ""}, "q": {"parameters": [["q", "number"], ["u", "string"]], "type": "string", "language": "jme", "definition": "siground(qty(q,u),4)"}, "offset": {"parameters": [["way", "number"], ["c", "vector"], ["d", "vector"]], "type": "string", "language": "jme", "definition": "let(dir,sign(d[way]-c[way]),r,siground(abs(c-d),3),\n    if(dir=0,'', c[way] + \"\\\\textrm{ \" + units + \"\\} \" + if(dir=1,\"+\",\"-\") + \n\"\\\\dfrac\\{4r\\}\\{3\\\\pi\\} =\"))"}}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": true, "customName": "{names[pts[0]]}", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>$I_\\var{axis}$ = &nbsp;[[0]] &nbsp;</p>\n<p>$k_\\var{axis}$ = &nbsp;[[1]] &nbsp;</p>\n<p><span data-jme-visible=\"debug\">{siground(I0,4)}</span> <span data-jme-visible=\"debug\">{siground(k0,4)}</span></p>", "gaps": [{"type": "engineering-answer", "useCustomName": true, "customName": "I", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "I0", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "k", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "k0", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}], "sortAnswers": false}, {"type": "gapfill", "useCustomName": true, "customName": "{names[pts[1]]}", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>$I_\\var{axis}$ = &nbsp;[[0]] &nbsp;</p>\n<p>$k_\\var{axis}$ = &nbsp;[[1]] &nbsp;</p>\n<p><span data-jme-visible=\"debug\">{siground(I1,4)}</span> <span data-jme-visible=\"debug\">{siground(k1,4)}</span></p>", "gaps": [{"type": "engineering-answer", "useCustomName": true, "customName": "I", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "I1", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "k", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "k1", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}], "sortAnswers": false}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}], "resources": []}]}], "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}]}