// Numbas version: finer_feedback_settings
{"name": "Centroidal Moment of inertia: T shape", "extensions": ["geogebra", "quantities"], "custom_part_types": [{"source": {"pk": 19, "author": {"name": "William Haynes", "pk": 2530}, "edit_page": "/part_type/19/edit"}, "name": "Engineering Accuracy with units", "short_name": "engineering-answer", "description": "<p>A value with units marked right if within an adjustable % error of the correct value. &nbsp;Marked close if within a wider margin of error.</p>", "help_url": "", "input_widget": "string", "input_options": {"correctAnswer": "siground(settings['correctAnswer'],4)", "hint": {"static": true, "value": ""}, "allowEmpty": {"static": true, "value": true}}, "can_be_gap": true, "can_be_step": true, "marking_script": "mark:\nswitch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)\n\ninterpreted_answer:\nqty(student_scalar, student_units)\n\n\n\ncorrect_quantity:\nsettings[\"correctAnswer\"]\n\n\n\ncorrect_units:\nunits(correct_quantity)\n\n\nallowed_notation_styles:\n[\"plain\",\"en\"]\n\nmatch_student_number:\nmatchnumber(studentAnswer,allowed_notation_styles)\n\nstudent_scalar:\nmatch_student_number[1]\n\nstudent_units:\nreplace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")\n\ngood_units:\ntry(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n\n\nstudent_quantity:\nswitch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n\n\npercent_error:\ntry(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   \n\nright:\npercent_error <= settings['right']\n\n\nclose:\nright_sign and percent_error <= settings['close']\n\nright_sign:\nsign(student_scalar) = sign(correct_quantity)", "marking_notes": [{"name": "mark", "description": "This is the main marking note. It should award credit and provide feedback based on the student's answer.", "definition": "switch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)"}, {"name": "interpreted_answer", "description": "A value representing the student's answer to this part.", "definition": "qty(student_scalar, student_units)\n\n"}, {"name": "correct_quantity", "description": "", "definition": "settings[\"correctAnswer\"]\n\n"}, {"name": "correct_units", "description": "", "definition": "units(correct_quantity)\n"}, {"name": "allowed_notation_styles", "description": "", "definition": "[\"plain\",\"en\"]"}, {"name": "match_student_number", "description": "", "definition": "matchnumber(studentAnswer,allowed_notation_styles)"}, {"name": "student_scalar", "description": "", "definition": "match_student_number[1]"}, {"name": "student_units", "description": "<p>Modify&nbsp;the&nbsp;unit portion of&nbsp;the student's answer by</p>\n<p>1. replacing \"ohms\" with \"ohm\" &nbsp;case insensitive</p>\n<p>2.&nbsp;replacing '-'&nbsp;with ' '&nbsp;</p>\n<p>3. replacing '&deg;' with ' deg'&nbsp;</p>\n<p>to allow answers like 10 ft-lb and 30&deg;</p>", "definition": "replace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")"}, {"name": "good_units", "description": "", "definition": "try(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n"}, {"name": "student_quantity", "description": "<p>This fixes the student answer for two common errors. &nbsp;</p>\n<p>If student_units are wrong&nbsp;&nbsp;- replace with correct units</p>\n<p>If student_scalar has the wrong sign - replace with right sign</p>\n<p>If student makes both errors, only one gets fixed.</p>", "definition": "switch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n"}, {"name": "percent_error", "description": "", "definition": "try(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   "}, {"name": "right", "description": "", "definition": "percent_error <= settings['right']\n"}, {"name": "close", "description": "<p>Only marked close if the student actually has the right sign.</p>", "definition": "right_sign and percent_error <= settings['close']"}, {"name": "right_sign", "description": "", "definition": "sign(student_scalar) = sign(correct_quantity) "}], "settings": [{"name": "correctAnswer", "label": "Correct Quantity.", "help_url": "", "hint": "The correct answer given as a JME quantity.", "input_type": "code", "default_value": "", "evaluate": true}, {"name": "right", "label": "% Accuracy for right.", "help_url": "", "hint": "Question will be considered correct if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "0.2", "evaluate": true}, {"name": "close", "label": "% Accuracy for close.", "help_url": "", "hint": "Question will be considered close if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "1.0", "evaluate": true}, {"name": "C1", "label": "Close with units.", "help_url": "", "hint": "Partial Credit for close value with appropriate units. &nbsp;if correct answer is 100 N and close is &plusmn;1%,<br />99 &nbsp;N is accepted.", "input_type": "percent", "default_value": "75"}, {"name": "C2", "label": "No units or wrong sign", "help_url": "", "hint": "Partial credit for forgetting units or using wrong sign.<br />If the correct answer is 100 N, both 100 and -100 N are accepted.", "input_type": "percent", "default_value": "50"}, {"name": "C3", "label": "Close, no units.", "help_url": "", "hint": "Partial Credit for close value but forgotten units.<br />This value would be close&nbsp;if the expected units were provided. &nbsp;If the correct answer is 100 N, and close is &plusmn;1%,<br />99 is accepted.", "input_type": "percent", "default_value": "25"}], "public_availability": "always", "published": true, "extensions": ["quantities"]}], "resources": [["question-resources/Tshape.ggb", "Tshape.ggb"]], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Centroidal Moment of inertia: T shape", "tags": ["centroid", "Mechanics", "mechanics", "moment of inertia", "Statics", "statics"], "metadata": {"description": "<p>Find the centroidal moment of inertia of a sideways T shape. This requires first locating the centroid, then applying&nbsp;the parallel axis theorem.</p>", "licence": "Creative Commons Attribution-NonCommercial 4.0 International"}, "statement": "<p>Determine the moment of inertia of the <strong>T</strong> shape with respect to a vertical axis passing through the shape's centroid. &nbsp;Let this axis be designated $a$. Grid units are [{units}].</p>\n<table style=\"margin-left: auto; margin-right: auto;\">\n<tbody>\n<tr>\n<td style=\"text-align: center;\">{geogebra_applet('wfhyktjn',['A':P1, 'B': P2])}</td>\n</tr>\n</tbody>\n</table>", "advice": "<p></p>\n<p>Number the parts: &nbsp;</p>\n<p style=\"padding-left: 30px;\">Part 1 horizontal rectangle: $b_1$ = {qty(b1,units)}, $h_1$ = {qty(h1,units)}.</p>\n<p style=\"padding-left: 30px;\">Part 2 vertical rectangle: $b_2$ = {qty(b2,units)}, $h_2$ = {qty(h2,units)}.</p>\n<p style=\"padding-left: 30px;\"></p>\n<p>Calculate the areas:</p>\n<p style=\"padding-left: 30px;\">$ A_1 = b_1 h_1=\\var{qty(A1,units+'^2')}$</p>\n<p style=\"padding-left: 30px;\">$&nbsp;A_2 =&nbsp;b_2 h_2 = \\var{qty(A2,units+'^2')}$</p>\n<p>Locate the centroids of the parts.</p>\n<p style=\"padding-left: 30px;\">&nbsp; &nbsp; $ \\bar{x}_1 = b_1/2 = \\var{qty(d1,units)} \\qquad \\bar{x}_2 = b_1 + b_2/2 = \\var{qty(d2,units)}$</p>\n<p>Locate the centroid of the whole.</p>\n<p style=\"padding-left: 30px;\">&nbsp; &nbsp;$\\bar{x} = \\dfrac{ A_1 \\bar{x}_1 + A_2 \\bar{x}_2} {A_1 + A_2} = \\var{qty(siground(xbar,4),units)}$</p>\n<p style=\"padding-left: 30px;\">&nbsp; &nbsp;$\\bar{y} = \\var{qty(0, units)} $, by symmetry</p>\n<p>Calculate the moment of inertia with respect to the <em>y</em>-axis for&nbsp;the parts and the whole.</p>\n<p style=\"padding-left: 30px;\">$\\begin{align} \\\\{I_y}_1 &amp;= \\tfrac{1}{3} h_1 b_1^3 \\\\&amp;= \\tfrac{1}{3}( \\var{h1})(\\var{b1})^3 \\\\&amp;= \\var{qty(display(Iy1),units+'^4')} \\\\\\\\{I_y}_2 &amp;= [\\bar{I} + A d^2]_2\\\\ &amp;= \\frac{1}{12} h_2 b_2^3 + A_2 d_2^2 \\\\&amp;= \\var{display(Ibar2)} + (\\var{display(A2)}) (\\var{display(d2)})^2\\\\&amp;=\\var{qty(display(Iy2),units+'^4')} \\\\\\\\I_y &amp;= {I_y}_1+{I_y}_2\\\\&amp;=\\var{display(Iy1)}+\\var{display(Iy2)}\\\\&amp;=\\var{qty(display(Iy),units+'^4')}\\end{align}$</p>\n<p>Use the parallel axis theorem $I &nbsp;= &nbsp;\\bar{I} + A d^2$ backwards to find the centroidal moment of inertia.</p>\n<p style=\"padding-left: 30px;\">$\\begin{align} \\bar{I}_a &amp;= I_y - A (\\bar{x})^2 \\\\&amp;= \\var{display(Iy)} - \\var{display(A)} (\\var{display(xbar)})^2\\\\&amp;=\\var{qty(display(I_a),units+'^4')}\\end{align}$</p>\n<p></p>", "rulesets": {}, "extensions": ["geogebra", "quantities"], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"A2": {"name": "A2", "group": "Area", "definition": "b2 * h2 ", "description": "", "templateType": "anything", "can_override": false}, "A1": {"name": "A1", "group": "Area", "definition": "b1 h1", "description": "", "templateType": "anything", "can_override": false}, "Iy2": {"name": "Iy2", "group": "Iy", "definition": "h2  b2^3/12 + A2 * (b1 + b2/2)^2", "description": "", "templateType": "anything", "can_override": false}, "ky": {"name": "ky", "group": "Iy", "definition": "sqrt(Iy/A)", "description": "", "templateType": "anything", "can_override": false}, "applet": {"name": "applet", "group": "Inputs", "definition": "geogebra_file('Tshape.ggb',['A':P1, 'B': P2])\n//geogebra_applet('wfhyktjn',['A':P1, 'B': P2])", "description": "", "templateType": "anything", "can_override": false}, "A": {"name": "A", "group": "Area", "definition": "A1 + A2", "description": "", "templateType": "anything", "can_override": false}, "units": {"name": "units", "group": "Inputs", "definition": "random('in','cm')", "description": "<p>width of P1</p>", "templateType": "anything", "can_override": false}, "debug": {"name": "debug", "group": "Inputs", "definition": "false", "description": "<p>true</p>", "templateType": "anything", "can_override": false}, "Iy1": {"name": "Iy1", "group": "Iy", "definition": "h1 b1^3/12 + A1 * (b1/2)^2", "description": "", "templateType": "anything", "can_override": false}, "P1": {"name": "P1", "group": "Inputs", "definition": "vector(random(5..12),random(1..2))", "description": "<p>Point A in GGB</p>", "templateType": "anything", "can_override": false}, "P2": {"name": "P2", "group": "Inputs", "definition": "P1 + vector(random(1..3),random(4..6))", "description": "", "templateType": "anything", "can_override": false}, "b1": {"name": "b1", "group": "Inputs", "definition": "P1[0]", "description": "", "templateType": "anything", "can_override": false}, "h1": {"name": "h1", "group": "Inputs", "definition": "2 P1[1]", "description": "<p>height of p1</p>", "templateType": "anything", "can_override": false}, "b2": {"name": "b2", "group": "Inputs", "definition": "p2[0]-p1[0]", "description": "<p>width of p2</p>", "templateType": "anything", "can_override": false}, "h2": {"name": "h2", "group": "Inputs", "definition": "2 p2[1]", "description": "<p>height of p2</p>", "templateType": "anything", "can_override": false}, "ibar1": {"name": "ibar1", "group": "check", "definition": "h1 b1^3/12", "description": "<p>centroidal moi</p>", "templateType": "anything", "can_override": false}, "Ibar2": {"name": "Ibar2", "group": "check", "definition": "h2 b2^3/12", "description": "", "templateType": "anything", "can_override": false}, "d1": {"name": "d1", "group": "check", "definition": "b1/2", "description": "", "templateType": "anything", "can_override": false}, "d2": {"name": "d2", "group": "check", "definition": "b1 + b2/2", "description": "", "templateType": "anything", "can_override": false}, "Iy1'": {"name": "Iy1'", "group": "check", "definition": "ibar1 + A1 d1^2", "description": "<p>alternate calculation of Iy</p>", "templateType": "anything", "can_override": false}, "Iy2'": {"name": "Iy2'", "group": "check", "definition": "Ibar2 + a2 d2^2", "description": "", "templateType": "anything", "can_override": false}, "Iy": {"name": "Iy", "group": "check", "definition": "Iy1' + Iy2'", "description": "<p>Moment of intertia about y axis</p>", "templateType": "anything", "can_override": false}, "Qy": {"name": "Qy", "group": "check", "definition": "A1 d1 + A2 d2", "description": "", "templateType": "anything", "can_override": false}, "xbar": {"name": "xbar", "group": "check", "definition": "Qy/A", "description": "", "templateType": "anything", "can_override": false}, "I_a": {"name": "I_a", "group": "check", "definition": "Iy - A xbar^2", "description": "<p>Centroidal moment of inertia of total shape</p>", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": [], "variable_groups": [{"name": "Inputs", "variables": ["units", "debug", "applet", "P1", "P2", "b1", "h1", "b2", "h2"]}, {"name": "Area", "variables": ["A1", "A2", "A"]}, {"name": "Iy", "variables": ["Iy1", "Iy2", "ky"]}, {"name": "check", "variables": ["ibar1", "Ibar2", "d1", "d2", "Iy1'", "Iy2'", "Iy", "Qy", "xbar", "I_a"]}], "functions": {"display": {"parameters": [["q", "number"]], "type": "number", "language": "jme", "definition": "(siground(q,4))"}}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": true, "customName": "Centroid", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>$\\bar{x} = $ [[0]]&nbsp;</p>\n<p>$\\bar{y} = $ [[1]]</p>", "gaps": [{"type": "engineering-answer", "useCustomName": true, "customName": "xbar", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(xbar,units)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "ybar", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(0,units)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}], "sortAnswers": false}, {"type": "gapfill", "useCustomName": true, "customName": "For the $a$ axis", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>$\\bar{I}_a$ = [[0]]&nbsp;<span data-jme-visible=\"debug\">{display(I_a)}</span></p>", "gaps": [{"type": "engineering-answer", "useCustomName": true, "customName": "$\\bar{I}_{a}$", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(I_a, units+'^4')", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}], "sortAnswers": false}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}], "resources": ["question-resources/Tshape.ggb"]}]}], "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}]}