// Numbas version: finer_feedback_settings
{"name": "Method of sections: Pratt Truss", "extensions": ["geogebra", "quantities"], "custom_part_types": [{"source": {"pk": 19, "author": {"name": "William Haynes", "pk": 2530}, "edit_page": "/part_type/19/edit"}, "name": "Engineering Accuracy with units", "short_name": "engineering-answer", "description": "<p>A value with units marked right if within an adjustable % error of the correct value. &nbsp;Marked close if within a wider margin of error.</p>", "help_url": "", "input_widget": "string", "input_options": {"correctAnswer": "siground(settings['correctAnswer'],4)", "hint": {"static": true, "value": ""}, "allowEmpty": {"static": true, "value": true}}, "can_be_gap": true, "can_be_step": true, "marking_script": "mark:\nswitch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)\n\ninterpreted_answer:\nqty(student_scalar, student_units)\n\n\n\ncorrect_quantity:\nsettings[\"correctAnswer\"]\n\n\n\ncorrect_units:\nunits(correct_quantity)\n\n\nallowed_notation_styles:\n[\"plain\",\"en\"]\n\nmatch_student_number:\nmatchnumber(studentAnswer,allowed_notation_styles)\n\nstudent_scalar:\nmatch_student_number[1]\n\nstudent_units:\nreplace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")\n\ngood_units:\ntry(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n\n\nstudent_quantity:\nswitch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n\n\npercent_error:\ntry(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   \n\nright:\npercent_error <= settings['right']\n\n\nclose:\nright_sign and percent_error <= settings['close']\n\nright_sign:\nsign(student_scalar) = sign(correct_quantity)", "marking_notes": [{"name": "mark", "description": "This is the main marking note. It should award credit and provide feedback based on the student's answer.", "definition": "switch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)"}, {"name": "interpreted_answer", "description": "A value representing the student's answer to this part.", "definition": "qty(student_scalar, student_units)\n\n"}, {"name": "correct_quantity", "description": "", "definition": "settings[\"correctAnswer\"]\n\n"}, {"name": "correct_units", "description": "", "definition": "units(correct_quantity)\n"}, {"name": "allowed_notation_styles", "description": "", "definition": "[\"plain\",\"en\"]"}, {"name": "match_student_number", "description": "", "definition": "matchnumber(studentAnswer,allowed_notation_styles)"}, {"name": "student_scalar", "description": "", "definition": "match_student_number[1]"}, {"name": "student_units", "description": "<p>Modify&nbsp;the&nbsp;unit portion of&nbsp;the student's answer by</p>\n<p>1. replacing \"ohms\" with \"ohm\" &nbsp;case insensitive</p>\n<p>2.&nbsp;replacing '-'&nbsp;with ' '&nbsp;</p>\n<p>3. replacing '&deg;' with ' deg'&nbsp;</p>\n<p>to allow answers like 10 ft-lb and 30&deg;</p>", "definition": "replace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")"}, {"name": "good_units", "description": "", "definition": "try(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n"}, {"name": "student_quantity", "description": "<p>This fixes the student answer for two common errors. &nbsp;</p>\n<p>If student_units are wrong&nbsp;&nbsp;- replace with correct units</p>\n<p>If student_scalar has the wrong sign - replace with right sign</p>\n<p>If student makes both errors, only one gets fixed.</p>", "definition": "switch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n"}, {"name": "percent_error", "description": "", "definition": "try(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   "}, {"name": "right", "description": "", "definition": "percent_error <= settings['right']\n"}, {"name": "close", "description": "<p>Only marked close if the student actually has the right sign.</p>", "definition": "right_sign and percent_error <= settings['close']"}, {"name": "right_sign", "description": "", "definition": "sign(student_scalar) = sign(correct_quantity) "}], "settings": [{"name": "correctAnswer", "label": "Correct Quantity.", "help_url": "", "hint": "The correct answer given as a JME quantity.", "input_type": "code", "default_value": "", "evaluate": true}, {"name": "right", "label": "% Accuracy for right.", "help_url": "", "hint": "Question will be considered correct if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "0.2", "evaluate": true}, {"name": "close", "label": "% Accuracy for close.", "help_url": "", "hint": "Question will be considered close if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "1.0", "evaluate": true}, {"name": "C1", "label": "Close with units.", "help_url": "", "hint": "Partial Credit for close value with appropriate units. &nbsp;if correct answer is 100 N and close is &plusmn;1%,<br />99 &nbsp;N is accepted.", "input_type": "percent", "default_value": "75"}, {"name": "C2", "label": "No units or wrong sign", "help_url": "", "hint": "Partial credit for forgetting units or using wrong sign.<br />If the correct answer is 100 N, both 100 and -100 N are accepted.", "input_type": "percent", "default_value": "50"}, {"name": "C3", "label": "Close, no units.", "help_url": "", "hint": "Partial Credit for close value but forgotten units.<br />This value would be close&nbsp;if the expected units were provided. &nbsp;If the correct answer is 100 N, and close is &plusmn;1%,<br />99 is accepted.", "input_type": "percent", "default_value": "25"}], "public_availability": "always", "published": true, "extensions": ["quantities"]}], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Method of sections: Pratt Truss", "tags": ["Mechanics", "mechanics", "method of sections", "Statics", "statics", "truss"], "metadata": {"description": "<p>Solve for the internal force in three members of a truss.</p>", "licence": "Creative Commons Attribution-NonCommercial 4.0 International"}, "statement": "<p>The Pratt truss are shown supports three loads: {load_html(0,units[0])}, {load_html(1,units[0])}, and {load_html(2, units[0])}. Determine the reactions and the forces in members <span data-jme-visible=\"version=0\">$CD$, $HJ$ and $CJ$</span> <span data-jme-visible=\"version=1\">$DE$, $EJ$, and $JK$</span>&nbsp;using the method of sections. Indicate tension or compression.</p>\n<table style=\"margin-left: auto; margin-right: auto;\">\n<tbody>\n<tr>\n<td>\n<p>{geogebra_applet('yrdq38xt',[['width',width],['height',height],['units','\"'+units[1]+'\"']]+ggb_forces)}</p>\n</td>\n</tr>\n</tbody>\n</table>", "advice": "<ol>\n<li>Draw a free body diagram of the whole truss and take moments at <em>A</em> to find force <em>G</em>, and&nbsp;take moments at <em>G</em> to find <em>A</em>. &nbsp;Check your answer by ensuring that $\\Sigma F_y = 0$.</li>\n<li>Take an imaginary cut through the three members of interest to divide the truss into two sections.</li>\n<li>Draw a free body diagram of one side of the truss. &nbsp;Indicate forces at the cut members as tension (pulling away from the joint). &nbsp;Show only the loads which act on the isolated portion of the truss</li>\n<li>Apply three equations of equilibrium to find the forces in the cut members. &nbsp;Negative values will indicate compression. &nbsp;It is good idea to take moments at a point where two unknowns intersect when possible.</li>\n</ol>", "rulesets": {}, "extensions": ["geogebra", "quantities"], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"ans_name": {"name": "ans_name", "group": "magnitudes", "definition": "[[\"CD\",\"DE\"][version],[\"HJ\",\"EJ\"][version],[\"CJ\",\"JK\"][version]]", "description": "", "templateType": "anything", "can_override": false}, "HJ": {"name": "HJ", "group": "magnitudes", "definition": "(A_y width - A_x height)/height", "description": "", "templateType": "anything", "can_override": false}, "units": {"name": "units", "group": "Inputs", "definition": "random(['N','m'],['lb','ft'])", "description": "", "templateType": "anything", "can_override": false}, "DE": {"name": "DE", "group": "magnitudes", "definition": "((F[\"E\"] + F[\"K\"] - 2 G) width/height + F[\"F\"])", "description": "", "templateType": "anything", "can_override": false}, "CD": {"name": "CD", "group": "magnitudes", "definition": "((F[\"C\"]+F[\"H\"]-2 A_y) width + F[\"B\"]* height)/height", "description": "<p>sigma M_j = 0&nbsp;</p>", "templateType": "anything", "can_override": false}, "version": {"name": "version", "group": "Inputs", "definition": "random(0,1)", "description": "", "templateType": "anything", "can_override": false}, "JK": {"name": "JK", "group": "magnitudes", "definition": "G width / height", "description": "", "templateType": "anything", "can_override": false}, "EJ_y": {"name": "EJ_y", "group": "magnitudes", "definition": "G - F[\"E\"] - f[\"K\"]", "description": "<p>sigma f_y=0</p>", "templateType": "anything", "can_override": false}, "ggb_forces": {"name": "ggb_forces", "group": "Inputs", "definition": "map([lower(n),F[n]],n,loads)", "description": "", "templateType": "anything", "can_override": false}, "CJ": {"name": "CJ", "group": "magnitudes", "definition": "CJ_y/sin(radians(theta))", "description": "", "templateType": "anything", "can_override": false}, "ans": {"name": "ans", "group": "magnitudes", "definition": "[[CD,DE][version],[HJ,EJ][version],[CJ,JK][version]]\n\n", "description": "", "templateType": "anything", "can_override": false}, "G": {"name": "G", "group": "magnitudes", "definition": "sum(map(F[k]*D[k],k,loads))/4", "description": "<p>sum of moments about A / dperp&nbsp;</p>", "templateType": "anything", "can_override": false}, "F": {"name": "F", "group": "Inputs", "definition": "dict(map([n,if(n in loads,random(100..500#25),0)],n,names))", "description": "<p>force magnitudes&nbsp;</p>", "templateType": "anything", "can_override": false}, "D": {"name": "D", "group": "Inputs", "definition": "dict(map([names[n],[0,-height/width,1,1,2,2,3,3,height/width,4][n]],n,0..9))", "description": "<p>Normalized perpendicular&nbsp;distances from point A</p>\n<p>multiply this by width to get&nbsp;actual distance</p>\n<p>positive moments except for force B which acts left</p>", "templateType": "anything", "can_override": false}, "width": {"name": "width", "group": "Inputs", "definition": "random(3,4,5,6)", "description": "", "templateType": "anything", "can_override": false}, "height": {"name": "height", "group": "Inputs", "definition": " width + random(-1.5..1.5#0.5)\n", "description": "", "templateType": "anything", "can_override": false}, "EJ": {"name": "EJ", "group": "magnitudes", "definition": "EJ_y/sin(radians(theta))", "description": "", "templateType": "anything", "can_override": false}, "CJ_y": {"name": "CJ_y", "group": "magnitudes", "definition": "A_y - F[\"C\"]-F[\"H\"]", "description": "<p>sum of moments about A/dperp</p>", "templateType": "anything", "can_override": false}, "theta": {"name": "theta", "group": "Inputs", "definition": "degrees(arctan(height/width))", "description": "<p>Angle of diagonal members</p>", "templateType": "anything", "can_override": false}, "A": {"name": "A", "group": "magnitudes", "definition": "sqrt(A_x^2+A_y^2)", "description": "", "templateType": "anything", "can_override": false}, "names": {"name": "names", "group": "Inputs", "definition": "[\"A\",\"B\",\"H\",\"C\",\"J\",\"D\",\"K\",\"E\",\"F\",\"G\"]", "description": "", "templateType": "anything", "can_override": false}, "A_y": {"name": "A_y", "group": "magnitudes", "definition": "sum(map(F[k],k,[\"C\",\"D\",\"E\",\"H\",\"J\",\"K\"]))-G", "description": "", "templateType": "anything", "can_override": false}, "loads": {"name": "loads", "group": "Inputs", "definition": "sort(shuffle(names[1..9])[0..3])", "description": "<p>only apply loads at middle 8&nbsp;(out of 10) points.</p>", "templateType": "anything", "can_override": false}, "debug": {"name": "debug", "group": "Inputs", "definition": "false", "description": "", "templateType": "anything", "can_override": false}, "A_x": {"name": "A_x", "group": "magnitudes", "definition": "f[\"B\"]-f[\"F\"]", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": [], "variable_groups": [{"name": "Inputs", "variables": ["height", "width", "units", "ggb_forces", "names", "loads", "F", "D", "version", "debug", "theta"]}, {"name": "magnitudes", "variables": ["G", "A_y", "A_x", "A", "CD", "HJ", "CJ_y", "CJ", "DE", "JK", "EJ_y", "EJ", "ans", "ans_name"]}, {"name": "quantities", "variables": []}], "functions": {"load_html": {"parameters": [["k", "number"], ["u", "string"]], "type": "html", "language": "jme", "definition": "\"<i>\" + loads[k] + \"</i> = \" + qty(f[loads[k]],u)"}, "display": {"parameters": [["q", "number"]], "type": "string", "language": "jme", "definition": "string(siground(q,4))"}, "ans": {"parameters": [["f", "number"], ["u", "string"]], "type": "string", "language": "jme", "definition": "string(siground(qty(abs(f),u),4)) + if(f<0,' (C)', ' (T)')"}}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p style=\"padding-left: 40px;\">$A_x&nbsp;=$&nbsp;[[0]]&nbsp;&nbsp;[[1]]&nbsp; <span data-jme-visible=\"debug\">{siground(qty(A_x,units[0]),4)} &nbsp; &nbsp;{map(if(sign(A_x)=s,2,0),s,[1,-1,0])}</span></p>\n<p style=\"padding-left: 40px;\">$A_y=$&nbsp;[[2]]&nbsp;&nbsp;[[3]] <span data-jme-visible=\"debug\">{siground(qty(A_y,units[0]),4)} &nbsp; &nbsp;{map(if(sign(A_y)=s,2,0),s,[1,-1,0])}</span></p>\n<p style=\"padding-left: 40px;\">$G=$&nbsp;[[4]]&nbsp;&nbsp;[[5]]&nbsp; <span data-jme-visible=\"debug\">{siground(qty(G,units[0]),4)} &nbsp; &nbsp;{map(if(sign(G)=s,2,0),s,[1,-1,0])}</span></p>\n<p style=\"padding-left: 40px;\">$\\var{latex(ans_name[0])}$ = [[6]] &nbsp;[[7]] &nbsp;<span data-jme-visible=\"debug\">{siground(qty(ans[0],units[0]),4)}</span></p>\n<p style=\"padding-left: 40px;\">$\\var{latex(ans_name[1])}$ = [[8]] &nbsp;[[9]] &nbsp;<span data-jme-visible=\"debug\">{siground(qty(ans[0],units[1]),4)}</span></p>\n<p style=\"padding-left: 40px;\">$\\var{latex(ans_name[2])}$ = [[10]] &nbsp;[[11]] &nbsp;<span data-jme-visible=\"debug\">{siground(qty(ans[0],units[2]),4)}</span></p>", "gaps": [{"type": "engineering-answer", "useCustomName": true, "customName": "A_x", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(abs(A_x),units[0])", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "choices": ["Right", "Left", "Neither"], "matrix": "map(if(sign(A_x)=s,1,0),s,[1,-1,0])"}, {"type": "engineering-answer", "useCustomName": true, "customName": "A_y", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(abs(A_y),units[0])", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "choices": ["Up", "Down", "Neither"], "matrix": "map(if(sign(A_y)=s,1,0),s,[1,-1,0])"}, {"type": "engineering-answer", "useCustomName": true, "customName": "G", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(abs(G),units[0])", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "choices": ["Up", "Down", "Neither"], "matrix": "map(if(sign(G)=s,1,0),s,[1,-1,0])"}, {"type": "engineering-answer", "useCustomName": true, "customName": "a0", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(abs(ans[0]),units[0])", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "choices": ["Tension", "Compression", "Neither"], "matrix": ["if(ans[0]>0,1,0)", "if(ans[0]<0,1,0)", "if(ans[0]=0,1,0)"], "distractors": ["", "", ""]}, {"type": "engineering-answer", "useCustomName": true, "customName": "a1", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(abs(ans[1]),units[0])", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "choices": ["Tension", "Compression", "Neither"], "matrix": ["if(ans[1]>0,1,0)", "if(ans[1]<0,1,0)", "if(ans[1]=0,1,0)"], "distractors": ["", "", ""]}, {"type": "engineering-answer", "useCustomName": true, "customName": "a2", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(abs(ans[2]),units[0])", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "choices": ["Tension", "Compression", "Neither"], "matrix": ["if(ans[2]>0,1,0)", "if(ans[2]<0,1,0)", "if(ans[2]=0,1,0)"], "distractors": ["", "", ""]}], "sortAnswers": false}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}], "resources": []}]}], "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}]}