// Numbas version: finer_feedback_settings
{"name": "Particle Equilibrium: Use slopes to find components", "extensions": ["linear-algebra", "geogebra", "quantities"], "custom_part_types": [{"source": {"pk": 19, "author": {"name": "William Haynes", "pk": 2530}, "edit_page": "/part_type/19/edit"}, "name": "Engineering Accuracy with units", "short_name": "engineering-answer", "description": "<p>A value with units marked right if within an adjustable % error of the correct value. &nbsp;Marked close if within a wider margin of error.</p>", "help_url": "", "input_widget": "string", "input_options": {"correctAnswer": "siground(settings['correctAnswer'],4)", "hint": {"static": true, "value": ""}, "allowEmpty": {"static": true, "value": true}}, "can_be_gap": true, "can_be_step": true, "marking_script": "mark:\nswitch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)\n\ninterpreted_answer:\nqty(student_scalar, student_units)\n\n\n\ncorrect_quantity:\nsettings[\"correctAnswer\"]\n\n\n\ncorrect_units:\nunits(correct_quantity)\n\n\nallowed_notation_styles:\n[\"plain\",\"en\"]\n\nmatch_student_number:\nmatchnumber(studentAnswer,allowed_notation_styles)\n\nstudent_scalar:\nmatch_student_number[1]\n\nstudent_units:\nreplace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")\n\ngood_units:\ntry(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n\n\nstudent_quantity:\nswitch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n\n\npercent_error:\ntry(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   \n\nright:\npercent_error <= settings['right']\n\n\nclose:\nright_sign and percent_error <= settings['close']\n\nright_sign:\nsign(student_scalar) = sign(correct_quantity)", "marking_notes": [{"name": "mark", "description": "This is the main marking note. It should award credit and provide feedback based on the student's answer.", "definition": "switch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)"}, {"name": "interpreted_answer", "description": "A value representing the student's answer to this part.", "definition": "qty(student_scalar, student_units)\n\n"}, {"name": "correct_quantity", "description": "", "definition": "settings[\"correctAnswer\"]\n\n"}, {"name": "correct_units", "description": "", "definition": "units(correct_quantity)\n"}, {"name": "allowed_notation_styles", "description": "", "definition": "[\"plain\",\"en\"]"}, {"name": "match_student_number", "description": "", "definition": "matchnumber(studentAnswer,allowed_notation_styles)"}, {"name": "student_scalar", "description": "", "definition": "match_student_number[1]"}, {"name": "student_units", "description": "<p>Modify&nbsp;the&nbsp;unit portion of&nbsp;the student's answer by</p>\n<p>1. replacing \"ohms\" with \"ohm\" &nbsp;case insensitive</p>\n<p>2.&nbsp;replacing '-'&nbsp;with ' '&nbsp;</p>\n<p>3. replacing '&deg;' with ' deg'&nbsp;</p>\n<p>to allow answers like 10 ft-lb and 30&deg;</p>", "definition": "replace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")"}, {"name": "good_units", "description": "", "definition": "try(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n"}, {"name": "student_quantity", "description": "<p>This fixes the student answer for two common errors. &nbsp;</p>\n<p>If student_units are wrong&nbsp;&nbsp;- replace with correct units</p>\n<p>If student_scalar has the wrong sign - replace with right sign</p>\n<p>If student makes both errors, only one gets fixed.</p>", "definition": "switch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n"}, {"name": "percent_error", "description": "", "definition": "try(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   "}, {"name": "right", "description": "", "definition": "percent_error <= settings['right']\n"}, {"name": "close", "description": "<p>Only marked close if the student actually has the right sign.</p>", "definition": "right_sign and percent_error <= settings['close']"}, {"name": "right_sign", "description": "", "definition": "sign(student_scalar) = sign(correct_quantity) "}], "settings": [{"name": "correctAnswer", "label": "Correct Quantity.", "help_url": "", "hint": "The correct answer given as a JME quantity.", "input_type": "code", "default_value": "", "evaluate": true}, {"name": "right", "label": "% Accuracy for right.", "help_url": "", "hint": "Question will be considered correct if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "0.2", "evaluate": true}, {"name": "close", "label": "% Accuracy for close.", "help_url": "", "hint": "Question will be considered close if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "1.0", "evaluate": true}, {"name": "C1", "label": "Close with units.", "help_url": "", "hint": "Partial Credit for close value with appropriate units. &nbsp;if correct answer is 100 N and close is &plusmn;1%,<br />99 &nbsp;N is accepted.", "input_type": "percent", "default_value": "75"}, {"name": "C2", "label": "No units or wrong sign", "help_url": "", "hint": "Partial credit for forgetting units or using wrong sign.<br />If the correct answer is 100 N, both 100 and -100 N are accepted.", "input_type": "percent", "default_value": "50"}, {"name": "C3", "label": "Close, no units.", "help_url": "", "hint": "Partial Credit for close value but forgotten units.<br />This value would be close&nbsp;if the expected units were provided. &nbsp;If the correct answer is 100 N, and close is &plusmn;1%,<br />99 is accepted.", "input_type": "percent", "default_value": "25"}], "public_availability": "always", "published": true, "extensions": ["quantities"]}], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Particle Equilibrium: Use slopes to find components", "tags": ["Equilibrium", "equilibrium", "Mechanics", "mechanics", "Particle", "particle", "Statics", "statics"], "metadata": {"description": "<p>Two dimensional particle equilibrium problem. &nbsp;Advice shows how to use how to use slope triangles to find sines and cosines, rather than finding the angle and using that.</p>", "licence": "Creative Commons Attribution-NonCommercial 4.0 International"}, "statement": "<p>A $\\var{qty(F,\"N\")}$ force $F$ is applied in the direction shown to a small ring, which is held in equilibrium by two inextensible ropes attached to fixed points $A$ and $B$. Determine the tensions in the two ropes.</p>\n<table style=\"margin-left: auto; margin-right: auto;\">\n<tbody>\n<tr>\n<td>{applet1(\"G\", 300)}&nbsp;</td>\n</tr>\n</tbody>\n</table>", "advice": "<p>1. &nbsp;Draw FBD of the Particle</p>\n<p>{applet1(\"D\", 250)}</p>\n<p>2. &nbsp;Determine slope triangles from geometry.</p>\n<table style=\"margin-left: 40px; text-align: center;\">\n<tbody>\n<tr style=\"border-bottom: solid;\">\n<td style=\"width: 52.989132px;\">Force</td>\n<td style=\"width: 57.989132px;\">Rise</td>\n<td style=\"width: 42.989132px;\">Run</td>\n<td style=\"width: 97.989136px;\">Hypotenuse</td>\n</tr>\n<tr>\n<td style=\"width: 52.989132px;\">$F$</td>\n<td style=\"width: 57.989132px;\">$\\var{dyF}$</td>\n<td style=\"width: 42.989132px;\">$\\var{dxF}$</td>\n<td style=\"width: 97.989136px;\">$\\sqrt{\\var{rF2}}$</td>\n</tr>\n<tr>\n<td style=\"width: 52.989132px;\">$A$</td>\n<td style=\"width: 57.989132px;\">$\\var{dyA}$</td>\n<td style=\"width: 42.989132px;\">$\\var{dxA}$</td>\n<td style=\"width: 97.989136px;\">$\\sqrt{\\var{rA2}}$</td>\n</tr>\n<tr>\n<td style=\"width: 52.989132px;\">$B$</td>\n<td style=\"width: 57.989132px;\">$\\var{dyB}$</td>\n<td style=\"width: 42.989132px;\">$\\var{dxB}$</td>\n<td style=\"width: 97.989136px;\">$\\sqrt{\\var{rB2}}$</td>\n</tr>\n</tbody>\n</table>\n<p>3. Determine scalar components in terms of $A$, $B$, and $F$ from slopes and given $F = \\var{qty(F,'N')}$</p>\n<p style=\"padding-left: 40px;\">$\\begin{align}<br/>F_x &amp;= \\left(\\frac{\\var{dxF}}{\\sqrt{\\var{rF2}}}\\right) F = \\var{precround(qty(vecF[0], 'N'),4)} &amp; F_y &amp;= \\left(\\frac{\\var{dyF}}{\\sqrt{\\var{rF2}}}\\right) F = \\var{precround(qty(vecF[1], 'N'),4)}\\\\<br/>A_x &amp;= \\left(\\frac{\\var{dxA}}{\\sqrt{\\var{rA2}}}\\right) A = \\var{precround(cos(radians(alpha)),3)} A &amp; A_y &amp;= \\left(\\frac{\\var{dyA}}{\\sqrt{\\var{rA2}}}\\right) A&nbsp;= \\var{precround(sin(radians(alpha)),3)} A\\\\<br/>B_x &amp;= \\left(\\frac{\\var{dxB}}{\\sqrt{\\var{rB2}}}\\right) B = \\var{precround(cos(radians(beta )),3)} B &amp; B_y &amp;= \\left(\\frac{\\var{dyB}}{\\sqrt{\\var{rB2}}}\\right) B =&nbsp;\\var{precround(sin(radians(beta )),3)} B<br/>\\end{align}$</p>\n<p>4. &nbsp;Set up two equations of equilibrium.</p>\n<p style=\"padding-left: 40px;\">$\\begin{align}<br/>\\Sigma F_x &amp;= 0 &amp; \\Sigma F_y &amp;= 0\\\\<br/>A_x + B_x + F_x &amp;=0&amp; A_y + B_y +F_y &amp;=0\\\\<br/>\\simplify[!noLeadingMinus]{{precround(cos(radians(alpha)),3)} A + {precround(cos(radians(beta)),3)} B} &amp; =\\var{precround(qty(- vecF[0], 'N'),4)} &amp;<br/>\\simplify[!noLeadingMinus]{{precround(sin(radians(alpha)),3)} A + {precround(sin(radians(beta)),3)} B} &amp; =\\var{precround(qty(- vecF[1], 'N'),4)}\\\\<br/>\\end{align}$</p>\n<p>5. Solve simultaneously to get:</p>\n<p style=\"padding-left: 40px;\">$\\begin{align}<br/>A &amp;= \\var{siground(qty(abs(FA),'N'),4)} &amp; B &amp;= \\var{siground(qty(abs(FB),'N'),4)}<br/>\\end{align}$</p>\n<p style=\"padding-left: 40px;\"></p>", "rulesets": {}, "extensions": ["geogebra", "linear-algebra", "quantities"], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"applet": {"name": "applet", "group": "Ungrouped variables", "definition": "geogebra_applet('saaa7mqd')", "description": "", "templateType": "anything", "can_override": false}, "A": {"name": "A", "group": "Ungrouped variables", "definition": "random(vector(random(0,10), random(0..-10)), vector(random(0..10), random(0,-10)))", "description": "<p>Grid top left at (0,0) bottom right at(10,-10)</p>", "templateType": "anything", "can_override": false}, "B": {"name": "B", "group": "Ungrouped variables", "definition": "random(vector(random(0,10), random(0..-10)), vector(random(0..10), random(0,-10)))", "description": "", "templateType": "anything", "can_override": false}, "P": {"name": "P", "group": "Ungrouped variables", "definition": "vector(random(3..7),random(-3..-7))", "description": "", "templateType": "anything", "can_override": false}, "alpha": {"name": "alpha", "group": "angles", "definition": "degrees(atan2(A[1]-P[1], A[0]-P[0]))", "description": "<p>Direction of Force A will only point in quadrant 2 or 4</p>", "templateType": "anything", "can_override": false}, "beta": {"name": "beta", "group": "angles", "definition": "degrees(atan2(B[1]-P[1], B[0]-P[0]))", "description": "<p>Direction of force B will only point in quadrant 1 or 2</p>", "templateType": "anything", "can_override": false}, "F": {"name": "F", "group": "Ungrouped variables", "definition": "random(50..600#25)", "description": "", "templateType": "anything", "can_override": false}, "FA": {"name": "FA", "group": "Ungrouped variables", "definition": "soln[0][2]", "description": "", "templateType": "anything", "can_override": false}, "FB": {"name": "FB", "group": "Ungrouped variables", "definition": "soln[1][2]", "description": "", "templateType": "anything", "can_override": false}, "C": {"name": "C", "group": "Ungrouped variables", "definition": "random(vector(random(0,10), random(0..-10)), vector(random(0..10), random(0,-10)))", "description": "", "templateType": "anything", "can_override": false}, "phi": {"name": "phi", "group": "angles", "definition": "degrees(atan2(C[1]-P[1], C[0]-P[0]))", "description": "<p>Direction of force F (will only point in quadrant 3 or 4)</p>", "templateType": "anything", "can_override": false}, "soln": {"name": "soln", "group": "Ungrouped variables", "definition": "reduced_row_echelon_form(   \nmatrix([cos(radians(alpha)),cos(radians(beta)),- F cos(radians(phi))],\n  [sin(radians(alpha)),sin(radians(beta)),- F sin(radians(phi))]))", "description": "", "templateType": "anything", "can_override": false}, "vecA": {"name": "vecA", "group": "Ungrouped variables", "definition": "vector(cos(radians(alpha)),sin(radians(alpha))) FA", "description": "", "templateType": "anything", "can_override": false}, "vecB": {"name": "vecB", "group": "Ungrouped variables", "definition": "vector(cos(radians(beta)),sin(radians(beta))) FB", "description": "", "templateType": "anything", "can_override": false}, "vecF": {"name": "vecF", "group": "Ungrouped variables", "definition": "siground(vector(cos(radians(phi)),sin(radians(phi))) F,4)", "description": "", "templateType": "anything", "can_override": false}, "Check": {"name": "Check", "group": "Ungrouped variables", "definition": "vecA + vecB + vecF", "description": "", "templateType": "anything", "can_override": false}, "dxA": {"name": "dxA", "group": "angles", "definition": "(A[0]-P[0])", "description": "<p>run for force A</p>", "templateType": "anything", "can_override": false}, "dyA": {"name": "dyA", "group": "angles", "definition": "(A[1]-P[1])", "description": "<p>Rise for force A</p>", "templateType": "anything", "can_override": false}, "dxB": {"name": "dxB", "group": "angles", "definition": "(B[0]-P[0])", "description": "", "templateType": "anything", "can_override": false}, "dyB": {"name": "dyB", "group": "angles", "definition": "(B[1]-P[1])", "description": "", "templateType": "anything", "can_override": false}, "dxF": {"name": "dxF", "group": "angles", "definition": "(C[0]-P[0])", "description": "", "templateType": "anything", "can_override": false}, "dyF": {"name": "dyF", "group": "angles", "definition": "(C[1]-P[1])", "description": "", "templateType": "anything", "can_override": false}, "rA2": {"name": "rA2", "group": "angles", "definition": "(dxA^2+dyA^2)", "description": "", "templateType": "anything", "can_override": false}, "rB2": {"name": "rB2", "group": "angles", "definition": "(dxB^2+dyB^2)", "description": "", "templateType": "anything", "can_override": false}, "rF2": {"name": "rF2", "group": "angles", "definition": "(dxF^2+dyF^2)", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "//avoid forces on same line of action\ntan(radians(alpha)) <> tan(radians(beta)) and\ntan(radians(beta)) <> tan(radians(phi)) and\ntan(radians(phi)) <> tan(radians(alpha)) and\nFA > 0 and FB > 0 and // only tension please\nFA < 50 F and FB < 50F // and not too big", "maxRuns": 100}, "ungrouped_variables": ["applet", "A", "B", "P", "C", "F", "FA", "FB", "soln", "vecA", "vecB", "vecF", "Check"], "variable_groups": [{"name": "angles", "variables": ["alpha", "beta", "phi", "dxA", "dyA", "dxB", "dyB", "dxF", "dyF", "rA2", "rB2", "rF2"]}], "functions": {"applet1": {"parameters": [["perspective", "string"], ["width", "number"]], "type": "anything", "language": "javascript", "definition": "// Create the worksheet. \n// This function returns an object with a container `element` and a `promise` resolving to a GeoGebra applet.\n\nvar params = {\n    material_id: 'saaa7mqd'\n};\nvar result = Numbas.extensions.geogebra.createGeogebraApplet(params);\n\n// Once the applet has loaded, run some commands to manipulate the worksheet.\nresult.promise.then(function(d) {\n    var app = d.app;\n    question.applet = d;\n    app.setPerspective(perspective);\n    app.setWidth(width);\n    app.setHeight(width);\n    //app.setAxesVisible(true,false);\n    //app.enableShiftDragZoom(false);\n\n    function setGGBPoint(name, nname = name) {\n        // moves point in GGB to Numbas value\n        var x = Numbas.jme.unwrapValue(question.scope.getVariable(nname));\n        app.setFixed(name, false, false); //fixed and selection allowed\n        app.setCoords(name, x[0], x[1]);\n        app.setFixed(name, true, false);\n    }\n\n    setGGBPoint('A');\n    setGGBPoint('B');\n    setGGBPoint('C');\n    setGGBPoint('P');\n    function setGGBAngle(gname, nname = gname) {\n        // Sets angle in GGB to a Numbas Variable given in degrees.\n        var v = Math.PI / 180 * Numbas.jme.unwrapValue(question.scope.getVariable(nname));\n        app.setValue(gname, v);\n    }\n\n    //setGGBAngle('\u03b8','theta');\n    //setGGBAngle('\u03b1','alpha');\n    //app.setValue('tension',Numbas.jme.unwrapValue(question.scope.getVariable('tension')));\n\n});\n\n// This function returns the result of `createGeogebraApplet` as an object \n// with the JME data type 'ggbapplet', which can be substituted into the question's content.\nreturn new Numbas.jme.types.ggbapplet(result);\n"}}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p style=\"padding-left: 40px;\">$A = $ [[0]]</p>\n<p style=\"padding-left: 40px;\">$B = $ [[1]]</p>", "gaps": [{"type": "engineering-answer", "useCustomName": true, "customName": "A", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "siground(qty(FA,\"N\"),3)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "B", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "siground(qty(FB,\"N\"),3)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}], "sortAnswers": false}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "Chris Graham", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/369/"}, {"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}], "resources": []}]}], "contributors": [{"name": "Chris Graham", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/369/"}, {"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}]}