// Numbas version: finer_feedback_settings
{"name": "Method of joints: cantilever truss", "extensions": ["quantities"], "custom_part_types": [{"source": {"pk": 19, "author": {"name": "William Haynes", "pk": 2530}, "edit_page": "/part_type/19/edit"}, "name": "Engineering Accuracy with units", "short_name": "engineering-answer", "description": "<p>A value with units marked right if within an adjustable % error of the correct value. &nbsp;Marked close if within a wider margin of error.</p>", "help_url": "", "input_widget": "string", "input_options": {"correctAnswer": "siground(settings['correctAnswer'],4)", "hint": {"static": true, "value": ""}, "allowEmpty": {"static": true, "value": true}}, "can_be_gap": true, "can_be_step": true, "marking_script": "mark:\nswitch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)\n\ninterpreted_answer:\nqty(student_scalar, student_units)\n\n\n\ncorrect_quantity:\nsettings[\"correctAnswer\"]\n\n\n\ncorrect_units:\nunits(correct_quantity)\n\n\nallowed_notation_styles:\n[\"plain\",\"en\"]\n\nmatch_student_number:\nmatchnumber(studentAnswer,allowed_notation_styles)\n\nstudent_scalar:\nmatch_student_number[1]\n\nstudent_units:\nreplace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")\n\ngood_units:\ntry(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n\n\nstudent_quantity:\nswitch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n\n\npercent_error:\ntry(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   \n\nright:\npercent_error <= settings['right']\n\n\nclose:\nright_sign and percent_error <= settings['close']\n\nright_sign:\nsign(student_scalar) = sign(correct_quantity)", "marking_notes": [{"name": "mark", "description": "This is the main marking note. It should award credit and provide feedback based on the student's answer.", "definition": "switch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)"}, {"name": "interpreted_answer", "description": "A value representing the student's answer to this part.", "definition": "qty(student_scalar, student_units)\n\n"}, {"name": "correct_quantity", "description": "", "definition": "settings[\"correctAnswer\"]\n\n"}, {"name": "correct_units", "description": "", "definition": "units(correct_quantity)\n"}, {"name": "allowed_notation_styles", "description": "", "definition": "[\"plain\",\"en\"]"}, {"name": "match_student_number", "description": "", "definition": "matchnumber(studentAnswer,allowed_notation_styles)"}, {"name": "student_scalar", "description": "", "definition": "match_student_number[1]"}, {"name": "student_units", "description": "<p>Modify&nbsp;the&nbsp;unit portion of&nbsp;the student's answer by</p>\n<p>1. replacing \"ohms\" with \"ohm\" &nbsp;case insensitive</p>\n<p>2.&nbsp;replacing '-'&nbsp;with ' '&nbsp;</p>\n<p>3. replacing '&deg;' with ' deg'&nbsp;</p>\n<p>to allow answers like 10 ft-lb and 30&deg;</p>", "definition": "replace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")"}, {"name": "good_units", "description": "", "definition": "try(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n"}, {"name": "student_quantity", "description": "<p>This fixes the student answer for two common errors. &nbsp;</p>\n<p>If student_units are wrong&nbsp;&nbsp;- replace with correct units</p>\n<p>If student_scalar has the wrong sign - replace with right sign</p>\n<p>If student makes both errors, only one gets fixed.</p>", "definition": "switch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n"}, {"name": "percent_error", "description": "", "definition": "try(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   "}, {"name": "right", "description": "", "definition": "percent_error <= settings['right']\n"}, {"name": "close", "description": "<p>Only marked close if the student actually has the right sign.</p>", "definition": "right_sign and percent_error <= settings['close']"}, {"name": "right_sign", "description": "", "definition": "sign(student_scalar) = sign(correct_quantity) "}], "settings": [{"name": "correctAnswer", "label": "Correct Quantity.", "help_url": "", "hint": "The correct answer given as a JME quantity.", "input_type": "code", "default_value": "", "evaluate": true}, {"name": "right", "label": "% Accuracy for right.", "help_url": "", "hint": "Question will be considered correct if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "0.2", "evaluate": true}, {"name": "close", "label": "% Accuracy for close.", "help_url": "", "hint": "Question will be considered close if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "1.0", "evaluate": true}, {"name": "C1", "label": "Close with units.", "help_url": "", "hint": "Partial Credit for close value with appropriate units. &nbsp;if correct answer is 100 N and close is &plusmn;1%,<br />99 &nbsp;N is accepted.", "input_type": "percent", "default_value": "75"}, {"name": "C2", "label": "No units or wrong sign", "help_url": "", "hint": "Partial credit for forgetting units or using wrong sign.<br />If the correct answer is 100 N, both 100 and -100 N are accepted.", "input_type": "percent", "default_value": "50"}, {"name": "C3", "label": "Close, no units.", "help_url": "", "hint": "Partial Credit for close value but forgotten units.<br />This value would be close&nbsp;if the expected units were provided. &nbsp;If the correct answer is 100 N, and close is &plusmn;1%,<br />99 is accepted.", "input_type": "percent", "default_value": "25"}], "public_availability": "always", "published": true, "extensions": ["quantities"]}], "resources": [["question-resources/truss_KUTOuc7.png", "truss_KUTOuc7.png"], ["question-resources/truss.vwx", "truss.vwx"], ["question-resources/jointA_i2nJN9c.png", "jointA_i2nJN9c.png"], ["question-resources/jointB_N4t3sYw.png", "jointB_N4t3sYw.png"], ["question-resources/jointC_5MFLZd4.png", "jointC_5MFLZd4.png"]], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Method of joints: cantilever truss", "tags": ["Equilibrium", "equilibrium", "Mechanics", "mechanics", "method of joints", "particle equilibrium", "Statics", "statics", "truss"], "metadata": {"description": "Use the method of joints to solve for the forces in a cantilever truss.", "licence": "Creative Commons Attribution-NonCommercial 4.0 International"}, "statement": "<p>The truss shown supports three loads, $A = \\var{A}$, $B=\\var{B}$, and $C=\\var{C}$.</p>\n<table style=\"margin-left: auto; margin-right: auto;\">\n<tbody>\n<tr>\n<td><img src=\"resources/question-resources/truss_KUTOuc7.png\" width=\"325\" height=\"271\" style=\"display: block; margin-left: auto; margin-right: auto;\"/></td>\n</tr>\n</tbody>\n</table>\n<p>&nbsp;Use the method of joints to find the forces in all members. &nbsp;Indicate tension or compression.</p>", "advice": "<p><strong>Procedure</strong></p>\n<p>Start with a joint with two unknowns, in this case $A$. &nbsp;Draw a free body diagram of the joint and solve for the two unknowns. Move to an adjacent joint with two unknowns and repeat. &nbsp;Continue in this manner until all the unknowns are found.</p>\n<p>When drawing free body diagrams, the force of&nbsp;a member acts along the axis of the member. &nbsp;A force pushing towards the joint indicates compression, and a force pulling away indicates tension. &nbsp;An incorrect assumed direction will result in a negative value.</p>\n<p>After all the forces in the members are found&nbsp;you can check your answer by calculating the&nbsp;external&nbsp;reactions at $D$ and $E$, and verifying that &nbsp;joints $D$ and $E$ are in equilibrium.</p>\n<p></p>\n<p><img src=\"resources/question-resources/jointA_i2nJN9c.png\" width=\"200px/\"/></p>\n<p>$\\begin{align}A\\!: \\Sigma F_y &amp;= 0 &amp; &nbsp;A\\!: \\Sigma F_x &amp;= 0 &nbsp; \\\\AB_y&nbsp;&amp;= A&amp; &nbsp; &nbsp; &nbsp;AC&amp;= AB_x &nbsp; &nbsp;\\\\AB \\sin 30&deg; &amp;= A &nbsp;&amp; &nbsp; &nbsp;&amp;= AB \\cos 30&deg;&nbsp; &nbsp; &nbsp; &nbsp; \\\\AB &amp;= \\dfrac{\\var{disp(A)}}{\\sin 30&deg;} &amp; &nbsp; &amp;= \\var{disp(AB)}&nbsp;\\cos 30&deg;&nbsp; &nbsp; &nbsp; &nbsp; &nbsp; &nbsp; &nbsp; &nbsp; &nbsp; \\\\&amp;= \\var{siground(AB,4)} &nbsp;\\text{ (T)} &amp; &amp;= \\var{siground(AC,4)} &nbsp;\\text{ (C)}\\end{align}$</p>\n<p></p>\n<p></p>\n<p><img src=\"resources/question-resources/jointB_N4t3sYw.png\" width=\"200px\"/></p>\n<p>$\\begin{align}B\\!: \\Sigma F_{y'} &amp;= 0&amp;&nbsp;B\\!: \\Sigma F_{x'}&amp;= 0\\\\BC &amp;= B_{y'}&amp;BD&amp;= B_{x'} + AB&nbsp;\\\\ &amp;= B&nbsp;\\cos 30&deg; &amp;&nbsp;&amp;= B \\sin 30&deg;+ AB\\\\ &amp;= {\\var{disp(B)}} \\cos&nbsp;30&deg;&amp;&amp;= \\var{disp(B)}&nbsp;\\sin 30&deg;+ {AB}\\\\&amp;= \\var{siground(BC,4)} &nbsp;\\text{ (C)}&amp; &amp;= \\var{siground(BD,4)} &nbsp;\\text{ (T)}\\ \\\\\\end{align}$</p>\n<p></p>\n<p></p>\n<p><img src=\"resources/question-resources/jointC_5MFLZd4.png\" width=\"200px/\"/></p>\n<p>$\\begin{align}C\\!: \\Sigma {F_y} &amp;= 0 &amp; &nbsp; &nbsp; &nbsp; &nbsp;&nbsp;C\\!: \\Sigma {F_x} &amp;= 0\\\\ &nbsp; &nbsp; &nbsp; &nbsp; &nbsp; &nbsp;-CD_y-BC_{y} -C&nbsp;&amp;=0 &nbsp; &nbsp; &nbsp;&amp; &nbsp; CE + CD_x &amp;= BC_x +&nbsp;AC\\\\ &nbsp; &nbsp;CD \\sin 60&deg;&amp;= -(BC \\sin&nbsp;60&deg; +C)&amp; &nbsp; &nbsp;CE&nbsp;&amp;= AC + BC_x - CD_x\\\\CD&amp;= -\\frac{{\\var{disp(BC)}} \\sin 60&deg; + \\var{C}}{\\sin 60&deg;} &amp;&nbsp;&amp;= \\var{disp(AC)} + \\var{disp(BC)} \\cos 60&deg; -(\\var{disp(-CD)}) \\cos 60&deg; \\\\ &nbsp;&amp;= \\var{-disp(CD)} &nbsp;\\text{ (C)} &nbsp;&amp; &amp;= \\var{siground(CE,4)}&nbsp; \\text{ (C)} \\\\&amp;= \\var{siground(CD,4)} &nbsp;\\text{ (T)} &nbsp;\\end{align}$</p>", "rulesets": {}, "extensions": ["quantities"], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"C": {"name": "C", "group": "Ungrouped variables", "definition": "qty(random(250..1000#50),units)", "description": "", "templateType": "anything", "can_override": false}, "CD": {"name": "CD", "group": "Ungrouped variables", "definition": "(BC sin(radians(60)) + C )/sin(radians(60))", "description": "", "templateType": "anything", "can_override": false}, "BC": {"name": "BC", "group": "Ungrouped variables", "definition": "B cos(radians(30))", "description": "", "templateType": "anything", "can_override": false}, "check": {"name": "check", "group": "Ungrouped variables", "definition": "(3 A + 1.5 B + C)/sqrt(3) - ce", "description": "", "templateType": "anything", "can_override": false}, "AC": {"name": "AC", "group": "Ungrouped variables", "definition": "AB cos(radians(30))", "description": "", "templateType": "anything", "can_override": false}, "AB": {"name": "AB", "group": "Ungrouped variables", "definition": "A/sin(radians(30))", "description": "", "templateType": "anything", "can_override": false}, "B": {"name": "B", "group": "Ungrouped variables", "definition": "qty(random(250..1000#50),units)", "description": "", "templateType": "anything", "can_override": false}, "A": {"name": "A", "group": "Ungrouped variables", "definition": "qty(random(250..1000#50),units)", "description": "", "templateType": "anything", "can_override": false}, "BD": {"name": "BD", "group": "Ungrouped variables", "definition": "B sin(radians(30)) + AB", "description": "", "templateType": "anything", "can_override": false}, "units": {"name": "units", "group": "Ungrouped variables", "definition": "random(['N', 'kN', 'lb'])", "description": "", "templateType": "anything", "can_override": false}, "CE": {"name": "CE", "group": "Ungrouped variables", "definition": "ac + bc cos(radians(60)) + cd cos(radians(60))", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["units", "A", "B", "C", "AB", "AC", "BC", "BD", "CD", "CE", "check"], "variable_groups": [], "functions": {"disp": {"parameters": [["q", "quantity"]], "type": "string", "language": "jme", "definition": "siground(scalar(q),4)"}}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>1. Analyze joint $A$ and find $AB$ and $AC$.</p>\n<p style=\"padding-left: 40px;\">$AB =$ &nbsp;[[0]]&nbsp;[[1]]&nbsp;</p>\n<p style=\"padding-left: 40px;\">$AC =$&nbsp;[[2]]&nbsp;[[3]]</p>\n<p>2. Analyze joint $B$ and find $BC$ and $BD$.</p>\n<p style=\"padding-left: 40px;\">$BC =$ [[4]]&nbsp;[[5]]</p>\n<p style=\"padding-left: 40px;\">$BD =$[[6]]&nbsp;[[7]]</p>\n<p>3. Analyze joint $C$ and find $CD$ and $CE$.</p>\n<p style=\"padding-left: 40px;\">$CD = $ [[8]]&nbsp;[[9]]</p>\n<p style=\"padding-left: 40px;\">$CE =$ [[10]]&nbsp;[[11]]</p>", "gaps": [{"type": "engineering-answer", "useCustomName": true, "customName": "AB", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "AB", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "choices": ["Tension", "Compression"], "matrix": ["1", 0], "distractors": ["", ""]}, {"type": "engineering-answer", "useCustomName": true, "customName": "AC", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "AC", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "choices": ["Tension", "Compression"], "matrix": ["0", "1"], "distractors": ["", ""]}, {"type": "engineering-answer", "useCustomName": true, "customName": "BC", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "BC", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "choices": ["Tension", "Compression"], "matrix": ["0", "1"], "distractors": ["", ""]}, {"type": "engineering-answer", "useCustomName": true, "customName": "BD", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "BD", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "choices": ["Tension", "Compression"], "matrix": ["1", "0"], "distractors": ["", ""]}, {"type": "engineering-answer", "useCustomName": true, "customName": "CD", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "CD", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "choices": ["Tension", "Compression"], "matrix": ["1", "0"], "distractors": ["", ""]}, {"type": "engineering-answer", "useCustomName": true, "customName": "CE", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "CE", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "choices": ["Tension", "Compression"], "matrix": ["0", "1"], "distractors": ["", ""]}], "sortAnswers": false}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}], "resources": ["question-resources/truss_KUTOuc7.png", "question-resources/truss.vwx", "question-resources/jointA_i2nJN9c.png", "question-resources/jointB_N4t3sYw.png", "question-resources/jointC_5MFLZd4.png"]}]}], "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}]}