// Numbas version: finer_feedback_settings
{"name": "Moment of inertia: 2 channels and 2 plates", "extensions": ["geogebra", "quantities"], "custom_part_types": [{"source": {"pk": 19, "author": {"name": "William Haynes", "pk": 2530}, "edit_page": "/part_type/19/edit"}, "name": "Engineering Accuracy with units", "short_name": "engineering-answer", "description": "<p>A value with units marked right if within an adjustable % error of the correct value. &nbsp;Marked close if within a wider margin of error.</p>", "help_url": "", "input_widget": "string", "input_options": {"correctAnswer": "siground(settings['correctAnswer'],4)", "hint": {"static": true, "value": ""}, "allowEmpty": {"static": true, "value": true}}, "can_be_gap": true, "can_be_step": true, "marking_script": "mark:\nswitch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)\n\ninterpreted_answer:\nqty(student_scalar, student_units)\n\n\n\ncorrect_quantity:\nsettings[\"correctAnswer\"]\n\n\n\ncorrect_units:\nunits(correct_quantity)\n\n\nallowed_notation_styles:\n[\"plain\",\"en\"]\n\nmatch_student_number:\nmatchnumber(studentAnswer,allowed_notation_styles)\n\nstudent_scalar:\nmatch_student_number[1]\n\nstudent_units:\nreplace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")\n\ngood_units:\ntry(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n\n\nstudent_quantity:\nswitch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n\n\npercent_error:\ntry(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   \n\nright:\npercent_error <= settings['right']\n\n\nclose:\nright_sign and percent_error <= settings['close']\n\nright_sign:\nsign(student_scalar) = sign(correct_quantity)", "marking_notes": [{"name": "mark", "description": "This is the main marking note. It should award credit and provide feedback based on the student's answer.", "definition": "switch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)"}, {"name": "interpreted_answer", "description": "A value representing the student's answer to this part.", "definition": "qty(student_scalar, student_units)\n\n"}, {"name": "correct_quantity", "description": "", "definition": "settings[\"correctAnswer\"]\n\n"}, {"name": "correct_units", "description": "", "definition": "units(correct_quantity)\n"}, {"name": "allowed_notation_styles", "description": "", "definition": "[\"plain\",\"en\"]"}, {"name": "match_student_number", "description": "", "definition": "matchnumber(studentAnswer,allowed_notation_styles)"}, {"name": "student_scalar", "description": "", "definition": "match_student_number[1]"}, {"name": "student_units", "description": "<p>Modify&nbsp;the&nbsp;unit portion of&nbsp;the student's answer by</p>\n<p>1. replacing \"ohms\" with \"ohm\" &nbsp;case insensitive</p>\n<p>2.&nbsp;replacing '-'&nbsp;with ' '&nbsp;</p>\n<p>3. replacing '&deg;' with ' deg'&nbsp;</p>\n<p>to allow answers like 10 ft-lb and 30&deg;</p>", "definition": "replace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")"}, {"name": "good_units", "description": "", "definition": "try(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n"}, {"name": "student_quantity", "description": "<p>This fixes the student answer for two common errors. &nbsp;</p>\n<p>If student_units are wrong&nbsp;&nbsp;- replace with correct units</p>\n<p>If student_scalar has the wrong sign - replace with right sign</p>\n<p>If student makes both errors, only one gets fixed.</p>", "definition": "switch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n"}, {"name": "percent_error", "description": "", "definition": "try(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   "}, {"name": "right", "description": "", "definition": "percent_error <= settings['right']\n"}, {"name": "close", "description": "<p>Only marked close if the student actually has the right sign.</p>", "definition": "right_sign and percent_error <= settings['close']"}, {"name": "right_sign", "description": "", "definition": "sign(student_scalar) = sign(correct_quantity) "}], "settings": [{"name": "correctAnswer", "label": "Correct Quantity.", "help_url": "", "hint": "The correct answer given as a JME quantity.", "input_type": "code", "default_value": "", "evaluate": true}, {"name": "right", "label": "% Accuracy for right.", "help_url": "", "hint": "Question will be considered correct if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "0.2", "evaluate": true}, {"name": "close", "label": "% Accuracy for close.", "help_url": "", "hint": "Question will be considered close if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "1.0", "evaluate": true}, {"name": "C1", "label": "Close with units.", "help_url": "", "hint": "Partial Credit for close value with appropriate units. &nbsp;if correct answer is 100 N and close is &plusmn;1%,<br />99 &nbsp;N is accepted.", "input_type": "percent", "default_value": "75"}, {"name": "C2", "label": "No units or wrong sign", "help_url": "", "hint": "Partial credit for forgetting units or using wrong sign.<br />If the correct answer is 100 N, both 100 and -100 N are accepted.", "input_type": "percent", "default_value": "50"}, {"name": "C3", "label": "Close, no units.", "help_url": "", "hint": "Partial Credit for close value but forgotten units.<br />This value would be close&nbsp;if the expected units were provided. &nbsp;If the correct answer is 100 N, and close is &plusmn;1%,<br />99 is accepted.", "input_type": "percent", "default_value": "25"}], "public_availability": "always", "published": true, "extensions": ["quantities"]}], "resources": [["question-resources/channel_EfJjxiR.png", "channel_EfJjxiR.png"], ["question-resources/builtupbeam.ggb", "builtupbeam.ggb"]], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Moment of inertia: 2 channels and 2 plates", "tags": ["built-up beam", "Mechanics", "mechanics", "moment of inertia", "radius of gyration", "standard sections", "Statics", "statics"], "metadata": {"description": "<p>Find moment of inertia and radius of gyration for a built-up beam made of two channels and two plates.</p>", "licence": "Creative Commons Attribution-NonCommercial 4.0 International"}, "statement": "<p>A built-up beam is constructed by welding two channels separated by a distance $b = \\var{b}$ to two ({w} $\\times$ {h}) plates.</p>\n<table height=\"34\" style=\"width: 548px; margin-left: auto; margin-right: auto;\">\n<tbody>\n<tr>\n<td style=\"width: 223.390625px; text-align: left;\">{geogebra_applet('rgft4vzy',[\"b\":ggb(b),'bf': ggb(bf),'w': ggb(w) ,'d': ggb(d),'t': ggb(t),'h': ggb(h)])}</td>\n<td style=\"width: 501.609375px; text-align: left;\">\n<div style=\"float: right;\">\n<p><strong>Channel properties:</strong></p>\n<p>$\\begin{align}\\var{desc}\\\\\\text{Area}&amp;= \\var{A_C} \\\\D &amp;= \\var{d} \\\\b_f &amp;= \\var{bf}\\\\\\bar{x} &amp;= \\var{xbar}\\\\\\bar{I}_{xx} &amp;= \\var{Ibarx_C} \\\\\\bar{I}_{yy} &amp;= \\var{Ibary_C}\\end{align}$</p>\n</div>\n</td>\n</tr>\n</tbody>\n</table>\n<p>Determine the moment of inertia and radii of gyration of the composite beam with respect to the $x$- and $y$-axes.</p>\n<p>Note: $D$ is the depth of the channel, $b_f$ is the width of the flange.</p>\n<p></p>", "advice": "<p><strong></strong><strong>Determine moment of inertia with respect to the <em>x</em>-axis</strong></p>\n<p style=\"padding-left: 30px;\"><strong></strong>For the channel, since the&nbsp;x-axis passes through its centroid:</p>\n<p style=\"padding-left: 60px;\">$(I_x)_C = \\bar{I}_{xx} =&nbsp;\\var{Ibarx_C} $</p>\n<p style=\"padding-left: 30px;\">For the plate, the <em>x</em>-axis does not pass through its&nbsp;centroid&nbsp;so&nbsp;the parallel axis theorem must be used.</p>\n<p style=\"padding-left: 60px;\">$(I_x)_R = [\\bar{I} + A d^2] $ where,</p>\n<p style=\"padding-left: 90px;\">$\\bar{I} = \\dfrac{b h^3}{12} = \\dfrac{(\\var{w})(\\var{h})^3}{12} = \\var{siground(Ibarx_R,4)}$</p>\n<p style=\"padding-left: 90px;\">$A &nbsp;= b h = (\\var{w})\\,(\\var{h}) = \\var{A_r}$</p>\n<p style=\"padding-left: 90px;\">$d = D/2 + h/2 = \\dfrac{\\var{d}+ \\var{h}}{2} = \\var{d_y} $</p>\n<p style=\"padding-left: 60px;\">$(I_x)_R = \\var{disp(Ix_R)}$</p>\n<p style=\"padding-left: 30px;\">For the composite shape,</p>\n<p style=\"padding-left: 60px;\">$A = 2 [A_C + A_R] = \\var{disp(A_T)}$</p>\n<p style=\"padding-left: 60px;\">$I_x = 2 [ (I_x)_C + (I_x)_R ] = \\var{disp(Ix)}$, and &nbsp;$k_x = \\sqrt{\\dfrac{I_x}{A}} = \\var{disp(kx)}$</p>\n<p style=\"padding-left: 90px;\"></p>\n<p><strong>Determine moment of inertia with respect to the <em>y-</em>axis</strong></p>\n<p style=\"padding-left: 30px;\">For the rectangle, since the y-axis passes through its centroid:</p>\n<p style=\"padding-left: 60px;\">$(I_y)_R = \\dfrac{h b^3}{12}&nbsp;= \\dfrac{(\\var{h}) (\\var{w})^3}{12} = \\var{disp(Ibary_r)}$</p>\n<p style=\"padding-left: 30px;\">For the channel, the y-axis does not pass through its centroid so the parallel axis theorem must be used.</p>\n<p style=\"padding-left: 60px;\">$(I_y)_C = [\\bar{I} + A d^2] $ where,</p>\n<p style=\"padding-left: 90px;\">$\\bar{I} = \\bar{I}_{yy} = \\var{Ibary_C}$</p>\n<p style=\"padding-left: 90px;\">$A = A_C&nbsp;= \\var{A_C}$</p>\n<p style=\"padding-left: 90px;\">$d = b/2 + \\bar{x}&nbsp;= \\dfrac{\\var{b}}{2} + \\var{xbar}= \\var{d_x} $</p>\n<p style=\"padding-left: 90px;\">$(I_y)_C = \\var{disp(Iy_C)}$</p>\n<p style=\"padding-left: 30px;\">For the composite shape,</p>\n<p style=\"padding-left: 60px;\">$A = 2 [A_C + A_R] = \\var{disp(A_T)}$</p>\n<p style=\"padding-left: 60px;\">$I_y = 2 [ (I_y)_C + (I_y)_R ] = \\var{disp(Iy)}$, and&nbsp;$k_y = \\sqrt{\\dfrac{I_y}{A}} = \\var{disp(ky)}$</p>\n<p style=\"padding-left: 60px;\"><strong>&nbsp;</strong></p>\n<p style=\"padding-left: 60px;\"></p>", "rulesets": {}, "extensions": ["geogebra", "quantities"], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"xbar": {"name": "xbar", "group": "Channel Properties", "definition": "unit1(C_data['x'][unit])", "description": "<p>horizontal distance to centroid from flange</p>", "templateType": "anything", "can_override": false}, "A_t": {"name": "A_t", "group": "Solution", "definition": "2(A_r+A_C)", "description": "", "templateType": "anything", "can_override": false}, "Ibarx_c": {"name": "Ibarx_c", "group": "Channel Properties", "definition": "unit4(C_data['Ix'][unit])", "description": "", "templateType": "anything", "can_override": false}, "ky": {"name": "ky", "group": "Solution", "definition": "unit1(sqrt(scalar(Iy/A_t)))", "description": "", "templateType": "anything", "can_override": false}, "d_y": {"name": "d_y", "group": "Solution", "definition": "d/2 + h/2", "description": "", "templateType": "anything", "can_override": false}, "index": {"name": "index", "group": "Ungrouped variables", "definition": "random(0..length(data)-1)", "description": "<p>random channel index</p>", "templateType": "anything", "can_override": false}, "A_r": {"name": "A_r", "group": "Solution", "definition": "w h", "description": "", "templateType": "anything", "can_override": false}, "C_data": {"name": "C_data", "group": "Channel Properties", "definition": "data[index]", "description": "<p>hash of data for selected channel</p>", "templateType": "anything", "can_override": false}, "kx": {"name": "kx", "group": "Solution", "definition": "unit1(sqrt(scalar(Ix/A_t)))", "description": "", "templateType": "anything", "can_override": false}, "A_C": {"name": "A_C", "group": "Channel Properties", "definition": "unit2(C_data['A'][unit])", "description": "<p>area of channel</p>", "templateType": "anything", "can_override": false}, "unit": {"name": "unit", "group": "Ungrouped variables", "definition": "random(0,1)", "description": "<p>0 = inches, 1 = mm</p>", "templateType": "anything", "can_override": false}, "Iy_C": {"name": "Iy_C", "group": "Solution", "definition": "Ibary_C + A_c d_x d_x", "description": "", "templateType": "anything", "can_override": false}, "Ix": {"name": "Ix", "group": "Solution", "definition": "2 (Ibarx_c + Ix_r)", "description": "", "templateType": "anything", "can_override": false}, "bf": {"name": "bf", "group": "Channel Properties", "definition": "unit1(C_data['bf'][unit])", "description": "<p>flange width</p>", "templateType": "anything", "can_override": false}, "Iy": {"name": "Iy", "group": "Solution", "definition": "2 (Iy_C + Ibary_R)", "description": "", "templateType": "anything", "can_override": false}, "w": {"name": "w", "group": "Ungrouped variables", "definition": "round ((b/2 + bf + t ) 2 , if(unit=0,'1 in', '10 mm'))", "description": "<p>width&nbsp;of plate &nbsp;a little longer than necessary</p>", "templateType": "anything", "can_override": false}, "t": {"name": "t", "group": "Channel Properties", "definition": "unit1(C_data['tf'][unit])", "description": "", "templateType": "anything", "can_override": false}, "d_x": {"name": "d_x", "group": "Solution", "definition": "b/2 + xbar", "description": "", "templateType": "anything", "can_override": false}, "d": {"name": "d", "group": "Channel Properties", "definition": "unit1(C_data['d'][unit])", "description": "<p>depth of channel</p>", "templateType": "anything", "can_override": false}, "Ibary_c": {"name": "Ibary_c", "group": "Channel Properties", "definition": "unit4(C_data['Iy'][unit])", "description": "<p>centroidal momentof inertia</p>", "templateType": "anything", "can_override": false}, "Desc": {"name": "Desc", "group": "Channel Properties", "definition": "C_data['DESC'][unit] ", "description": "<p>designation</p>", "templateType": "anything", "can_override": false}, "Ix_R": {"name": "Ix_R", "group": "Solution", "definition": "Ibarx_r + A_r d_y d_y", "description": "", "templateType": "anything", "can_override": false}, "h": {"name": "h", "group": "Ungrouped variables", "definition": "round(random(1..3#0.1)t,if(unit=0,'0.125 in', '5 mm'))", "description": "<p>width of plate</p>", "templateType": "anything", "can_override": false}, "Ibarx_r": {"name": "Ibarx_r", "group": "Solution", "definition": "w h h h /12", "description": "", "templateType": "anything", "can_override": false}, "units": {"name": "units", "group": "Ungrouped variables", "definition": "['in','mm'][unit]", "description": "", "templateType": "anything", "can_override": false}, "Ibary_R": {"name": "Ibary_R", "group": "Solution", "definition": "h w w w /12", "description": "", "templateType": "anything", "can_override": false}, "data": {"name": "data", "group": "Channel Properties", "definition": "json_decode(safe(\"[\\n  {\\n    \\\"DESC\\\": [\\n      \\\"C15$\\\\\\\\times$50\\\",\\n      \\\"C380$\\\\\\\\times$74\\\"\\n    ],\\n    \\\"W\\\": [\\n      50,\\n      74\\n    ],\\n    \\\"A\\\": [\\n      14.7,\\n      9480\\n    ],\\n    \\\"d\\\": [\\n      15,\\n      381\\n    ],\\n    \\\"bf\\\": [\\n      3.72,\\n      94.5\\n    ],\\n    \\\"tw\\\": [\\n      0.716,\\n      18.2\\n    ],\\n    \\\"tf\\\": [\\n      0.65,\\n      16.5\\n    ],\\n    \\\"x\\\": [\\n      0.799,\\n      20.3\\n    ],\\n    \\\"Ix\\\": [\\n      404,\\n      168\\n    ],\\n    \\\"Iy\\\": [\\n      11,\\n      4.58\\n    ]\\n  },\\n  {\\n    \\\"DESC\\\": [\\n      \\\"C15$\\\\\\\\times$40\\\",\\n      \\\"C380$\\\\\\\\times$60\\\"\\n    ],\\n    \\\"W\\\": [\\n      40,\\n      60\\n    ],\\n    \\\"A\\\": [\\n      11.8,\\n      7610\\n    ],\\n    \\\"d\\\": [\\n      15,\\n      381\\n    ],\\n    \\\"bf\\\": [\\n      3.52,\\n      89.4\\n    ],\\n    \\\"tw\\\": [\\n      0.52,\\n      13.2\\n    ],\\n    \\\"tf\\\": [\\n      0.65,\\n      16.5\\n    ],\\n    \\\"x\\\": [\\n      0.778,\\n      19.8\\n    ],\\n    \\\"Ix\\\": [\\n      348,\\n      145\\n    ],\\n    \\\"Iy\\\": [\\n      9.17,\\n      3.82\\n    ]\\n  },\\n  {\\n    \\\"DESC\\\": [\\n      \\\"C15$\\\\\\\\times$33.9\\\",\\n      \\\"C380$\\\\\\\\times$50.4\\\"\\n    ],\\n    \\\"W\\\": [\\n      33.9,\\n      50.4\\n    ],\\n    \\\"A\\\": [\\n      10,\\n      6450\\n    ],\\n    \\\"d\\\": [\\n      15,\\n      381\\n    ],\\n    \\\"bf\\\": [\\n      3.4,\\n      86.4\\n    ],\\n    \\\"tw\\\": [\\n      0.4,\\n      10.2\\n    ],\\n    \\\"tf\\\": [\\n      0.65,\\n      16.5\\n    ],\\n    \\\"x\\\": [\\n      0.788,\\n      20\\n    ],\\n    \\\"Ix\\\": [\\n      315,\\n      131\\n    ],\\n    \\\"Iy\\\": [\\n      8.07,\\n      3.36\\n    ]\\n  },\\n  {\\n    \\\"DESC\\\": [\\n      \\\"C12$\\\\\\\\times$30\\\",\\n      \\\"C310$\\\\\\\\times$45\\\"\\n    ],\\n    \\\"W\\\": [\\n      30,\\n      45\\n    ],\\n    \\\"A\\\": [\\n      8.81,\\n      5680\\n    ],\\n    \\\"d\\\": [\\n      12,\\n      305\\n    ],\\n    \\\"bf\\\": [\\n      3.17,\\n      80.5\\n    ],\\n    \\\"tw\\\": [\\n      0.51,\\n      13\\n    ],\\n    \\\"tf\\\": [\\n      0.501,\\n      12.7\\n    ],\\n    \\\"x\\\": [\\n      0.674,\\n      17.1\\n    ],\\n    \\\"Ix\\\": [\\n      162,\\n      67.4\\n    ],\\n    \\\"Iy\\\": [\\n      5.12,\\n      2.13\\n    ]\\n  },\\n  {\\n    \\\"DESC\\\": [\\n      \\\"C12$\\\\\\\\times$25\\\",\\n      \\\"C310$\\\\\\\\times$37\\\"\\n    ],\\n    \\\"W\\\": [\\n      25,\\n      37\\n    ],\\n    \\\"A\\\": [\\n      7.34,\\n      4740\\n    ],\\n    \\\"d\\\": [\\n      12,\\n      305\\n    ],\\n    \\\"bf\\\": [\\n      3.05,\\n      77.5\\n    ],\\n    \\\"tw\\\": [\\n      0.387,\\n      9.83\\n    ],\\n    \\\"tf\\\": [\\n      0.501,\\n      12.7\\n    ],\\n    \\\"x\\\": [\\n      0.674,\\n      17.1\\n    ],\\n    \\\"Ix\\\": [\\n      144,\\n      59.9\\n    ],\\n    \\\"Iy\\\": [\\n      4.45,\\n      1.85\\n    ]\\n  },\\n  {\\n    \\\"DESC\\\": [\\n      \\\"C12$\\\\\\\\times$20.7\\\",\\n      \\\"C310$\\\\\\\\times$30.8\\\"\\n    ],\\n    \\\"W\\\": [\\n      20.7,\\n      30.8\\n    ],\\n    \\\"A\\\": [\\n      6.08,\\n      3920\\n    ],\\n    \\\"d\\\": [\\n      12,\\n      305\\n    ],\\n    \\\"bf\\\": [\\n      2.94,\\n      74.7\\n    ],\\n    \\\"tw\\\": [\\n      0.282,\\n      7.16\\n    ],\\n    \\\"tf\\\": [\\n      0.501,\\n      12.7\\n    ],\\n    \\\"x\\\": [\\n      0.698,\\n      17.7\\n    ],\\n    \\\"Ix\\\": [\\n      129,\\n      53.7\\n    ],\\n    \\\"Iy\\\": [\\n      3.86,\\n      1.61\\n    ]\\n  },\\n  {\\n    \\\"DESC\\\": [\\n      \\\"C10$\\\\\\\\times$30\\\",\\n      \\\"C250$\\\\\\\\times$45\\\"\\n    ],\\n    \\\"W\\\": [\\n      30,\\n      45\\n    ],\\n    \\\"A\\\": [\\n      8.81,\\n      5680\\n    ],\\n    \\\"d\\\": [\\n      10,\\n      254\\n    ],\\n    \\\"bf\\\": [\\n      3.03,\\n      77\\n    ],\\n    \\\"tw\\\": [\\n      0.673,\\n      17.1\\n    ],\\n    \\\"tf\\\": [\\n      0.436,\\n      11.1\\n    ],\\n    \\\"x\\\": [\\n      0.649,\\n      16.5\\n    ],\\n    \\\"Ix\\\": [\\n      103,\\n      42.9\\n    ],\\n    \\\"Iy\\\": [\\n      3.93,\\n      1.64\\n    ]\\n  },\\n  {\\n    \\\"DESC\\\": [\\n      \\\"C10$\\\\\\\\times$25\\\",\\n      \\\"C250$\\\\\\\\times$37\\\"\\n    ],\\n    \\\"W\\\": [\\n      25,\\n      37\\n    ],\\n    \\\"A\\\": [\\n      7.35,\\n      4740\\n    ],\\n    \\\"d\\\": [\\n      10,\\n      254\\n    ],\\n    \\\"bf\\\": [\\n      2.89,\\n      73.4\\n    ],\\n    \\\"tw\\\": [\\n      0.526,\\n      13.4\\n    ],\\n    \\\"tf\\\": [\\n      0.436,\\n      11.1\\n    ],\\n    \\\"x\\\": [\\n      0.617,\\n      15.7\\n    ],\\n    \\\"Ix\\\": [\\n      91.1,\\n      37.9\\n    ],\\n    \\\"Iy\\\": [\\n      3.34,\\n      1.39\\n    ]\\n  },\\n  {\\n    \\\"DESC\\\": [\\n      \\\"C10$\\\\\\\\times$20\\\",\\n      \\\"C250$\\\\\\\\times$30\\\"\\n    ],\\n    \\\"W\\\": [\\n      20,\\n      30\\n    ],\\n    \\\"A\\\": [\\n      5.87,\\n      3790\\n    ],\\n    \\\"d\\\": [\\n      10,\\n      254\\n    ],\\n    \\\"bf\\\": [\\n      2.74,\\n      69.6\\n    ],\\n    \\\"tw\\\": [\\n      0.379,\\n      9.63\\n    ],\\n    \\\"tf\\\": [\\n      0.436,\\n      11.1\\n    ],\\n    \\\"x\\\": [\\n      0.606,\\n      15.4\\n    ],\\n    \\\"Ix\\\": [\\n      78.9,\\n      32.8\\n    ],\\n    \\\"Iy\\\": [\\n      2.8,\\n      1.17\\n    ]\\n  },\\n  {\\n    \\\"DESC\\\": [\\n      \\\"C10$\\\\\\\\times$15.3\\\",\\n      \\\"C250$\\\\\\\\times$22.8\\\"\\n    ],\\n    \\\"W\\\": [\\n      15.3,\\n      22.8\\n    ],\\n    \\\"A\\\": [\\n      4.48,\\n      2890\\n    ],\\n    \\\"d\\\": [\\n      10,\\n      254\\n    ],\\n    \\\"bf\\\": [\\n      2.6,\\n      66\\n    ],\\n    \\\"tw\\\": [\\n      0.24,\\n      6.1\\n    ],\\n    \\\"tf\\\": [\\n      0.436,\\n      11.1\\n    ],\\n    \\\"x\\\": [\\n      0.634,\\n      16.1\\n    ],\\n    \\\"Ix\\\": [\\n      67.3,\\n      28\\n    ],\\n    \\\"Iy\\\": [\\n      2.27,\\n      0.945\\n    ]\\n  },\\n  {\\n    \\\"DESC\\\": [\\n      \\\"C9$\\\\\\\\times$20\\\",\\n      \\\"C230$\\\\\\\\times$30\\\"\\n    ],\\n    \\\"W\\\": [\\n      20,\\n      30\\n    ],\\n    \\\"A\\\": [\\n      5.87,\\n      3790\\n    ],\\n    \\\"d\\\": [\\n      9,\\n      229\\n    ],\\n    \\\"bf\\\": [\\n      2.65,\\n      67.3\\n    ],\\n    \\\"tw\\\": [\\n      0.448,\\n      11.4\\n    ],\\n    \\\"tf\\\": [\\n      0.413,\\n      10.5\\n    ],\\n    \\\"x\\\": [\\n      0.583,\\n      14.8\\n    ],\\n    \\\"Ix\\\": [\\n      60.9,\\n      25.3\\n    ],\\n    \\\"Iy\\\": [\\n      2.41,\\n      1\\n    ]\\n  },\\n  {\\n    \\\"DESC\\\": [\\n      \\\"C9$\\\\\\\\times$15\\\",\\n      \\\"C230$\\\\\\\\times$22\\\"\\n    ],\\n    \\\"W\\\": [\\n      15,\\n      22\\n    ],\\n    \\\"A\\\": [\\n      4.4,\\n      2840\\n    ],\\n    \\\"d\\\": [\\n      9,\\n      229\\n    ],\\n    \\\"bf\\\": [\\n      2.49,\\n      63.2\\n    ],\\n    \\\"tw\\\": [\\n      0.285,\\n      7.24\\n    ],\\n    \\\"tf\\\": [\\n      0.413,\\n      10.5\\n    ],\\n    \\\"x\\\": [\\n      0.586,\\n      14.9\\n    ],\\n    \\\"Ix\\\": [\\n      51,\\n      21.2\\n    ],\\n    \\\"Iy\\\": [\\n      1.91,\\n      0.795\\n    ]\\n  },\\n  {\\n    \\\"DESC\\\": [\\n      \\\"C9$\\\\\\\\times$13.4\\\",\\n      \\\"C230$\\\\\\\\times$19.9\\\"\\n    ],\\n    \\\"W\\\": [\\n      13.4,\\n      19.9\\n    ],\\n    \\\"A\\\": [\\n      3.94,\\n      2540\\n    ],\\n    \\\"d\\\": [\\n      9,\\n      229\\n    ],\\n    \\\"bf\\\": [\\n      2.43,\\n      61.7\\n    ],\\n    \\\"tw\\\": [\\n      0.233,\\n      5.92\\n    ],\\n    \\\"tf\\\": [\\n      0.413,\\n      10.5\\n    ],\\n    \\\"x\\\": [\\n      0.601,\\n      15.3\\n    ],\\n    \\\"Ix\\\": [\\n      47.8,\\n      19.9\\n    ],\\n    \\\"Iy\\\": [\\n      1.75,\\n      0.728\\n    ]\\n  },\\n  {\\n    \\\"DESC\\\": [\\n      \\\"C8$\\\\\\\\times$18.75\\\",\\n      \\\"C200$\\\\\\\\times$27.9\\\"\\n    ],\\n    \\\"W\\\": [\\n      18.75,\\n      27.9\\n    ],\\n    \\\"A\\\": [\\n      5.51,\\n      3550\\n    ],\\n    \\\"d\\\": [\\n      8,\\n      203\\n    ],\\n    \\\"bf\\\": [\\n      2.53,\\n      64.3\\n    ],\\n    \\\"tw\\\": [\\n      0.487,\\n      12.4\\n    ],\\n    \\\"tf\\\": [\\n      0.39,\\n      9.91\\n    ],\\n    \\\"x\\\": [\\n      0.565,\\n      14.4\\n    ],\\n    \\\"Ix\\\": [\\n      43.9,\\n      18.3\\n    ],\\n    \\\"Iy\\\": [\\n      1.97,\\n      0.82\\n    ]\\n  },\\n  {\\n    \\\"DESC\\\": [\\n      \\\"C8$\\\\\\\\times$13.75\\\",\\n      \\\"C200$\\\\\\\\times$20.5\\\"\\n    ],\\n    \\\"W\\\": [\\n      13.75,\\n      20.5\\n    ],\\n    \\\"A\\\": [\\n      4.03,\\n      2600\\n    ],\\n    \\\"d\\\": [\\n      8,\\n      203\\n    ],\\n    \\\"bf\\\": [\\n      2.34,\\n      59.4\\n    ],\\n    \\\"tw\\\": [\\n      0.303,\\n      7.7\\n    ],\\n    \\\"tf\\\": [\\n      0.39,\\n      9.91\\n    ],\\n    \\\"x\\\": [\\n      0.554,\\n      14.1\\n    ],\\n    \\\"Ix\\\": [\\n      36.1,\\n      15\\n    ],\\n    \\\"Iy\\\": [\\n      1.52,\\n      0.633\\n    ]\\n  },\\n  {\\n    \\\"DESC\\\": [\\n      \\\"C8$\\\\\\\\times$11.5\\\",\\n      \\\"C200$\\\\\\\\times$17.1\\\"\\n    ],\\n    \\\"W\\\": [\\n      11.5,\\n      17.1\\n    ],\\n    \\\"A\\\": [\\n      3.37,\\n      2170\\n    ],\\n    \\\"d\\\": [\\n      8,\\n      203\\n    ],\\n    \\\"bf\\\": [\\n      2.26,\\n      57.4\\n    ],\\n    \\\"tw\\\": [\\n      0.22,\\n      5.59\\n    ],\\n    \\\"tf\\\": [\\n      0.39,\\n      9.91\\n    ],\\n    \\\"x\\\": [\\n      0.572,\\n      14.5\\n    ],\\n    \\\"Ix\\\": [\\n      32.5,\\n      13.5\\n    ],\\n    \\\"Iy\\\": [\\n      1.31,\\n      0.545\\n    ]\\n  },\\n  {\\n    \\\"DESC\\\": [\\n      \\\"C7$\\\\\\\\times$14.75\\\",\\n      \\\"C180$\\\\\\\\times$22\\\"\\n    ],\\n    \\\"W\\\": [\\n      14.75,\\n      22\\n    ],\\n    \\\"A\\\": [\\n      4.33,\\n      2790\\n    ],\\n    \\\"d\\\": [\\n      7,\\n      178\\n    ],\\n    \\\"bf\\\": [\\n      2.3,\\n      58.4\\n    ],\\n    \\\"tw\\\": [\\n      0.419,\\n      10.6\\n    ],\\n    \\\"tf\\\": [\\n      0.366,\\n      9.3\\n    ],\\n    \\\"x\\\": [\\n      0.532,\\n      13.5\\n    ],\\n    \\\"Ix\\\": [\\n      27.2,\\n      11.3\\n    ],\\n    \\\"Iy\\\": [\\n      1.37,\\n      0.57\\n    ]\\n  },\\n  {\\n    \\\"DESC\\\": [\\n      \\\"C7$\\\\\\\\times$12.25\\\",\\n      \\\"C180$\\\\\\\\times$18.2\\\"\\n    ],\\n    \\\"W\\\": [\\n      12.25,\\n      18.2\\n    ],\\n    \\\"A\\\": [\\n      3.59,\\n      2320\\n    ],\\n    \\\"d\\\": [\\n      7,\\n      178\\n    ],\\n    \\\"bf\\\": [\\n      2.19,\\n      55.6\\n    ],\\n    \\\"tw\\\": [\\n      0.314,\\n      7.98\\n    ],\\n    \\\"tf\\\": [\\n      0.366,\\n      9.3\\n    ],\\n    \\\"x\\\": [\\n      0.525,\\n      13.3\\n    ],\\n    \\\"Ix\\\": [\\n      24.2,\\n      10.1\\n    ],\\n    \\\"Iy\\\": [\\n      1.16,\\n      0.483\\n    ]\\n  },\\n  {\\n    \\\"DESC\\\": [\\n      \\\"C7$\\\\\\\\times$9.8\\\",\\n      \\\"C180$\\\\\\\\times$14.6\\\"\\n    ],\\n    \\\"W\\\": [\\n      9.8,\\n      14.6\\n    ],\\n    \\\"A\\\": [\\n      2.87,\\n      1850\\n    ],\\n    \\\"d\\\": [\\n      7,\\n      178\\n    ],\\n    \\\"bf\\\": [\\n      2.09,\\n      53.1\\n    ],\\n    \\\"tw\\\": [\\n      0.21,\\n      5.33\\n    ],\\n    \\\"tf\\\": [\\n      0.366,\\n      9.3\\n    ],\\n    \\\"x\\\": [\\n      0.541,\\n      13.7\\n    ],\\n    \\\"Ix\\\": [\\n      21.2,\\n      8.82\\n    ],\\n    \\\"Iy\\\": [\\n      0.957,\\n      0.398\\n    ]\\n  },\\n  {\\n    \\\"DESC\\\": [\\n      \\\"C6$\\\\\\\\times$13\\\",\\n      \\\"C150$\\\\\\\\times$19.3\\\"\\n    ],\\n    \\\"W\\\": [\\n      13,\\n      19.3\\n    ],\\n    \\\"A\\\": [\\n      3.82,\\n      2460\\n    ],\\n    \\\"d\\\": [\\n      6,\\n      152\\n    ],\\n    \\\"bf\\\": [\\n      2.16,\\n      54.9\\n    ],\\n    \\\"tw\\\": [\\n      0.437,\\n      11.1\\n    ],\\n    \\\"tf\\\": [\\n      0.343,\\n      8.71\\n    ],\\n    \\\"x\\\": [\\n      0.514,\\n      13.1\\n    ],\\n    \\\"Ix\\\": [\\n      17.3,\\n      7.2\\n    ],\\n    \\\"Iy\\\": [\\n      1.05,\\n      0.437\\n    ]\\n  },\\n  {\\n    \\\"DESC\\\": [\\n      \\\"C6$\\\\\\\\times$10.5\\\",\\n      \\\"C150$\\\\\\\\times$15.6\\\"\\n    ],\\n    \\\"W\\\": [\\n      10.5,\\n      15.6\\n    ],\\n    \\\"A\\\": [\\n      3.07,\\n      1980\\n    ],\\n    \\\"d\\\": [\\n      6,\\n      152\\n    ],\\n    \\\"bf\\\": [\\n      2.03,\\n      51.6\\n    ],\\n    \\\"tw\\\": [\\n      0.314,\\n      7.98\\n    ],\\n    \\\"tf\\\": [\\n      0.343,\\n      8.71\\n    ],\\n    \\\"x\\\": [\\n      0.5,\\n      12.7\\n    ],\\n    \\\"Ix\\\": [\\n      15.1,\\n      6.29\\n    ],\\n    \\\"Iy\\\": [\\n      0.86,\\n      0.358\\n    ]\\n  },\\n  {\\n    \\\"DESC\\\": [\\n      \\\"C6$\\\\\\\\times$8.2\\\",\\n      \\\"C150$\\\\\\\\times$12.2\\\"\\n    ],\\n    \\\"W\\\": [\\n      8.2,\\n      12.2\\n    ],\\n    \\\"A\\\": [\\n      2.39,\\n      1540\\n    ],\\n    \\\"d\\\": [\\n      6,\\n      152\\n    ],\\n    \\\"bf\\\": [\\n      1.92,\\n      48.8\\n    ],\\n    \\\"tw\\\": [\\n      0.2,\\n      5.08\\n    ],\\n    \\\"tf\\\": [\\n      0.343,\\n      8.71\\n    ],\\n    \\\"x\\\": [\\n      0.512,\\n      13\\n    ],\\n    \\\"Ix\\\": [\\n      13.1,\\n      5.45\\n    ],\\n    \\\"Iy\\\": [\\n      0.687,\\n      0.286\\n    ]\\n  },\\n  {\\n    \\\"DESC\\\": [\\n      \\\"C5$\\\\\\\\times$9\\\",\\n      \\\"C130$\\\\\\\\times$13\\\"\\n    ],\\n    \\\"W\\\": [\\n      9,\\n      13\\n    ],\\n    \\\"A\\\": [\\n      2.64,\\n      1700\\n    ],\\n    \\\"d\\\": [\\n      5,\\n      127\\n    ],\\n    \\\"bf\\\": [\\n      1.89,\\n      48\\n    ],\\n    \\\"tw\\\": [\\n      0.325,\\n      8.26\\n    ],\\n    \\\"tf\\\": [\\n      0.32,\\n      8.13\\n    ],\\n    \\\"x\\\": [\\n      0.478,\\n      12.1\\n    ],\\n    \\\"Ix\\\": [\\n      8.89,\\n      3.7\\n    ],\\n    \\\"Iy\\\": [\\n      0.624,\\n      0.26\\n    ]\\n  },\\n  {\\n    \\\"DESC\\\": [\\n      \\\"C5$\\\\\\\\times$6.7\\\",\\n      \\\"C130$\\\\\\\\times$10.4\\\"\\n    ],\\n    \\\"W\\\": [\\n      6.7,\\n      10.4\\n    ],\\n    \\\"A\\\": [\\n      1.97,\\n      1270\\n    ],\\n    \\\"d\\\": [\\n      5,\\n      127\\n    ],\\n    \\\"bf\\\": [\\n      1.75,\\n      44.5\\n    ],\\n    \\\"tw\\\": [\\n      0.19,\\n      4.83\\n    ],\\n    \\\"tf\\\": [\\n      0.32,\\n      8.13\\n    ],\\n    \\\"x\\\": [\\n      0.484,\\n      12.3\\n    ],\\n    \\\"Ix\\\": [\\n      7.48,\\n      3.11\\n    ],\\n    \\\"Iy\\\": [\\n      0.47,\\n      0.196\\n    ]\\n  },\\n  {\\n    \\\"DESC\\\": [\\n      \\\"C4$\\\\\\\\times$7.25\\\",\\n      \\\"C100$\\\\\\\\times$10.8\\\"\\n    ],\\n    \\\"W\\\": [\\n      7.25,\\n      10.8\\n    ],\\n    \\\"A\\\": [\\n      2.13,\\n      1370\\n    ],\\n    \\\"d\\\": [\\n      4,\\n      102\\n    ],\\n    \\\"bf\\\": [\\n      1.72,\\n      43.7\\n    ],\\n    \\\"tw\\\": [\\n      0.321,\\n      8.15\\n    ],\\n    \\\"tf\\\": [\\n      0.296,\\n      7.52\\n    ],\\n    \\\"x\\\": [\\n      0.459,\\n      11.7\\n    ],\\n    \\\"Ix\\\": [\\n      4.58,\\n      1.91\\n    ],\\n    \\\"Iy\\\": [\\n      0.425,\\n      0.177\\n    ]\\n  },\\n  {\\n    \\\"DESC\\\": [\\n      \\\"C4$\\\\\\\\times$6.25\\\",\\n      \\\"C100$\\\\\\\\times$9.3\\\"\\n    ],\\n    \\\"W\\\": [\\n      6.25,\\n      9.3\\n    ],\\n    \\\"A\\\": [\\n      1.84,\\n      1190\\n    ],\\n    \\\"d\\\": [\\n      4,\\n      102\\n    ],\\n    \\\"bf\\\": [\\n      1.65,\\n      41.9\\n    ],\\n    \\\"tw\\\": [\\n      0.247,\\n      6.27\\n    ],\\n    \\\"tf\\\": [\\n      0.296,\\n      7.52\\n    ],\\n    \\\"x\\\": [\\n      0.453,\\n      11.5\\n    ],\\n    \\\"Ix\\\": [\\n      4.19,\\n      1.74\\n    ],\\n    \\\"Iy\\\": [\\n      0.374,\\n      0.156\\n    ]\\n  },\\n  {\\n    \\\"DESC\\\": [\\n      \\\"C4$\\\\\\\\times$5.4\\\",\\n      \\\"C100$\\\\\\\\times$8\\\"\\n    ],\\n    \\\"W\\\": [\\n      5.4,\\n      8\\n    ],\\n    \\\"A\\\": [\\n      1.58,\\n      1020\\n    ],\\n    \\\"d\\\": [\\n      4,\\n      102\\n    ],\\n    \\\"bf\\\": [\\n      1.58,\\n      40.1\\n    ],\\n    \\\"tw\\\": [\\n      0.184,\\n      4.67\\n    ],\\n    \\\"tf\\\": [\\n      0.296,\\n      7.52\\n    ],\\n    \\\"x\\\": [\\n      0.457,\\n      11.6\\n    ],\\n    \\\"Ix\\\": [\\n      3.85,\\n      1.6\\n    ],\\n    \\\"Iy\\\": [\\n      0.312,\\n      0.13\\n    ]\\n  },\\n  {\\n    \\\"DESC\\\": [\\n      \\\"C4$\\\\\\\\times$4.5\\\",\\n      \\\"C100$\\\\\\\\times$6.7\\\"\\n    ],\\n    \\\"W\\\": [\\n      4.5,\\n      6.7\\n    ],\\n    \\\"A\\\": [\\n      1.34,\\n      865\\n    ],\\n    \\\"d\\\": [\\n      4,\\n      102\\n    ],\\n    \\\"bf\\\": [\\n      1.52,\\n      38.6\\n    ],\\n    \\\"tw\\\": [\\n      0.125,\\n      3.18\\n    ],\\n    \\\"tf\\\": [\\n      0.296,\\n      7.52\\n    ],\\n    \\\"x\\\": [\\n      0.473,\\n      12\\n    ],\\n    \\\"Ix\\\": [\\n      3.53,\\n      1.47\\n    ],\\n    \\\"Iy\\\": [\\n      0.265,\\n      0.11\\n    ]\\n  },\\n  {\\n    \\\"DESC\\\": [\\n      \\\"C3$\\\\\\\\times$6\\\",\\n      \\\"C75$\\\\\\\\times$8.9\\\"\\n    ],\\n    \\\"W\\\": [\\n      6,\\n      8.9\\n    ],\\n    \\\"A\\\": [\\n      1.76,\\n      1140\\n    ],\\n    \\\"d\\\": [\\n      3,\\n      76.2\\n    ],\\n    \\\"bf\\\": [\\n      1.6,\\n      40.6\\n    ],\\n    \\\"tw\\\": [\\n      0.356,\\n      9.04\\n    ],\\n    \\\"tf\\\": [\\n      0.273,\\n      6.93\\n    ],\\n    \\\"x\\\": [\\n      0.455,\\n      11.6\\n    ],\\n    \\\"Ix\\\": [\\n      2.07,\\n      0.862\\n    ],\\n    \\\"Iy\\\": [\\n      0.3,\\n      0.125\\n    ]\\n  },\\n  {\\n    \\\"DESC\\\": [\\n      \\\"C3$\\\\\\\\times$5\\\",\\n      \\\"C75$\\\\\\\\times$7.4\\\"\\n    ],\\n    \\\"W\\\": [\\n      5,\\n      7.4\\n    ],\\n    \\\"A\\\": [\\n      1.47,\\n      948\\n    ],\\n    \\\"d\\\": [\\n      3,\\n      76.2\\n    ],\\n    \\\"bf\\\": [\\n      1.5,\\n      38.1\\n    ],\\n    \\\"tw\\\": [\\n      0.258,\\n      6.55\\n    ],\\n    \\\"tf\\\": [\\n      0.273,\\n      6.93\\n    ],\\n    \\\"x\\\": [\\n      0.439,\\n      11.2\\n    ],\\n    \\\"Ix\\\": [\\n      1.85,\\n      0.77\\n    ],\\n    \\\"Iy\\\": [\\n      0.241,\\n      0.1\\n    ]\\n  },\\n  {\\n    \\\"DESC\\\": [\\n      \\\"C3$\\\\\\\\times$4.1\\\",\\n      \\\"C75$\\\\\\\\times$6.1\\\"\\n    ],\\n    \\\"W\\\": [\\n      4.1,\\n      6.1\\n    ],\\n    \\\"A\\\": [\\n      1.2,\\n      774\\n    ],\\n    \\\"d\\\": [\\n      3,\\n      76.2\\n    ],\\n    \\\"bf\\\": [\\n      1.41,\\n      35.8\\n    ],\\n    \\\"tw\\\": [\\n      0.17,\\n      4.32\\n    ],\\n    \\\"tf\\\": [\\n      0.273,\\n      6.93\\n    ],\\n    \\\"x\\\": [\\n      0.437,\\n      11.1\\n    ],\\n    \\\"Ix\\\": [\\n      1.65,\\n      0.687\\n    ],\\n    \\\"Iy\\\": [\\n      0.191,\\n      0.0795\\n    ]\\n  },\\n  {\\n    \\\"DESC\\\": [\\n      \\\"C3$\\\\\\\\times$3.5\\\",\\n      \\\"C75$\\\\\\\\times$5.2\\\"\\n    ],\\n    \\\"W\\\": [\\n      3.5,\\n      5.2\\n    ],\\n    \\\"A\\\": [\\n      1.09,\\n      703\\n    ],\\n    \\\"d\\\": [\\n      3,\\n      76.2\\n    ],\\n    \\\"bf\\\": [\\n      1.37,\\n      34.8\\n    ],\\n    \\\"tw\\\": [\\n      0.132,\\n      3.35\\n    ],\\n    \\\"tf\\\": [\\n      0.273,\\n      6.93\\n    ],\\n    \\\"x\\\": [\\n      0.443,\\n      11.3\\n    ],\\n    \\\"Ix\\\": [\\n      1.57,\\n      0.653\\n    ],\\n    \\\"Iy\\\": [\\n      0.169,\\n      0.0703\\n    ]\\n  }\\n]\"))", "description": "<p>AISC data for channels</p>", "templateType": "json", "can_override": false}, "b": {"name": "b", "group": "Ungrouped variables", "definition": "round( random(0.4.. 1.2) d, if(unit=0,'0.25 in', '10 mm'))", "description": "<p>distance between channels</p>", "templateType": "anything", "can_override": false}, "unit4": {"name": "unit4", "group": "Ungrouped variables", "definition": "v->siground(qty(v,units+'^4'),4) *[1,10^6][unit]", "description": "", "templateType": "anything", "can_override": false}, "unit1": {"name": "unit1", "group": "Ungrouped variables", "definition": "v -> siground(qty(v,units),4)", "description": "", "templateType": "anything", "can_override": false}, "unit2": {"name": "unit2", "group": "Ungrouped variables", "definition": "v->siground(qty(v,units+'^2'),4)", "description": "", "templateType": "anything", "can_override": false}, "unit3": {"name": "unit3", "group": "Ungrouped variables", "definition": "v->siground(qty(v,units+'^3'),4)", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["index", "b", "w", "h", "unit", "units", "unit1", "unit2", "unit3", "unit4"], "variable_groups": [{"name": "Channel Properties", "variables": ["data", "C_data", "Desc", "Ibary_c", "Ibarx_c", "A_C", "xbar", "d", "bf", "t"]}, {"name": "Solution", "variables": ["Ibarx_r", "A_r", "d_y", "Ix_R", "Ix", "A_t", "kx", "Ibary_R", "d_x", "Iy_C", "Iy", "ky"]}], "functions": {"disp": {"parameters": [["q", "quantity"]], "type": "number", "language": "jme", "definition": "siground(q,4)"}, "ggb": {"parameters": [["q", "quantity"]], "type": "number", "language": "jme", "definition": "siground(scalar(q in 'in'),4)"}}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p style=\"padding-left: 40px;\">$I_x = $[[0]]</p>\n<p style=\"padding-left: 40px;\">$k_x = $[[1]]</p>\n<p style=\"padding-left: 40px;\">$I_y = $[[2]]</p>\n<p style=\"padding-left: 40px;\">$k_y = $[[3]]</p>", "gaps": [{"type": "engineering-answer", "useCustomName": true, "customName": "Ix", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "Ix", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "kx", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "kx", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "Iy", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "Iy", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "ky", "marks": "5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "ky", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}], "sortAnswers": false}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}], "resources": ["question-resources/channel_EfJjxiR.png", "question-resources/builtupbeam.ggb"]}]}], "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}]}