// Numbas version: finer_feedback_settings
{"name": "Moment of Inertia: Basic Shapes from the MOI Table", "extensions": ["geogebra", "quantities"], "custom_part_types": [{"source": {"pk": 19, "author": {"name": "William Haynes", "pk": 2530}, "edit_page": "/part_type/19/edit"}, "name": "Engineering Accuracy with units", "short_name": "engineering-answer", "description": "<p>A value with units marked right if within an adjustable % error of the correct value. &nbsp;Marked close if within a wider margin of error.</p>", "help_url": "", "input_widget": "string", "input_options": {"correctAnswer": "siground(settings['correctAnswer'],4)", "hint": {"static": true, "value": ""}, "allowEmpty": {"static": true, "value": true}}, "can_be_gap": true, "can_be_step": true, "marking_script": "mark:\nswitch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)\n\ninterpreted_answer:\nqty(student_scalar, student_units)\n\n\n\ncorrect_quantity:\nsettings[\"correctAnswer\"]\n\n\n\ncorrect_units:\nunits(correct_quantity)\n\n\nallowed_notation_styles:\n[\"plain\",\"en\"]\n\nmatch_student_number:\nmatchnumber(studentAnswer,allowed_notation_styles)\n\nstudent_scalar:\nmatch_student_number[1]\n\nstudent_units:\nreplace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")\n\ngood_units:\ntry(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n\n\nstudent_quantity:\nswitch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n\n\npercent_error:\ntry(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   \n\nright:\npercent_error <= settings['right']\n\n\nclose:\nright_sign and percent_error <= settings['close']\n\nright_sign:\nsign(student_scalar) = sign(correct_quantity)", "marking_notes": [{"name": "mark", "description": "This is the main marking note. It should award credit and provide feedback based on the student's answer.", "definition": "switch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)"}, {"name": "interpreted_answer", "description": "A value representing the student's answer to this part.", "definition": "qty(student_scalar, student_units)\n\n"}, {"name": "correct_quantity", "description": "", "definition": "settings[\"correctAnswer\"]\n\n"}, {"name": "correct_units", "description": "", "definition": "units(correct_quantity)\n"}, {"name": "allowed_notation_styles", "description": "", "definition": "[\"plain\",\"en\"]"}, {"name": "match_student_number", "description": "", "definition": "matchnumber(studentAnswer,allowed_notation_styles)"}, {"name": "student_scalar", "description": "", "definition": "match_student_number[1]"}, {"name": "student_units", "description": "<p>Modify&nbsp;the&nbsp;unit portion of&nbsp;the student's answer by</p>\n<p>1. replacing \"ohms\" with \"ohm\" &nbsp;case insensitive</p>\n<p>2.&nbsp;replacing '-'&nbsp;with ' '&nbsp;</p>\n<p>3. replacing '&deg;' with ' deg'&nbsp;</p>\n<p>to allow answers like 10 ft-lb and 30&deg;</p>", "definition": "replace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")"}, {"name": "good_units", "description": "", "definition": "try(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n"}, {"name": "student_quantity", "description": "<p>This fixes the student answer for two common errors. &nbsp;</p>\n<p>If student_units are wrong&nbsp;&nbsp;- replace with correct units</p>\n<p>If student_scalar has the wrong sign - replace with right sign</p>\n<p>If student makes both errors, only one gets fixed.</p>", "definition": "switch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n"}, {"name": "percent_error", "description": "", "definition": "try(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   "}, {"name": "right", "description": "", "definition": "percent_error <= settings['right']\n"}, {"name": "close", "description": "<p>Only marked close if the student actually has the right sign.</p>", "definition": "right_sign and percent_error <= settings['close']"}, {"name": "right_sign", "description": "", "definition": "sign(student_scalar) = sign(correct_quantity) "}], "settings": [{"name": "correctAnswer", "label": "Correct Quantity.", "help_url": "", "hint": "The correct answer given as a JME quantity.", "input_type": "code", "default_value": "", "evaluate": true}, {"name": "right", "label": "% Accuracy for right.", "help_url": "", "hint": "Question will be considered correct if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "0.2", "evaluate": true}, {"name": "close", "label": "% Accuracy for close.", "help_url": "", "hint": "Question will be considered close if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "1.0", "evaluate": true}, {"name": "C1", "label": "Close with units.", "help_url": "", "hint": "Partial Credit for close value with appropriate units. &nbsp;if correct answer is 100 N and close is &plusmn;1%,<br />99 &nbsp;N is accepted.", "input_type": "percent", "default_value": "75"}, {"name": "C2", "label": "No units or wrong sign", "help_url": "", "hint": "Partial credit for forgetting units or using wrong sign.<br />If the correct answer is 100 N, both 100 and -100 N are accepted.", "input_type": "percent", "default_value": "50"}, {"name": "C3", "label": "Close, no units.", "help_url": "", "hint": "Partial Credit for close value but forgotten units.<br />This value would be close&nbsp;if the expected units were provided. &nbsp;If the correct answer is 100 N, and close is &plusmn;1%,<br />99 is accepted.", "input_type": "percent", "default_value": "25"}], "public_availability": "always", "published": true, "extensions": ["quantities"]}], "resources": [["question-resources/moi.png", "moi.png"], ["question-resources/PastedGraphic-1.png", "PastedGraphic-1.png"], ["question-resources/basicShapes.ggb", "basicShapes.ggb"]], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Moment of Inertia: Basic Shapes from the MOI Table", "tags": ["Area moment of inertia", "area moment of inertia", "Mechanics", "mechanics", "Statics", "statics", "table of properties"], "metadata": {"description": "<p>Use a table of properties to find the Area Moment of inertia for simple shapes: rectangle, triangle, circle, semicircle, and quarter circle.</p>\n<p>&nbsp;The parallel axis theorem is not required for any of these shapes. &nbsp;One situation requires subtracting a triangle from a rectangle however.</p>\n<p>Distinguish between centroidal and non-centroidal moments of inertia.</p>", "licence": "Creative Commons Attribution-ShareAlike 4.0 International"}, "statement": "<p>&nbsp;A {shape} ({dimensions}) and its centroid are shown.</p>\n<table style=\"margin-left: auto; margin-right: auto;\">\n<tbody>\n<tr>\n<td style=\"text-align: center;\">{geogebra_applet('ndkfbkp9',['b': b, 'h': h, 'i': index, 'Z': Z])}</td>\n</tr>\n</tbody>\n</table>\n<p>Determine the moment of inertia of the {shape} with respect to the $x$ and $y$ axis. Are either of these values <em>centroidal</em> moments of inertia?</p>", "advice": "<p></p>\n<p>Use the information in the table below.</p>\n<p><img src=\"resources/question-resources/moi.png\" width=\"364\" height=\"1017\"/></p>", "rulesets": {}, "extensions": ["geogebra", "quantities"], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"b": {"name": "b", "group": "inputs", "definition": "random(-8..8 except 0)", "description": "<p>horizontal dimension of rectangles and triangles</p>", "templateType": "anything", "can_override": false}, "h": {"name": "h", "group": "inputs", "definition": "random(-8..8 except 0)", "description": "<p>vertical dimension of rectangles and triangles</p>", "templateType": "anything", "can_override": false}, "r": {"name": "r", "group": "inputs", "definition": "b", "description": "<p>radius of circle or semicircles.</p>", "templateType": "anything", "can_override": false}, "index": {"name": "index", "group": "inputs", "definition": "random(1..14)", "description": "<p>selects the shape -- see resources for diagram</p>", "templateType": "anything", "can_override": false}, "shape": {"name": "shape", "group": "Description", "definition": "if(index=1,'circle',if(index<=4,'rectangle',if(index <=10,'triangle',if(index<=12,'semicircle','quarter-circle'))))", "description": "", "templateType": "anything", "can_override": false}, "dimensions": {"name": "dimensions", "group": "Description", "definition": "if(isCircular,\n  \"$r = \\\\var{abs(b units) }$\",\n  \"$\\\\var{abs(b units)} \\\\times \\\\var{abs(h units)}$\"\n) ", "description": "", "templateType": "anything", "can_override": false}, "isCircular": {"name": "isCircular", "group": "Description", "definition": "index = 1 or  index > 10", "description": "", "templateType": "anything", "can_override": false}, "units": {"name": "units", "group": "inputs", "definition": "qty(random(['in', 'ft', 'mm', 'cm','m']))", "description": "", "templateType": "anything", "can_override": false}, "Ix": {"name": "Ix", "group": "Solution", "definition": "siground([ 0,pir4/4, \n  bh3/3, bh3/12, bh3/3, \n  bh3/12, bh3/12, bh3/36, bh3/36, bh3/12,bh3/36,\n  pir4/8,pir4/8,\n  pir4/16,pir4/16][index],4)", "description": "<p>0 is placeholder for zero-indexed array.</p>", "templateType": "anything", "can_override": false}, "PIR4": {"name": "PIR4", "group": "Solution", "definition": "pi r^4 * units * units * units * units", "description": "", "templateType": "anything", "can_override": false}, "bh3": {"name": "bh3", "group": "Solution", "definition": "abs(b h^3) * units * units * units * units", "description": "", "templateType": "anything", "can_override": false}, "hb3": {"name": "hb3", "group": "Solution", "definition": "abs(h b^3) * units * units * units * units", "description": "", "templateType": "anything", "can_override": false}, "Iy": {"name": "Iy", "group": "Solution", "definition": "siground([0, pir4/4,\n  hb3/3, hb3/3, hb3/12,\n  hb3/4, hb3/12, hb3/4, hb3/12, hb3/36, hb3/36,\n  pir4/8,pir4/8,\n  pir4/16,pir4/16][index],4)", "description": "", "templateType": "anything", "can_override": false}, "centroidalaxis": {"name": "centroidalaxis", "group": "Solution", "definition": "[[],both,neither, x, y, neither, neither, x, x, y, both, x, y, neither, neither][index]\n  ", "description": "<p>marking matrix for centroidal axis [x,y,both,neither]&nbsp;</p>", "templateType": "anything", "can_override": false}, "x": {"name": "x", "group": "Solution", "definition": "[10,0,0,0]", "description": "", "templateType": "anything", "can_override": false}, "y": {"name": "y", "group": "Solution", "definition": "[0,10,0,0]", "description": "", "templateType": "anything", "can_override": false}, "both": {"name": "both", "group": "Solution", "definition": "[0,0,10,0]", "description": "", "templateType": "anything", "can_override": false}, "neither": {"name": "neither", "group": "Solution", "definition": "[0,0,0,10]", "description": "", "templateType": "anything", "can_override": false}, "applet": {"name": "applet", "group": "Ungrouped variables", "definition": "\"applet()\"", "description": "<p>disabled calling the javascript applet</p>", "templateType": "anything", "can_override": false}, "Z": {"name": "Z", "group": "inputs", "definition": "10", "description": "<p>ZOOM</p>", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "abs(b)> 2 and abs(h) > 2", "maxRuns": 100}, "ungrouped_variables": ["applet"], "variable_groups": [{"name": "inputs", "variables": ["h", "r", "b", "index", "units", "Z"]}, {"name": "Description", "variables": ["shape", "dimensions", "isCircular"]}, {"name": "Solution", "variables": ["Ix", "PIR4", "bh3", "hb3", "Iy", "centroidalaxis", "x", "y", "both", "neither"]}], "functions": {"NotUsed": {"parameters": [], "type": "ggbapplet", "language": "javascript", "definition": "// Create the worksheet. \n// This function returns an object with a container `element` and a `promise` resolving to a GeoGebra applet.\nvar params = {\n  filename: 'resources/question-resources/basicShapes.ggb',\n  width: 250,\n  height: 250\n}\n\nvar result = Numbas.extensions.geogebra.createGeogebraApplet(params);\n\n// Once the applet has loaded, run some commands to manipulate the worksheet.\nresult.promise.then(function(d) {\n  var app = d.app;\n  question.applet = d;\n  \n   function setGGBPoint(g_name, n_name = g_name) {\n    // moves point in GGB to location of Numbas Vector Variable\n    // g_name = geogebra point, n_name = numbas vector\n    var pt = question.scope.evaluate(n_name).value\n    app.setFixed(g_name,false,false);\n    app.setCoords(g_name, pt[0], pt[1]);\n    app.setFixed(g_name,true,true);\n  }\n  \n   function setGGBNumber(g_name, n_name = g_name) {\n    // Sets number in GGB to a Numbas Variable\n    var n = question.scope.evaluate(n_name).value;\n    app.setValue(g_name,n);\n  }\n  \n  app.setValue(\"Z\",5);  //zoom\n  setGGBNumber(\"b\");\n  setGGBNumber(\"h\");\n  setGGBNumber(\"i\",\"index\")\n\n});\n\n// This function returns the result of `createGeogebraApplet` as an object \n// with the JME data type 'ggbapplet', which can be substituted into the question's content.\nreturn new Numbas.jme.types.ggbapplet(result);"}}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p style=\"padding-left: 40px;\">$I_x$ = &nbsp;[[0]]&nbsp;</p>\n<p style=\"padding-left: 40px;\">$I_y$ =&nbsp;[[1]]</p>\n<p>Are either&nbsp;of these values <em>centroidal</em> moments of inertia? [[2]]</p>", "gaps": [{"type": "engineering-answer", "useCustomName": true, "customName": "Ix", "marks": "20", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "Ix", "right": "0.2", "close": "1.0", "C1": "80", "C2": "80", "C3": "60"}}, {"type": "engineering-answer", "useCustomName": true, "customName": "Iy", "marks": "20", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "Iy", "right": "0.2", "close": "1.0", "C1": "80", "C2": "80", "C3": "60"}}, {"type": "1_n_2", "useCustomName": true, "customName": "centroidal", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "radiogroup", "displayColumns": "1", "showBlankOption": true, "showCellAnswerState": true, "choices": ["$I_x =\\bar{I}_x$", "$I_y = \\bar{I}_y$", "Both", "Neither"], "matrix": "centroidalaxis"}], "sortAnswers": false}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}], "resources": ["question-resources/moi.png", "question-resources/PastedGraphic-1.png", "question-resources/basicShapes.ggb"]}]}], "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}]}