// Numbas version: finer_feedback_settings
{"name": "Frame: A-Frame with load", "extensions": ["geogebra", "quantities"], "custom_part_types": [{"source": {"pk": 19, "author": {"name": "William Haynes", "pk": 2530}, "edit_page": "/part_type/19/edit"}, "name": "Engineering Accuracy with units", "short_name": "engineering-answer", "description": "<p>A value with units marked right if within an adjustable % error of the correct value. &nbsp;Marked close if within a wider margin of error.</p>", "help_url": "", "input_widget": "string", "input_options": {"correctAnswer": "siground(settings['correctAnswer'],4)", "hint": {"static": true, "value": ""}, "allowEmpty": {"static": true, "value": true}}, "can_be_gap": true, "can_be_step": true, "marking_script": "mark:\nswitch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)\n\ninterpreted_answer:\nqty(student_scalar, student_units)\n\n\n\ncorrect_quantity:\nsettings[\"correctAnswer\"]\n\n\n\ncorrect_units:\nunits(correct_quantity)\n\n\nallowed_notation_styles:\n[\"plain\",\"en\"]\n\nmatch_student_number:\nmatchnumber(studentAnswer,allowed_notation_styles)\n\nstudent_scalar:\nmatch_student_number[1]\n\nstudent_units:\nreplace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")\n\ngood_units:\ntry(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n\n\nstudent_quantity:\nswitch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n\n\npercent_error:\ntry(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   \n\nright:\npercent_error <= settings['right']\n\n\nclose:\nright_sign and percent_error <= settings['close']\n\nright_sign:\nsign(student_scalar) = sign(correct_quantity)", "marking_notes": [{"name": "mark", "description": "This is the main marking note. It should award credit and provide feedback based on the student's answer.", "definition": "switch( \n  right and good_units and right_sign, add_credit(1.0,'Correct.'),\n  right and good_units and not right_sign, add_credit(settings['C2'],'Wrong sign.'),\n  right and right_sign and not good_units, add_credit(settings['C2'],'Correct value, but wrong or missing units.'),\n  close and good_units, add_credit(settings['C1'],'Close.'),\n  close and not good_units, add_credit(settings['C3'],'Answer is close, but wrong or missing units.'),\n  incorrect('Wrong answer.')\n)"}, {"name": "interpreted_answer", "description": "A value representing the student's answer to this part.", "definition": "qty(student_scalar, student_units)\n\n"}, {"name": "correct_quantity", "description": "", "definition": "settings[\"correctAnswer\"]\n\n"}, {"name": "correct_units", "description": "", "definition": "units(correct_quantity)\n"}, {"name": "allowed_notation_styles", "description": "", "definition": "[\"plain\",\"en\"]"}, {"name": "match_student_number", "description": "", "definition": "matchnumber(studentAnswer,allowed_notation_styles)"}, {"name": "student_scalar", "description": "", "definition": "match_student_number[1]"}, {"name": "student_units", "description": "<p>Modify&nbsp;the&nbsp;unit portion of&nbsp;the student's answer by</p>\n<p>1. replacing \"ohms\" with \"ohm\" &nbsp;case insensitive</p>\n<p>2.&nbsp;replacing '-'&nbsp;with ' '&nbsp;</p>\n<p>3. replacing '&deg;' with ' deg'&nbsp;</p>\n<p>to allow answers like 10 ft-lb and 30&deg;</p>", "definition": "replace_regex('ohms','ohm',\n  replace_regex('\u00b0', ' deg',\n    replace_regex('-', ' ' ,\n      studentAnswer[len(match_student_number[0])..len(studentAnswer)])),\"i\")"}, {"name": "good_units", "description": "", "definition": "try(\ncompatible(quantity(1, student_units),correct_units),\nmsg,\nfeedback(msg);false)\n"}, {"name": "student_quantity", "description": "<p>This fixes the student answer for two common errors. &nbsp;</p>\n<p>If student_units are wrong&nbsp;&nbsp;- replace with correct units</p>\n<p>If student_scalar has the wrong sign - replace with right sign</p>\n<p>If student makes both errors, only one gets fixed.</p>", "definition": "switch(not good_units, \n       student_scalar * correct_units, \n       not right_sign,\n       -quantity(student_scalar, student_units),\n       quantity(student_scalar,student_units)\n)\n  \n"}, {"name": "percent_error", "description": "", "definition": "try(\nscalar(abs((correct_quantity - student_quantity)/correct_quantity))*100  \n,msg,\nif(student_quantity=correct_quantity,0,100))\n   "}, {"name": "right", "description": "", "definition": "percent_error <= settings['right']\n"}, {"name": "close", "description": "<p>Only marked close if the student actually has the right sign.</p>", "definition": "right_sign and percent_error <= settings['close']"}, {"name": "right_sign", "description": "", "definition": "sign(student_scalar) = sign(correct_quantity) "}], "settings": [{"name": "correctAnswer", "label": "Correct Quantity.", "help_url": "", "hint": "The correct answer given as a JME quantity.", "input_type": "code", "default_value": "", "evaluate": true}, {"name": "right", "label": "% Accuracy for right.", "help_url": "", "hint": "Question will be considered correct if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "0.2", "evaluate": true}, {"name": "close", "label": "% Accuracy for close.", "help_url": "", "hint": "Question will be considered close if the scalar part of the student's answer is within this % of correct value.", "input_type": "code", "default_value": "1.0", "evaluate": true}, {"name": "C1", "label": "Close with units.", "help_url": "", "hint": "Partial Credit for close value with appropriate units. &nbsp;if correct answer is 100 N and close is &plusmn;1%,<br />99 &nbsp;N is accepted.", "input_type": "percent", "default_value": "75"}, {"name": "C2", "label": "No units or wrong sign", "help_url": "", "hint": "Partial credit for forgetting units or using wrong sign.<br />If the correct answer is 100 N, both 100 and -100 N are accepted.", "input_type": "percent", "default_value": "50"}, {"name": "C3", "label": "Close, no units.", "help_url": "", "hint": "Partial Credit for close value but forgotten units.<br />This value would be close&nbsp;if the expected units were provided. &nbsp;If the correct answer is 100 N, and close is &plusmn;1%,<br />99 is accepted.", "input_type": "percent", "default_value": "25"}], "public_availability": "always", "published": true, "extensions": ["quantities"]}], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Frame: A-Frame with load", "tags": ["Equilibrium", "equilibrium", "frame", "Frame", "Mechanics", "mechanics", "Statics", "statics"], "metadata": {"description": "Solve for the internal forces at on a multipart frame.", "licence": "Creative Commons Attribution-NonCommercial 4.0 International"}, "statement": "<p>The frame shown supports a {F} load at point $F$. Determine the $x$- and $y$- vector components of the forces acting on member $ADC$ at pins $A$, $D$ and $C$.</p>\n<table style=\"margin-left: auto; margin-right: auto; width: 15px;\">\n<tbody>\n<tr>\n<td style=\"width: 0px;\">{geogebra_applet('wrp4gmsn',var_1)}</td>\n</tr>\n</tbody>\n</table>\n<p>Note: Vector components are vectors, they have magnitude and direction.</p>", "advice": "<p></p>\n<p><span data-jme-visible=\"debug\">$L = \\var{L}$ $b = \\var{b}$ $h =\\var{h}$ &nbsp;$d=\\var{d}$</span></p>\n<p><span data-jme-visible=\"debug\">$DE = \\var{DE}$</span></p>\n<p><span data-jme-visible=\"debug\">$EF = \\var{EF}$</span></p>\n<p>Start by drawing neat, labeled, consistent free body diagrams of the whole frame and each of its parts. &nbsp;Exposed internal forces must occur in equal and opposite pairs. &nbsp;In the diagram below, the external forces are drawn in red.</p>\n<table style=\"margin-left: auto; margin-right: auto;\">\n<tbody>\n<tr>\n<td>{geogebra_applet('x3npnvsb',var_2)}</td>\n</tr>\n</tbody>\n</table>\n<p>Count the unknowns on each diagram. &nbsp;If any diagram has three unknowns, consider starting there. &nbsp;In this case FBD I has three unknowns ($A_x, A_y$, and $B$).</p>\n<p><strong>FBD I</strong></p>\n<p>&nbsp;\\[\\begin{align} \\text{I: }\\Sigma M_A &amp;= 0\\\\ &nbsp;B ( \\simplify[]{{AD}+{DE} + {AD} }) &nbsp;&amp;= F(&nbsp;\\simplify[!collectNumbers]{{AD}+{DE} + {EF} })\\\\B&amp;=\\var{F} \\left( \\dfrac{\\var{AD+qty(L,unit_d)}}{\\var{qty(b,unit_d)}}\\right) \\\\ &amp;= \\var{b_y}\\uparrow \\\\\\\\\\text{I: }\\Sigma F_y &amp;=0\\\\A_y+B &amp;= F\\\\A_y &amp;= \\var{F} - \\var{b_y}\\\\&amp;=\\var{(a_y)}\\uparrow\\\\\\\\\\text{I: }\\Sigma F_x &amp;= 0\\\\A_x&amp;=0\\end{align}\\]</p>\n<p><strong>FBD II</strong></p>\n<p>With $A_x, A_y$, and $B$ known, the remaining free body diagrams all have four unknowns but, luckily in&nbsp;this problem, the lines of action of forces&nbsp;$E_x, E_y$, and $D_x$ all intersect at point $E$ on FBD II.&nbsp;&nbsp;By taking moments there, we can find $D_y$. &nbsp;If this was not the case, you would need to select two free body diagrams and take moments&nbsp;appropriately to yield&nbsp;two equations with two unknowns&nbsp;which can be solved simultaneously. &nbsp;</p>\n<p>\\[\\begin{align} \\text{II: }\\Sigma M_E &amp;=0\\\\ D_y(\\var{DE}) &amp;= F(\\var{EF})\\\\D_y &amp;= \\var{F}\\left(\\dfrac{\\var{EF}}{\\var{DE}}\\right)\\\\ &amp;=\\var{d(D_y)}\\downarrow\\end{align}\\]</p>\n<p><strong>FBD III</strong></p>\n<p>With $A_x, A_y$ and $D_y$ found, only three unknowns remain on FBD III, so it can be solved to satisfy the problem statement. &nbsp;Remember that negative results indicate that the force acts in the direction opposite to the one assumed on the free body diagram.</p>\n<p>\\[\\begin{align} \\text{III: }\\Sigma M_C &amp;= 0\\\\ - A_y&nbsp;( \\var{b/2}) -&nbsp;D_x(\\var{h-d}) -&nbsp;D_y(\\var{scalar(DE/2)}) &nbsp;&amp;= 0&nbsp;\\\\D_x&nbsp;&amp;=\\dfrac{-\\var{b/2} A_y - \\var{scalar(DE/2)}D_y}{\\var{h-d}} \\\\ &amp;=-\\dfrac{\\var{b/2} (\\var{d(A_y)}) +\\var{scalar(DE/2)}(\\var{d(D_y)})}{\\var{h-d}} \\\\&amp;= \\var{d(d_x)} \\leftarrow\\\\\\\\\\text{III: }\\Sigma F_y &amp;=0\\\\A_y+D_y - C_y&nbsp;&amp;= 0\\\\C_y &amp;= \\var{d(A_y)} +\\var{d(D_y)}\\\\&amp;=\\var{d(C_y) }\\downarrow\\\\\\\\\\text{III: }\\Sigma F_x &amp;= 0\\\\D_x-C_x&amp;=0\\\\C_x &amp;= \\var{d(C_x)} \\rightarrow \\end{align}\\]</p>\n<p><strong>FBD II</strong></p>\n<p>With $D_x$ determined, the remainin unknown on FBD II can be found.</p>\n<p>\\[\\begin{align} \\text{II: }\\Sigma F_x &amp;= 0\\\\D_x-E_x&amp;=0\\\\E_x &amp;= \\var{d(E_x)} \\leftarrow \\end{align}\\]</p>\n<p><strong>FBD IV</strong></p>\n<p>All the unknown reactions are now known. &nbsp;You can check your work by verifying that these force put member $CEB$ in equilibrium. &nbsp;If they don't, you have made an error.</p>\n<p>\\[\\begin{align} \\text{IV: }\\Sigma M_c &amp;\\stackrel{?}{=} 0\\\\&nbsp;0&amp;\\stackrel{?}{=} &nbsp;B_y&nbsp;(\\var{qty(b/2,unit_d)}) - E_y (\\var{DE/2}) + E_x (\\var{qty(h-d,unit_d)})\\\\0&amp;\\stackrel{?}{=} &nbsp; (\\var{d(B_y)})(\\var{qty(b/2,unit_d)}) - &nbsp;(\\var{d(E_y)})(\\var{DE/2}) &nbsp;+ (\\var{d(E_x)}) (\\var{qty(h-d,unit_d)})\\\\0&amp;\\stackrel{?}{=} &nbsp; (\\var{d(B_y qty(b/2,unit_d))}) - &nbsp;(\\var{d(E_y DE/2)}) &nbsp;+ (\\var{d(E_x qty(h-d,unit_d))})\\\\0&amp;= \\var{check} \\, \\checkmark\\end{align}\\]</p>\n<p></p>\n<p></p>", "rulesets": {}, "extensions": ["geogebra", "quantities"], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"c_y": {"name": "c_y", "group": "solution", "definition": "A_y + D_y", "description": "<p>member ADC sigma fy</p>", "templateType": "anything", "can_override": false}, "d_y": {"name": "d_y", "group": "solution", "definition": "f EF/DE", "description": "<p>member DEF sum of moment about E</p>", "templateType": "anything", "can_override": false}, "e_y": {"name": "e_y", "group": "solution", "definition": "d_y + f", "description": "<p>member DEF sum forces in y</p>", "templateType": "anything", "can_override": false}, "debug": {"name": "debug", "group": "Ungrouped variables", "definition": "false", "description": "", "templateType": "anything", "can_override": false}, "d": {"name": "d", "group": "Ungrouped variables", "definition": "random(1..3) h/4", "description": "<p>height of arm</p>", "templateType": "anything", "can_override": false}, "check": {"name": "check", "group": "solution", "definition": "scalar(b_y qty(b/2,unit_d) - e_y DE/2 + e_x qty(h-d,unit_d))", "description": "<p>member CEB sigma Mc&nbsp;</p>", "templateType": "anything", "can_override": false}, "b": {"name": "b", "group": "Ungrouped variables", "definition": "random(h/2..h#2)", "description": "<p>width between supports</p>", "templateType": "anything", "can_override": false}, "var_1": {"name": "var_1", "group": "Ungrouped variables", "definition": "[['b',b],['h',h],['L',L],\n['d',d],['unit','\"'+unit_d+'\"']]", "description": "<p>list passed to geogebra applet</p>", "templateType": "anything", "can_override": false}, "EF": {"name": "EF", "group": "solution", "definition": "qty(l,unit_d)-DE", "description": "", "templateType": "anything", "can_override": false}, "unit_F": {"name": "unit_F", "group": "Ungrouped variables", "definition": "if(unit_d='m','N','lb')", "description": "", "templateType": "anything", "can_override": false}, "F": {"name": "F", "group": "Ungrouped variables", "definition": "qty(random(10..500#50),unit_f)", "description": "", "templateType": "anything", "can_override": false}, "a_x": {"name": "a_x", "group": "solution", "definition": "qty(0,unit_f)", "description": "", "templateType": "anything", "can_override": false}, "DE": {"name": "DE", "group": "solution", "definition": "qty((h-d)/h b,unit_d)", "description": "", "templateType": "anything", "can_override": false}, "L": {"name": "L", "group": "Ungrouped variables", "definition": "random(scalar(DE)..b#0.5)", "description": "<p>length of horizontal arm</p>", "templateType": "anything", "can_override": false}, "d_x": {"name": "d_x", "group": "solution", "definition": "-(qty(b/2,unit_d) A_y + DE/2 D_y)/qty(h-d,unit_d)", "description": "<p>-\\dfrac{\\var{b/2} (\\var{A_y}) +&nbsp;\\var{scalar(DE/2)}(\\var{D_y})}{\\var{h-d}}&nbsp;</p>\n<p></p>\n<p>&nbsp;ex&nbsp;(b_y qty(b/2,unit_d) + e_y DE/2)/qty(h-d,unit_d)</p>\n<p>-(qty(b/2,unit_d) A_y + DE/2 D_y)/qty(h-d,unit_d)</p>", "templateType": "anything", "can_override": false}, "var_2": {"name": "var_2", "group": "Ungrouped variables", "definition": "[['b',b 6/h],['h',6],['L',L 6/h],\n['d',d 6/h],['unit','\"'+unit_d+'\"']]", "description": "<p>diagram scaled proportional to h =6</p>", "templateType": "anything", "can_override": false}, "unit_d": {"name": "unit_d", "group": "Ungrouped variables", "definition": "random('m','ft')", "description": "", "templateType": "anything", "can_override": false}, "h": {"name": "h", "group": "Ungrouped variables", "definition": "random(4..10#2)", "description": "<p>height of structure</p>", "templateType": "anything", "can_override": false}, "c_x": {"name": "c_x", "group": "solution", "definition": "d_x", "description": "", "templateType": "anything", "can_override": false}, "AD": {"name": "AD", "group": "solution", "definition": "qty(b/2 d/h,unit_d)", "description": "<p>sum of moment about A for entire frame</p>", "templateType": "anything", "can_override": false}, "e_x": {"name": "e_x", "group": "solution", "definition": "d_x", "description": "<p>fbd DEF sigma fx</p>", "templateType": "anything", "can_override": false}, "a_y": {"name": "a_y", "group": "solution", "definition": "f - b_y", "description": "<p>sum f_y for whole thing</p>", "templateType": "anything", "can_override": false}, "b_y": {"name": "b_y", "group": "solution", "definition": "f * (scalar(AD)+L)/b", "description": "<p>distance from a to force</p>", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "scalar(AD)+L <> b and scalar(DE) <> L // prevent force over roller or pin E", "maxRuns": 100}, "ungrouped_variables": ["b", "h", "L", "d", "unit_d", "var_1", "F", "unit_F", "var_2", "debug"], "variable_groups": [{"name": "solution", "variables": ["AD", "DE", "EF", "b_y", "a_y", "d_y", "d_x", "e_y", "e_x", "c_y", "c_x", "check", "a_x"]}], "functions": {"d": {"parameters": [["q", "quantity"]], "type": "string", "language": "jme", "definition": "siground(q,4)"}}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p style=\"padding-left: 40px;\">$\\mathbf{A}_x =$ [[0]]&nbsp;[[1]]&nbsp;<span data-jme-visible=\"debug\"> {a_x}&nbsp;</span></p>\n<p style=\"padding-left: 40px;\">$\\mathbf{A}_y =$ [[2]]&nbsp;[[3]]&nbsp;&nbsp;<span data-jme-visible=\"debug\">{d(a_y)}</span></p>\n<p style=\"padding-left: 40px;\">$\\mathbf{D}_x =$ [[4]]&nbsp;[[5]]&nbsp;<span data-jme-visible=\"debug\"> {d(d_x)}</span></p>\n<p style=\"padding-left: 40px;\">$\\mathbf{D}_y =$ [[6]]&nbsp;[[7]]&nbsp;<span data-jme-visible=\"debug\"> {d(d_y)}</span></p>\n<p style=\"padding-left: 40px;\">$\\mathbf{C}_x =$ [[8]]&nbsp;[[9]]&nbsp;<span data-jme-visible=\"debug\"> {d(C_x)}</span></p>\n<p style=\"padding-left: 40px;\">$\\mathbf{C}_y =$ [[10]]&nbsp;[[11]]&nbsp;<span data-jme-visible=\"debug\"> {d(c_y)}</span></p>", "gaps": [{"type": "engineering-answer", "useCustomName": true, "customName": "Ax", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": false, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "qty(0,unit_f)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": true, "customName": "L-R a", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "choices": ["&larr;", "&rarr;", "Neither"], "matrix": [0, 0, "1"], "distractors": ["", "", ""]}, {"type": "engineering-answer", "useCustomName": true, "customName": "Ay", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": false, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "abs(A_y)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": true, "customName": "U-D a", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "choices": ["&uarr;", "&darr;", "Neither"], "matrix": ["if(scalar(A_y)>0,1,0)", "if(scalar(A_y)<0,1,0)", "if(scalar(A_y)=0,1,0)"], "distractors": ["", "", ""]}, {"type": "engineering-answer", "useCustomName": true, "customName": "Dx", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": false, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "abs(D_x)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": true, "customName": "L-R d", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "choices": ["&larr;", "&rarr;", "Neither"], "matrix": ["if(scalar(D_x)>0,1,0)", "if(scalar(D_x)<0,1,0)", "if(scalar(D_x)=0,1,0)"], "distractors": ["", "", ""]}, {"type": "engineering-answer", "useCustomName": true, "customName": "Dy", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": false, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "abs(D_y)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": true, "customName": "U-D d", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "choices": ["&uarr;", "&darr;", "Neither"], "matrix": ["if(scalar(D_y)>0,1,0)", "if(scalar(D_y)<0,1,0)", "if(scalar(D_y)=0,1,0)"], "distractors": ["", "", ""]}, {"type": "engineering-answer", "useCustomName": true, "customName": "Cx", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": false, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "abs(c_x)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": true, "customName": "L-R c", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "choices": ["&larr;", "&rarr;", "Neither"], "matrix": ["if(scalar(C_x)>0,1,0)", "if(scalar(C_x)<0,1,0)", "if(scalar(C_x)=0,1,0)"], "distractors": ["", "", ""]}, {"type": "engineering-answer", "useCustomName": true, "customName": "Cy", "marks": "4", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": false, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "abs(C_y)", "right": "0.2", "close": "1.0", "C1": "75", "C2": "50", "C3": "25"}}, {"type": "1_n_2", "useCustomName": true, "customName": "U-D  c", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "dropdownlist", "displayColumns": 0, "showBlankOption": true, "showCellAnswerState": true, "choices": ["&uarr;", "&darr;", "Neither"], "matrix": ["if(scalar(C_y)>0,1,0)", "if(scalar(C_y)<0,1,0)", "if(scalar(C_y)=0,1,0)"], "distractors": ["", "", ""]}], "sortAnswers": false}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}], "resources": []}]}], "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}]}